<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Additive and multiplicative inverses</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find the additive and the multiplicative inverse of \(z = {@a@}+{@b@}i\).</p>
<p>a) \(w=\)  [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>
<p>b) \( z^{-1}=\) [[input:ans2]] [[validation:ans2]][[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked solution</h3>
<p>Given \(z = {@a@} + {@b@}i\), let's find its additive and multiplicative inverses:</p>
<p>a) Additive Inverse: </p><p>The additive inverse of \(z\) is simply \(-z\).</p><p> So, \(w=-z = -({@a@} + {@b@}i) = {@ta1@}\)</p>
<p>b) Multiplicative Inverse: </p><p>The multiplicative inverse of \(z\) is denoted as \(z^{-1}\) and is calculated by taking the reciprocal of the complex number.</p>
<p> In other words, \( \displaystyle z^{-1} = \displaystyle \frac{1}{z}\)</p>
<p>\( \qquad \displaystyle z^{-1} =  \displaystyle \frac{1}{z}=  \displaystyle {@z_bar@}\)</p>
<p>To rationalize the denominator, we'll multiply both the numerator and denominator by the conjugate of the denominator, which is \({@z_conj@}\):</p>
<p>\( \qquad   \displaystyle  z^{-1} = \displaystyle \left( \frac{1}{{@a@} + {@b@}i} \right) \left( \frac{{@z_conj@}}{{@z_conj@}} \right)=  \displaystyle \frac{{@num@}}{{@den@}} =  \displaystyle{@realpart(ta2)@}{@if imagpart(ta2)<0 then "-" else "+"@}{@abs(imagpart(ta2))*%i@}\)</p>
<p>So, the multiplicative inverse of \(  \displaystyle  z = {@a@} + {@b@}i \) is \(\displaystyle{@realpart(ta2)@}{@if imagpart(ta2)<0 then "-" else "+"@}{@abs(imagpart(ta2))*%i@}\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>a: rand(50)+1; 
b: rand(20)+1;     
z: a+b*%i;
z_bar:1/z;
z_conj:conjugate(z);
ta1:-1*z;
num:1*z_conj;
den:expand(z*z_conj);
ta2:num/den;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Find the additive and the multiplicative inverse of \(z = {@a@}+{@b@}i\).</p>
<p>\( \qquad -z= {@ta1@}\) </p>
<p>\( \qquad z^{-1}= {@ta2@}\) </p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1520685570</deployedseed>
    <deployedseed>982068185</deployedseed>
    <deployedseed>739390996</deployedseed>
    <deployedseed>839620447</deployedseed>
    <deployedseed>1350194223</deployedseed>
    <deployedseed>427718574</deployedseed>
    <deployedseed>1487651010</deployedseed>
    <deployedseed>844587567</deployedseed>
    <deployedseed>568672818</deployedseed>
    <deployedseed>2023804019</deployedseed>
    <deployedseed>1324317960</deployedseed>
    <deployedseed>751038094</deployedseed>
    <deployedseed>361706826</deployedseed>
    <deployedseed>918805962</deployedseed>
    <deployedseed>1447995348</deployedseed>
    <deployedseed>1601370359</deployedseed>
    <deployedseed>1577141284</deployedseed>
    <deployedseed>2118728058</deployedseed>
    <deployedseed>175310094</deployedseed>
    <deployedseed>581320244</deployedseed>
    <deployedseed>470970940</deployedseed>
    <deployedseed>1080259915</deployedseed>
    <deployedseed>18962543</deployedseed>
    <deployedseed>18400469</deployedseed>
    <deployedseed>127695472</deployedseed>
    <deployedseed>1147881798</deployedseed>
    <deployedseed>1664055885</deployedseed>
    <deployedseed>206768502</deployedseed>
    <deployedseed>770129203</deployedseed>
    <deployedseed>482016847</deployedseed>
    <deployedseed>115750920</deployedseed>
    <deployedseed>825989673</deployedseed>
    <deployedseed>61184306</deployedseed>
    <deployedseed>818468744</deployedseed>
    <deployedseed>1943278812</deployedseed>
    <deployedseed>744657441</deployedseed>
    <deployedseed>1189075861</deployedseed>
    <deployedseed>1726886586</deployedseed>
    <deployedseed>989042259</deployedseed>
    <deployedseed>588444867</deployedseed>
    <deployedseed>319421510</deployedseed>
    <deployedseed>1645174890</deployedseed>
    <deployedseed>178621438</deployedseed>
    <deployedseed>728760711</deployedseed>
    <deployedseed>2076075719</deployedseed>
    <deployedseed>362391512</deployedseed>
    <deployedseed>179220519</deployedseed>
    <deployedseed>1384992714</deployedseed>
    <deployedseed>638168262</deployedseed>
    <deployedseed>1794470451</deployedseed>
    <deployedseed>1124629709</deployedseed>
    <deployedseed>1055441983</deployedseed>
    <deployedseed>248005490</deployedseed>
    <deployedseed>2006963532</deployedseed>
    <deployedseed>862237309</deployedseed>
    <deployedseed>363601875</deployedseed>
    <deployedseed>1051802777</deployedseed>
    <deployedseed>1927408964</deployedseed>
    <deployedseed>224700485</deployedseed>
    <deployedseed>343074758</deployedseed>
    <deployedseed>1906164117</deployedseed>
    <deployedseed>1312206549</deployedseed>
    <deployedseed>853512991</deployedseed>
    <deployedseed>813343158</deployedseed>
    <deployedseed>1876124198</deployedseed>
    <deployedseed>1778592519</deployedseed>
    <deployedseed>357564172</deployedseed>
    <deployedseed>437375792</deployedseed>
    <deployedseed>1326515803</deployedseed>
    <deployedseed>539350151</deployedseed>
    <deployedseed>916864771</deployedseed>
    <deployedseed>1225266102</deployedseed>
    <deployedseed>960216516</deployedseed>
    <deployedseed>1012039125</deployedseed>
    <deployedseed>888444365</deployedseed>
    <deployedseed>2046830960</deployedseed>
    <deployedseed>162308207</deployedseed>
    <deployedseed>1646294166</deployedseed>
    <deployedseed>646392658</deployedseed>
    <deployedseed>1579430422</deployedseed>
    <deployedseed>491121044</deployedseed>
    <deployedseed>1686279221</deployedseed>
    <deployedseed>709083884</deployedseed>
    <deployedseed>2013703161</deployedseed>
    <deployedseed>1332869104</deployedseed>
    <deployedseed>35045002</deployedseed>
    <deployedseed>1911161224</deployedseed>
    <deployedseed>1106150802</deployedseed>
    <deployedseed>1334925010</deployedseed>
    <deployedseed>284516597</deployedseed>
    <deployedseed>1855433630</deployedseed>
    <deployedseed>2135801492</deployedseed>
    <deployedseed>1355742560</deployedseed>
    <deployedseed>1999599228</deployedseed>
    <deployedseed>129177504</deployedseed>
    <deployedseed>2116077838</deployedseed>
    <deployedseed>729139398</deployedseed>
    <deployedseed>636340030</deployedseed>
    <deployedseed>307248293</deployedseed>
    <deployedseed>1717522078</deployedseed>
    <deployedseed>1667634579</deployedseed>
    <deployedseed>212065442</deployedseed>
    <deployedseed>1741855823</deployedseed>
    <deployedseed>862837412</deployedseed>
    <deployedseed>694731122</deployedseed>
    <deployedseed>1607850557</deployedseed>
    <deployedseed>1835165757</deployedseed>
    <deployedseed>859161874</deployedseed>
    <deployedseed>446217188</deployedseed>
    <deployedseed>1866190419</deployedseed>
    <deployedseed>1021131716</deployedseed>
    <deployedseed>1742013427</deployedseed>
    <deployedseed>2030639111</deployedseed>
    <deployedseed>609344508</deployedseed>
    <deployedseed>809935943</deployedseed>
    <deployedseed>203015322</deployedseed>
    <deployedseed>1394427862</deployedseed>
    <deployedseed>1225519230</deployedseed>
    <deployedseed>34624215</deployedseed>
    <deployedseed>135917261</deployedseed>
    <deployedseed>1134165059</deployedseed>
    <deployedseed>1388924440</deployedseed>
    <deployedseed>879661944</deployedseed>
    <deployedseed>963316736</deployedseed>
    <deployedseed>1000293646</deployedseed>
    <deployedseed>1113678239</deployedseed>
    <deployedseed>279334422</deployedseed>
    <deployedseed>1865096829</deployedseed>
    <deployedseed>2027856703</deployedseed>
    <deployedseed>421042881</deployedseed>
    <deployedseed>986162428</deployedseed>
    <deployedseed>260328087</deployedseed>
    <deployedseed>909858425</deployedseed>
    <deployedseed>724340712</deployedseed>
    <deployedseed>856589878</deployedseed>
    <deployedseed>613731143</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
