<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Complex Number Operations in Trigonometric Form</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Given complex numbers:</p> 
<p>&nbsp; &nbsp; &nbsp; &nbsp; &nbsp;\( z_1= \displaystyle  {@aa@} (\cos({@angle@})+i\sin({@angle@}))\)</p>
<p>&nbsp; &nbsp; &nbsp; &nbsp; &nbsp;\( z_2= \displaystyle  {@bb@}(\cos({@angle2@})+i\sin({@angle2@})\))</p>
<p>Calculate in the form \(a+{@%i*b@} \) with \(a, b\) be real numbers;</p><p>&nbsp; &nbsp; &nbsp; a. \( \displaystyle  z_1&nbsp;\times z_2=\) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>
<p>&nbsp; &nbsp; &nbsp; b. \( \displaystyle \frac{ z_1 }{z_2} = \) [[input:ans2]] [[validation:ans2]][[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked solution</h3>
<p>Given complex numbers:</p>
<p>\[ z_1= {@aa@} (\cos({@angle@})+i\sin({@angle@}))\] </p>
<p> \[z_2= {@bb@}(\cos({@angle2@})+i\sin({@angle2@}))\]
<p> (a) Recall : \(A( \cos(\theta)+i\sin(\theta) ) \times B (\cos(\alpha)+i\sin(\alpha)) = (A \times B)(\cos(\theta+\alpha)+i\sin(\theta+\alpha))\)</p>
<p>In this case, the <em>real</em> parts are \( A={@aa@}\) and \(B={@bb@}\), so you need to multiply them. </p>
<p>Then add the angles which are {@angle@} and {@angle2@},</p>
<p>\( \begin{align}
\displaystyle z_{1} \times z_{2} & =  {@aa@}(\cos({@angle@})+i\sin({@angle@})) \times {@bb@} (\cos({@angle2@})+i\sin({@angle2@}) \\
 & = \displaystyle ({@aa@} \times {@bb@})((\cos({@angle@}+{@angle2@})+i\sin({@angle@}+{@angle2@})) \\
& = \displaystyle {@ss@} (\cos({@angle_sum@})+i\sin({@angle_sum@})) \\
& = \displaystyle {@ta@}
\end{align} \)</p>
<p> (b) Recall : For \( \displaystyle \frac{ z_1 }{z_2}=\frac{A(\cos(\theta)+i\sin(\theta))}{B(\cos(\alpha)+i\sin(\alpha))} = \frac{A}{B} (\cos(\theta-\alpha)+i\sin(\theta-\alpha)\)))</p>
<p>Therefore:</p>
<p>\( \begin{align}
\displaystyle \frac{ z_1 }{z_2} & = \frac{{@aa@}(\cos({@angle@})+i\sin({@angle@}))}{ {@bb@} (\cos({@angle2@})+i\sin({@angle2@})} \\
& =  \displaystyle  \left(  \frac{@aa@}{@bb@}\right) \left((\cos({@angle@}-{@angle2@})+i\sin({@angle@}-{@angle2@}) \right) \\
& =  \displaystyle {@ta2@}
\end{align} \)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*complex number 1*/
aa:3+rand(4);
angle: rand([%pi,%pi/3, %pi/4, %pi/6]);
zz: aa*(cos(angle)+%i*sin(angle));
z_real:realpart(zz);
z_imag:imagpart(zz);

/*complex number 2*/
bb:3+rand(3);
angle2: rand([%pi, %pi/2, %pi/3, %pi/4, %pi/6,%pi*5]);
z2: bb*(cos(angle2)+%i*sin(angle2));
z_real2:realpart(z2);
z_imag2:imagpart(z2);

/*Worked out solution 1*/
ss:aa*bb;
angle_sum:angle+angle2;
ta:ratsimp(ss*(cos(angle_sum)+%i*sin(angle_sum)));

/*Worked out solution 2*/
AB:aa/bb;
angle_diff:angle-angle2;
ta2:ratsimp(AB*(cos(angle_diff)+%i*sin(angle_diff)));</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Given complex numbers:</p> 
<p>&nbsp; &nbsp; &nbsp; &nbsp; &nbsp;\( z_1=\)&nbsp; \({@aa@} (\cos({@angle@})+i\sin({@angle@}))\)</p>
<p>&nbsp; &nbsp; &nbsp; &nbsp; &nbsp;\( z_2=\)&nbsp;{@bb@}(\(\cos({@angle2@})+i\sin({@angle2@})\))</p>
<p> a. \( z_1&nbsp;\times z_2={@ta@}\) </p>
<p> b. \( \displaystyle \frac{ z_1 }{z_2} ={@ta2@} \) </p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>16</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>18</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1520685570</deployedseed>
    <deployedseed>982068185</deployedseed>
    <deployedseed>739390996</deployedseed>
    <deployedseed>839620447</deployedseed>
    <deployedseed>1350194223</deployedseed>
    <deployedseed>427718574</deployedseed>
    <deployedseed>1487651010</deployedseed>
    <deployedseed>844587567</deployedseed>
    <deployedseed>568672818</deployedseed>
    <deployedseed>2023804019</deployedseed>
    <deployedseed>1324317960</deployedseed>
    <deployedseed>751038094</deployedseed>
    <deployedseed>361706826</deployedseed>
    <deployedseed>918805962</deployedseed>
    <deployedseed>1447995348</deployedseed>
    <deployedseed>1601370359</deployedseed>
    <deployedseed>1577141284</deployedseed>
    <deployedseed>2118728058</deployedseed>
    <deployedseed>175310094</deployedseed>
    <deployedseed>581320244</deployedseed>
    <deployedseed>470970940</deployedseed>
    <deployedseed>1080259915</deployedseed>
    <deployedseed>18962543</deployedseed>
    <deployedseed>18400469</deployedseed>
    <deployedseed>127695472</deployedseed>
    <deployedseed>1147881798</deployedseed>
    <deployedseed>1664055885</deployedseed>
    <deployedseed>206768502</deployedseed>
    <deployedseed>770129203</deployedseed>
    <deployedseed>482016847</deployedseed>
    <deployedseed>115750920</deployedseed>
    <deployedseed>825989673</deployedseed>
    <deployedseed>61184306</deployedseed>
    <deployedseed>818468744</deployedseed>
    <deployedseed>1943278812</deployedseed>
    <deployedseed>744657441</deployedseed>
    <deployedseed>1189075861</deployedseed>
    <deployedseed>1726886586</deployedseed>
    <deployedseed>989042259</deployedseed>
    <deployedseed>588444867</deployedseed>
    <deployedseed>319421510</deployedseed>
    <deployedseed>1645174890</deployedseed>
    <deployedseed>178621438</deployedseed>
    <deployedseed>728760711</deployedseed>
    <deployedseed>2076075719</deployedseed>
    <deployedseed>362391512</deployedseed>
    <deployedseed>179220519</deployedseed>
    <deployedseed>1384992714</deployedseed>
    <deployedseed>638168262</deployedseed>
    <deployedseed>1794470451</deployedseed>
    <deployedseed>1124629709</deployedseed>
    <deployedseed>1055441983</deployedseed>
    <deployedseed>248005490</deployedseed>
    <deployedseed>2006963532</deployedseed>
    <deployedseed>862237309</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
