<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the values of x from a matrix equal to zero</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Solve the equation \({@mat@} = 0\)</p>

<br>

<p>\(x = \) [[input:ans1]][[validation:ans1]] or \(x = \) [[input:ans2]] [[validation:ans2]]</p>
<p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution:</h3>
<p>As all of the entries in the first column are \(1\), it is simplest to calculate the determinant along the first column,</p>
<p style="margin-left: 30px">\( \begin{align}
{@mat@} & = 1 \cdot {@matrix([-2,4],[pp*x,qq*x^2])@} - 1 \cdot {@matrix([3,6],[pp*x,qq*x^2])@} + 1 \cdot {@matrix([3,6],[-2,4])@}\\
& = ({@-2*qq*x^2 - 4*pp*x@}) - ({@3*qq*x^2-6*pp*x@}) + (12 + 12) \\
& = {@fx2@} \\
& = {@fx1@}
\end{align}\)</p>

<br>

<p>Now, solving the given equation requires only that we solve a quadratic,</p>
<p style="margin-left: 30px">\(\begin{align*}
{@mat@} = 0 & \Rightarrow {@fx1@} = 0 \\
& \Rightarrow {@aa*x-bb@} = 0, \hspace{10px} {@cc*x-dd@} = 0
\end{align*}\)</p>

<br>

<p>Therefore, \( x = {@ta1@}\) or  \(x = {@ta2@} \)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>aa:rand([-5,-10]);
bb:rand([2,4,6,12,24]);
cc:rand(5)+2;
dd:24/bb;
qq:-aa*cc/5;
pp:-(aa*dd+bb*cc)/2;
mat:matrix([1,3,6],[1,-2,4],[1,pp*x,qq*x^2]);
ta1:bb/aa;
ta2:dd/cc;
fx1:(aa*x-bb)*(cc*x-dd);
fx2:expand(fx1);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>equation:\({@mat@} = 0\)</p>
<p>values of \(x\) are \(x={@ta1@}\) or \(x={@ta2@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>|</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[setta:{ta1,ta2};
setans:{ans1,ans2};
incansset:setdifference(setans,setta);
numans:cardinality(incansset);
incans:if numans=1 then listify(incansset)[1] else "";]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>Sets</answertest>
        <sans>setans</sans>
        <tans>setta</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ans2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You entered the same answer twice. This equation has two distinct solutions</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>numans</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer \(x={@listify(intersect(setta,setans))[1]@}\) is correct.
<br>However, \(x={@incans@}\) is not a solution to the equation given.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1418596212</deployedseed>
    <deployedseed>879237589</deployedseed>
    <deployedseed>219561990</deployedseed>
    <deployedseed>1119278566</deployedseed>
    <deployedseed>1222963785</deployedseed>
    <deployedseed>711869230</deployedseed>
    <deployedseed>491488254</deployedseed>
    <deployedseed>1090720675</deployedseed>
    <deployedseed>715003351</deployedseed>
    <deployedseed>1250267619</deployedseed>
    <deployedseed>752796057</deployedseed>
    <deployedseed>499996462</deployedseed>
    <deployedseed>855670333</deployedseed>
    <deployedseed>1654102671</deployedseed>
    <deployedseed>1816157451</deployedseed>
    <deployedseed>1068395125</deployedseed>
    <deployedseed>26866820</deployedseed>
    <deployedseed>1149927642</deployedseed>
    <deployedseed>948826918</deployedseed>
    <deployedseed>1184320530</deployedseed>
    <deployedseed>569273558</deployedseed>
    <deployedseed>490989880</deployedseed>
    <deployedseed>421282719</deployedseed>
    <deployedseed>1177057857</deployedseed>
    <deployedseed>692167803</deployedseed>
    <deployedseed>1530937861</deployedseed>
    <deployedseed>2031760369</deployedseed>
    <deployedseed>99219575</deployedseed>
    <deployedseed>845562602</deployedseed>
    <deployedseed>1195382353</deployedseed>
    <deployedseed>1369596817</deployedseed>
    <deployedseed>1274465021</deployedseed>
    <deployedseed>671249746</deployedseed>
    <deployedseed>1535456574</deployedseed>
    <deployedseed>408451525</deployedseed>
    <deployedseed>1952345376</deployedseed>
    <deployedseed>382695332</deployedseed>
    <deployedseed>1501425409</deployedseed>
    <deployedseed>1651696203</deployedseed>
    <deployedseed>1168831945</deployedseed>
    <deployedseed>2030115230</deployedseed>
    <deployedseed>146342383</deployedseed>
    <deployedseed>1302133862</deployedseed>
    <deployedseed>556260490</deployedseed>
    <deployedseed>1522628984</deployedseed>
    <deployedseed>871471475</deployedseed>
    <deployedseed>218191215</deployedseed>
    <deployedseed>1788667790</deployedseed>
    <deployedseed>2000385973</deployedseed>
    <deployedseed>1335355047</deployedseed>
    <deployedseed>166403417</deployedseed>
    <deployedseed>857400918</deployedseed>
    <deployedseed>1442173623</deployedseed>
    <deployedseed>691636810</deployedseed>
    <deployedseed>1330646846</deployedseed>
    <deployedseed>1160843700</deployedseed>
    <deployedseed>1432428283</deployedseed>
    <deployedseed>151037131</deployedseed>
    <deployedseed>1140378935</deployedseed>
    <deployedseed>353496990</deployedseed>
    <deployedseed>1970877173</deployedseed>
    <deployedseed>380212226</deployedseed>
    <deployedseed>68254765</deployedseed>
    <deployedseed>2141884733</deployedseed>
    <deployedseed>339759739</deployedseed>
    <deployedseed>2072387749</deployedseed>
    <deployedseed>86376042</deployedseed>
    <deployedseed>229400489</deployedseed>
    <deployedseed>336032306</deployedseed>
    <deployedseed>775763098</deployedseed>
    <deployedseed>1975409475</deployedseed>
    <deployedseed>314454135</deployedseed>
    <deployedseed>887714353</deployedseed>
    <deployedseed>1504233549</deployedseed>
    <deployedseed>1402878613</deployedseed>
    <deployedseed>1909483105</deployedseed>
    <deployedseed>759439224</deployedseed>
    <deployedseed>1274392459</deployedseed>
    <deployedseed>1094383291</deployedseed>
    <deployedseed>682961672</deployedseed>
    <deployedseed>343433460</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
