<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Plotting complex numbers on an Argand diagram</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>By dragging the points from the origin, plot the following complex numbers on the Argand diagram below.</p>
<p>(a) \( {@point1x + point1y*i@} \)</p>
<p>(b) \( {@if point2_rad=1 then "" else if integerp(point2_rad) then point2_rad else sconcat("\\sqrt{", point2_rad^2, "}")@} (\cos({@angles[1]@}) + i \sin({@angles[1]@})) \)</p>
<p>(c) \( {@if point3_rad=1 then "" else if integerp(point3_rad) then point3_rad else sconcat("\\sqrt{", point3_rad^2," }")@} \exp({@angles[2]@} i) \)</p>

<br>

[[jsxgraph input-ref-ans1="sR1" input-ref-ans2="sR2" input-ref-ans3="sR3" height="400px" width="400px"]] 
JXG.Options.text.useMathJax = true; 
var board = JXG.JSXGraph.initBoard(divid, {
  axis: true, 
  showCopyright: false, 
  boundingbox: [-6,6,6,-6], 
  keepaspectratio: true,
  showNavigation: false,
  grid:true
}); 

var pointStyle = {
  name:'',
  snapToGrid: true,
  size: 5,
  showInfobox: false
};

var p1 = board.create('point', [0,0], pointStyle); 
var p2 = board.create('point', [0,0], pointStyle); 
var p3 = board.create('point', [0,0], pointStyle); 

stack_jxg.bind_point(sR1,p1);
stack_jxg.bind_point(sR2,p2);
stack_jxg.bind_point(sR3,p3);

board.unsuspendUpdate();
[[/jsxgraph]]

<p hidden>[[input:ans1]] [[validation:ans1]]</p>
<p hidden>[[input:ans2]] [[validation:ans2]]</p>
<p hidden>[[input:ans3]] [[validation:ans3]]</p>
<p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked solution</h3>

<p>In an Argand diagram, the position of the point along the \(x\)-axis represents the real part of the complex number and the position on the \(y\)-axis represents the imaginary part.</p>

<br>

<p>(a) For the first number, we can determine the coordinates of the point directly by extracting the real and imaginary parts of the complex number.
<br>The point then has coordinates \(({@point1x@}, {@point1y@})\)</p>

<br>

<p>(b) For the second number, we evaluate the trigonometric functions to write the complex number in a form which is easier to plot on the Argand diagram.
<p style="margin-left:30px">\(\begin{align*}
{@if point2_rad=1 then "" else if integerp(point2_rad) then point2_rad else sconcat("\\sqrt{", point2_rad^2, "}")@} (\cos({@angles[1]@}) + i \sin({@angles[1]@}))
&= {@ta2[1] + ta2[2]*i@}
\end{align*}\)</p>
<p>Then, point then has coordinates \(({@ta2[1]@}, {@ta2[2]@})\)</p>

<br>

<p>(c) For the third number, we use Euler’s formula to write the complex number in a form which is easier to plot on the Argand diagram.</p>
<p style="margin-left:30px">\(\begin{align*}
{@if point3_rad=1 then "" else if integerp(point3_rad) then point3_rad else sconcat("\\sqrt{", point3_rad^2," }")@} \exp({@angles[2]@} i)
&= {@if point3_rad=1 then "" else if integerp(point3_rad) then point3_rad else sconcat("\\sqrt{", point3_rad^2, "}")@} (\cos({@angles[2]@}) + i \sin({@angles[2]@}) \\
&= {@if point3_rad=1 then "" else if integerp(point3_rad) then point3_rad else sconcat("\\sqrt{", point3_rad^2, "}")@} (({@cos(angles[2]) + sin(angles[2])*i@} )) \\
&= {@ta3[1] + ta3[2]*i@}
\end{align*}\)</p>
<p>Then, point then has coordinates \(({@ta3[1]@}, {@ta3[2]@})\)</p>

<br>

<p>Plotting these three points, the Argand diagram should be the following.</p>

<br>

[[jsxgraph height="300px" width="300px"]] 
JXG.Options.text.useMathJax = true; 
var board = JXG.JSXGraph.initBoard(divid, {
  axis: true, 
  showCopyright: false, 
  boundingbox: [-6,6,6,-6], 
  keepaspectratio: true,
  showNavigation: false,
  grid:true
}); 

var pointStyle = {
  name:'',
  snapToGrid: true,
  size: 5,
  showInfobox: false
};

var p1 = board.create('point', [{#point1x#},{#point1y#}], {fixed:true}, pointStyle); 
var p2 = board.create('point', [{#ta2[1]#},{#ta2[2]#}], {fixed:true}, pointStyle); 
var p3 = board.create('point', [{#ta3[1]#},{#ta3[2]#}], {fixed:true}, pointStyle); 
[[/jsxgraph]]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/* for the first point, randomise two distinct non-zero coordinates (so point 1 cannot be the same as 2 or 3) */
point1x: rand_with_prohib(-5,5,[0]);
point1y: rand_with_prohib(-5,5,[0]);

/* for the second and third points, randomise the angle and then the radius */
axes_angles: [-pi/2, pi/2, pi, 3*pi/2, 2*pi];
diag_angles: [-pi/4, pi/4, 3*pi/4, 5*pi/4, 7*pi/4];

angles: random_permutation([rand(axes_angles), rand(diag_angles)]);

point2_rad: if member(angles[1], axes_angles) then rand(5)+1 else if member(angles[1], diag_angles) then sqrt(2*(rand(5)+1)^2);
point3_rad: if member(angles[2], axes_angles) then rand(5)+1 else if member(angles[2], diag_angles) then sqrt(2*(rand(5)+1)^2);

/* teacher answers */
ta1:[point1x, point1y];
ta2:[point2_rad*cos(angles[1]), point2_rad*sin(angles[1])];
ta3:[point3_rad*cos(angles[2]), point3_rad*sin(angles[2])];</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Complex numbers:</p>
<p>(a) \( {@point1x + point1y*i@} \)</p>
<p>(b) \( {@point2_rad@} (\cos({@angles[1]@}) + i \sin({@angles[1]@})) \)</p>
<p>(c) \( {@point3_rad@} \exp({@angles[2]@} i) \)</p>
<p>Correct coordinates: {@ta1@}, {@ta2@}, {@ta3@}</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/* checks on first answer */
node2_test: if member(ans1, [ta1, ta2, ta3]) then 1 else 0;

/* checks on second answer */
node3_test: if ans2=ans1 then 1 else 0;
node4_test: if member(ans2, [ta1, ta2, ta3]) then 1 else 0;

/* checks on third answer */
node5_test: if ans3=ans1 or ans3=ans2 then 1 else 0;
node6_test: if member(ans3, [ta1, ta2, ta3]) then 1 else 0;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>{ans1, ans2, ans3}</sans>
        <tans>{ta1, ta2, ta3}</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>node2_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1/3</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>The point you have placed at {@ans1@} is <b>correct</b>.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The point you have placed at {@ans1@} is <b>incorrect</b>.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>node3_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>4</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have repeated the point {@ans2@} in your set of answers.
<br>The numbers given are represented by three <b>distinct points</b>.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>node4_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1/3</truescore>
        <truepenalty/>
        <truenextnode>4</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>The point you have placed at {@ans2@} is <b>correct</b>.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The point you have placed at {@ans2@} is <b>incorrect</b>.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>node5_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have repeated the point {@ans3@} in your set of answers.
<br>The numbers given are represented by three <b>distinct points</b>.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>5</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>5</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>node6_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1/3</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-6-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>The point you have placed at {@ans3@} is <b>correct</b>.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-6-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The point you have placed at {@ans3@} is <b>incorrect</b>.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1684517805</deployedseed>
    <deployedseed>835493141</deployedseed>
    <deployedseed>318457387</deployedseed>
    <deployedseed>772770352</deployedseed>
    <deployedseed>2072926034</deployedseed>
    <deployedseed>778065303</deployedseed>
    <deployedseed>155515298</deployedseed>
    <deployedseed>788439091</deployedseed>
    <deployedseed>882219682</deployedseed>
    <deployedseed>762352595</deployedseed>
    <deployedseed>1170866324</deployedseed>
    <deployedseed>349758999</deployedseed>
    <deployedseed>1744591337</deployedseed>
    <deployedseed>699965610</deployedseed>
    <deployedseed>1527530248</deployedseed>
    <deployedseed>429739576</deployedseed>
    <deployedseed>907074585</deployedseed>
    <deployedseed>1798739501</deployedseed>
    <deployedseed>1500653128</deployedseed>
    <deployedseed>1945442589</deployedseed>
    <deployedseed>832710758</deployedseed>
    <deployedseed>1899874221</deployedseed>
    <deployedseed>690041041</deployedseed>
    <deployedseed>308437775</deployedseed>
    <deployedseed>617747197</deployedseed>
    <deployedseed>309162135</deployedseed>
    <deployedseed>108355615</deployedseed>
    <deployedseed>1864069340</deployedseed>
    <deployedseed>236150865</deployedseed>
    <deployedseed>1385584651</deployedseed>
    <deployedseed>462160590</deployedseed>
    <deployedseed>555541462</deployedseed>
    <deployedseed>1386706888</deployedseed>
    <deployedseed>1967544016</deployedseed>
    <deployedseed>366876905</deployedseed>
    <deployedseed>1692562561</deployedseed>
    <deployedseed>1422333812</deployedseed>
    <deployedseed>198876628</deployedseed>
    <deployedseed>1634085324</deployedseed>
    <deployedseed>1756385844</deployedseed>
    <deployedseed>1080694765</deployedseed>
    <deployedseed>1693993522</deployedseed>
    <deployedseed>526356028</deployedseed>
    <deployedseed>614924940</deployedseed>
    <deployedseed>1597298855</deployedseed>
    <deployedseed>2146252969</deployedseed>
    <deployedseed>1108797128</deployedseed>
    <deployedseed>1507105494</deployedseed>
    <deployedseed>206203053</deployedseed>
    <deployedseed>1992156886</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
