<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Derivative of the sine of a rational function</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find the derivative of \(f(t) = {@fx@}\)</p>

<p style="margin-left: 30px">\(f'(t) = \) [[input:ans1]] [[validation:ans1]]</p>
<p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>In this question, we compute the derivative using a combination of the chain rule and the quotient rule. 
<br>Let's first look at {@quot@} and calculate its derivative as we need it for the chain rule.</p>
<p>Using the quotient rule we get that the derivative is,</p>

<p style="margin-left: 30px">\(\displaystyle 
\frac{d}{dx} \left( {@quot@} \right)
= \frac{{@aa@} ({@denom@}) - ({@numer@})}{({@denom@})^2}
= \frac{{@expand(denom*aa-numer)@}}{({@denom@})^2}\)</p>

<p>We now apply the chain rule as follows,</p>
<p style="margin-left: 30px">\(\displaystyle \begin{align*}
f'(t) &= \frac{d}{dx} \left( {@quot@} \right) \cdot \cos \left( {@quot@} \right) \\
&= \frac{{@expand(denom*aa-numer)@}}{({@denom@})^2} \cdot \cos \left( {@quot@} \right) \\
&= \frac{{@expand(denom*aa-numer)*cos(quot)@}}{({@denom@})^2}
\end{align*}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>aa:rand(4)+2;
pm1:rand_with_prohib(-1,1,[0]);

bb:pm1*(rand_with_prohib(1,10,[aa]));
pm2:rand_with_prohib(-1,1,[0]);

cc:pm2*(rand_with_prohib(1,10,[abs(bb),bb/aa]));

numer:aa*t+bb;
denom:t+cc;
quot:numer/denom;
fx:sin(quot);

ta:ratsimp(diff(fx,t));</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\(f(t) = {@fx@}\).</p>
<p>\(f'(t) ={@ta@} \)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1032228019</deployedseed>
    <deployedseed>1538804822</deployedseed>
    <deployedseed>695203845</deployedseed>
    <deployedseed>1888180557</deployedseed>
    <deployedseed>364926935</deployedseed>
    <deployedseed>1155815308</deployedseed>
    <deployedseed>1220439563</deployedseed>
    <deployedseed>1910912693</deployedseed>
    <deployedseed>1352788013</deployedseed>
    <deployedseed>2040763417</deployedseed>
    <deployedseed>121840961</deployedseed>
    <deployedseed>1797035522</deployedseed>
    <deployedseed>655065767</deployedseed>
    <deployedseed>1205206384</deployedseed>
    <deployedseed>600133012</deployedseed>
    <deployedseed>1252792289</deployedseed>
    <deployedseed>518011942</deployedseed>
    <deployedseed>1109885735</deployedseed>
    <deployedseed>1372341203</deployedseed>
    <deployedseed>592604934</deployedseed>
    <deployedseed>279608973</deployedseed>
    <deployedseed>413404369</deployedseed>
    <deployedseed>1216574018</deployedseed>
    <deployedseed>1661324330</deployedseed>
    <deployedseed>523685404</deployedseed>
    <deployedseed>443413442</deployedseed>
    <deployedseed>1773648551</deployedseed>
    <deployedseed>1763168419</deployedseed>
    <deployedseed>1741202683</deployedseed>
    <deployedseed>1025959030</deployedseed>
    <deployedseed>988400646</deployedseed>
    <deployedseed>1074924698</deployedseed>
    <deployedseed>2014241089</deployedseed>
    <deployedseed>1587161898</deployedseed>
    <deployedseed>1074274468</deployedseed>
    <deployedseed>436837970</deployedseed>
    <deployedseed>133416033</deployedseed>
    <deployedseed>190404465</deployedseed>
    <deployedseed>844279417</deployedseed>
    <deployedseed>317565800</deployedseed>
    <deployedseed>331024900</deployedseed>
    <deployedseed>1974751459</deployedseed>
    <deployedseed>1308011244</deployedseed>
    <deployedseed>1721000536</deployedseed>
    <deployedseed>1874915674</deployedseed>
    <deployedseed>1225941967</deployedseed>
    <deployedseed>129410112</deployedseed>
    <deployedseed>584882212</deployedseed>
    <deployedseed>1896156097</deployedseed>
    <deployedseed>984737129</deployedseed>
    <deployedseed>1602092866</deployedseed>
    <deployedseed>609491649</deployedseed>
    <deployedseed>365039032</deployedseed>
    <deployedseed>395897328</deployedseed>
    <deployedseed>1235267456</deployedseed>
    <deployedseed>673749694</deployedseed>
    <deployedseed>298512136</deployedseed>
    <deployedseed>1345910253</deployedseed>
    <deployedseed>970072835</deployedseed>
    <deployedseed>1223458340</deployedseed>
    <deployedseed>502116810</deployedseed>
    <deployedseed>83353534</deployedseed>
    <deployedseed>2030915208</deployedseed>
    <deployedseed>150059844</deployedseed>
    <deployedseed>1885098181</deployedseed>
    <deployedseed>1648142989</deployedseed>
    <deployedseed>1802533149</deployedseed>
    <deployedseed>668902296</deployedseed>
    <deployedseed>1215838727</deployedseed>
    <deployedseed>555483436</deployedseed>
    <deployedseed>1145251007</deployedseed>
    <deployedseed>1505836863</deployedseed>
    <deployedseed>385758303</deployedseed>
    <deployedseed>1274307142</deployedseed>
    <deployedseed>412214496</deployedseed>
    <deployedseed>1887770995</deployedseed>
    <deployedseed>720687338</deployedseed>
    <deployedseed>1228972780</deployedseed>
    <deployedseed>1478386371</deployedseed>
    <deployedseed>864098020</deployedseed>
    <deployedseed>1922370355</deployedseed>
    <deployedseed>1787711870</deployedseed>
    <deployedseed>1986341338</deployedseed>
    <deployedseed>595596899</deployedseed>
    <deployedseed>1717721385</deployedseed>
    <deployedseed>737516248</deployedseed>
    <deployedseed>90101867</deployedseed>
    <deployedseed>1234041668</deployedseed>
    <deployedseed>93438365</deployedseed>
    <deployedseed>730285730</deployedseed>
    <deployedseed>1444831418</deployedseed>
    <deployedseed>574958508</deployedseed>
    <deployedseed>482144499</deployedseed>
    <deployedseed>1493917056</deployedseed>
    <deployedseed>19577419</deployedseed>
    <deployedseed>91942297</deployedseed>
    <deployedseed>938678873</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
