<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Derivative of a Function</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find the derivative of \(f(x)={@fx@}\)</p>
<p>\(f'(x) = \) [[input:ans1]] [[validation:ans1]]</p>
<p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>To find the derivative of the function \( f(x) = \left(x^4 + {@aa@}x^3 + {@bb@}x + {@dd@}\right)^{@ee@} \) we wish to apply the chain rule.</p>

<br>

<p>Firstly, we identify the inner function and determine its derivative.</p>
<p>We will write the inner function as \( g(x) = x^4 + {@aa@}x^3 + {@bb@}x + {@dd@} \). 
<br>To find its derivative, we apply the power rule and the addition rule.</p>

<p style="margin-left: 30px">\(\begin{align*} 
g'(x) &= \textstyle \frac{d}{dx} \left( {@inside@} \right) \\
&= \textstyle \frac{d}{dx} \left(x^4\right) + \frac{d}{dx} \left({@aa@}x^3\right) + \frac{d}{dx} \left({@bb@}x\right) + \frac{d}{dx} \left({@dd@}\right) \\
&= {@dinside@}
\end{align*}\)</p>

<br>

<p>We can now apply the chain rule.</p>
<p style="margin-left: 30px">\(\begin{align*} 
f'(x) &= \textstyle \frac{d}{dx} \left(g(x)\right)^{@ee@} \\
&= {@ee@} \left( g(x) \right)^{@(ee - 1)@} \cdot g'(x) \\
&= {@ee@} \left( {@inside@} \right)^{@(ee - 1)@} \left(4x^3 + {@ 3*aa@}x^2 + {@bb@}\right) 
\end{align*}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>pm1:rand_with_prohib(-1,1,[0]);
aa:pm1*(rand(8)+2);
pm2:rand_with_prohib(-1,1,[0]);
bb:pm2*(rand(8)+2);
pm3:rand_with_prohib(-1,1,[0]);
cc:pm3*(rand(8)+2);
pm4:rand_with_prohib(-1,1,[0]);
dd:pm4*(rand(8)+2);
ee:rand(5)+5;
inside:x^4+aa*x^3+bb*x+dd;
dinside:diff(inside,x);
fx:inside^ee;
ta:diff(fx,x);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\( f(x) = {@fx@} \)</p>
<p>\(f'(x) ={@ta@} \)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct student answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2119894739</deployedseed>
    <deployedseed>535779768</deployedseed>
    <deployedseed>445113592</deployedseed>
    <deployedseed>776019969</deployedseed>
    <deployedseed>1142278394</deployedseed>
    <deployedseed>1970667015</deployedseed>
    <deployedseed>26398487</deployedseed>
    <deployedseed>279485340</deployedseed>
    <deployedseed>1544122061</deployedseed>
    <deployedseed>681151660</deployedseed>
    <deployedseed>1155845662</deployedseed>
    <deployedseed>878599606</deployedseed>
    <deployedseed>827442266</deployedseed>
    <deployedseed>53750439</deployedseed>
    <deployedseed>1210234730</deployedseed>
    <deployedseed>1653271205</deployedseed>
    <deployedseed>305036145</deployedseed>
    <deployedseed>1884912091</deployedseed>
    <deployedseed>1189095941</deployedseed>
    <deployedseed>1023610185</deployedseed>
    <deployedseed>1536594058</deployedseed>
    <deployedseed>1863535367</deployedseed>
    <deployedseed>177959984</deployedseed>
    <deployedseed>467370068</deployedseed>
    <deployedseed>833202244</deployedseed>
    <deployedseed>816099325</deployedseed>
    <deployedseed>122604161</deployedseed>
    <deployedseed>414456050</deployedseed>
    <deployedseed>1376214250</deployedseed>
    <deployedseed>1646937501</deployedseed>
    <deployedseed>1261726922</deployedseed>
    <deployedseed>1344816168</deployedseed>
    <deployedseed>1878428287</deployedseed>
    <deployedseed>1031308235</deployedseed>
    <deployedseed>140171027</deployedseed>
    <deployedseed>319442979</deployedseed>
    <deployedseed>1210410585</deployedseed>
    <deployedseed>437049631</deployedseed>
    <deployedseed>115505964</deployedseed>
    <deployedseed>497712657</deployedseed>
    <deployedseed>1114516655</deployedseed>
    <deployedseed>1644858124</deployedseed>
    <deployedseed>895244432</deployedseed>
    <deployedseed>347736887</deployedseed>
    <deployedseed>614534242</deployedseed>
    <deployedseed>1918188287</deployedseed>
    <deployedseed>610646611</deployedseed>
    <deployedseed>148022323</deployedseed>
    <deployedseed>1099629689</deployedseed>
    <deployedseed>1126752784</deployedseed>
    <deployedseed>827247955</deployedseed>
    <deployedseed>94700872</deployedseed>
    <deployedseed>225529178</deployedseed>
    <deployedseed>1847000133</deployedseed>
    <deployedseed>124570689</deployedseed>
    <deployedseed>1218748828</deployedseed>
    <deployedseed>851365087</deployedseed>
    <deployedseed>2050586969</deployedseed>
    <deployedseed>834574143</deployedseed>
    <deployedseed>1787436631</deployedseed>
    <deployedseed>984570498</deployedseed>
    <deployedseed>954650350</deployedseed>
    <deployedseed>751492102</deployedseed>
    <deployedseed>1621541199</deployedseed>
    <deployedseed>805972641</deployedseed>
    <deployedseed>1462707788</deployedseed>
    <deployedseed>1671263347</deployedseed>
    <deployedseed>504818758</deployedseed>
    <deployedseed>882195197</deployedseed>
    <deployedseed>191446162</deployedseed>
    <deployedseed>440458242</deployedseed>
    <deployedseed>1007512196</deployedseed>
    <deployedseed>960862987</deployedseed>
    <deployedseed>1142542017</deployedseed>
    <deployedseed>2020520731</deployedseed>
    <deployedseed>2081732886</deployedseed>
    <deployedseed>2024105460</deployedseed>
    <deployedseed>119244286</deployedseed>
    <deployedseed>1095895958</deployedseed>
    <deployedseed>1060408357</deployedseed>
    <deployedseed>1363310916</deployedseed>
    <deployedseed>1992983458</deployedseed>
    <deployedseed>1903835854</deployedseed>
    <deployedseed>1908403294</deployedseed>
    <deployedseed>1515043841</deployedseed>
    <deployedseed>1240222554</deployedseed>
    <deployedseed>884415340</deployedseed>
    <deployedseed>1942688913</deployedseed>
    <deployedseed>500499663</deployedseed>
    <deployedseed>282539946</deployedseed>
    <deployedseed>1130169294</deployedseed>
    <deployedseed>1966503902</deployedseed>
    <deployedseed>196188666</deployedseed>
    <deployedseed>118470664</deployedseed>
    <deployedseed>1285432481</deployedseed>
    <deployedseed>2057711322</deployedseed>
    <deployedseed>1166948357</deployedseed>
    <deployedseed>2119887117</deployedseed>
    <deployedseed>2061111478</deployedseed>
    <deployedseed>699817330</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
