<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the Derivative of a Square Root Function at a Specific Point</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p></p><p>Find the derivative of \(f(x)=\) {@fx@} at \(x={@cc@}\) using limits.</p><p>[[input:ans1]] [[validation:ans1]]&nbsp;[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p><h3><b>Worked solution</b></h3></p>
<p>We find the derivative of \(f(x)=\) {@fx@} at \(x={@cc@}\) using limits.</p>
<p>We have;</p>
<p>\(
\begin{align}
f'(x) & = \displaystyle \lim_{h\rightarrow 0} \frac{f(x+h)-f(x)}{h} \quad (\text{By definition}) \\
& =  \displaystyle \lim_{h\rightarrow 0} \frac{{@aa@} \sqrt{{@bb@} (x+h)} -{{@aa@} \sqrt{{@bb@}x}}}{h}  \\
&= \displaystyle \lim_{h\rightarrow 0} \left( \frac{{@aa@}  \sqrt{{@bb@}   (x+h)} -{{@aa@}  \sqrt{{@bb@}  x}}}{h} \right)  \left( \frac{{@aa@}  \sqrt{{@bb@}   (x+h)} +{{@aa@}  \sqrt{{@bb@}  x}}}{{@aa@}  \sqrt{{@bb@}   (x+h)} +{{@aa@}  \sqrt{{@bb@}  x}}} \right)\\
&= \displaystyle \lim_{h\rightarrow 0} \frac{({@aa@}  \sqrt{{@bb@} (x+h)})^2-({@aa@}  \sqrt{{@bb@}  x})^2}{h  ({@aa@}  \sqrt{{@bb@}   (x+h)} +{{@aa@}  \sqrt{{@bb@}  x}})} \\
&=\displaystyle \lim_{h\rightarrow 0} \frac{({@aa^2@}) ( {@bb@}  (x+h))-({@aa^2@})({@bb@}x)}{h  ({@aa@}  \sqrt{{@bb@}   (x+h)} +{{@aa@}  \sqrt{{@bb@}  x}})} \\
&=\displaystyle \lim_{h\rightarrow 0} \frac{({@aa^2@})({@bb@}h)}{h  ({@aa@}  \sqrt{{@bb@}   (x+h)} +{{@aa@}  \sqrt{{@bb@}  x}})} \\
&= \displaystyle \lim_{h\rightarrow 0} \frac{({@aa^2@})({@bb@})}{{@aa@}  \sqrt{{@bb@}   (x+h)} +{ {@aa@}  \sqrt{{@bb@}  x}}} \\
&= \displaystyle \lim_{h\rightarrow 0} \frac{({@aa@})({@bb@})}{\sqrt{{@bb@}   (x+h)} +{  \sqrt{{@bb@}  x}}} \\
&= \displaystyle \lim_{h\rightarrow 0} \frac{{@ev(aa*bb,simp)@}}{\sqrt{{@bb@}   (x+h)} +{  \sqrt{{@bb@}  x}}} \\
& = \frac{{@ev(aa*bb,simp)@}}{2 \sqrt{{@bb@} x}} 
\end{align}
\)</p>
<p>Therefore, \(f'({@cc@}) = \frac{{@ev(aa*bb,simp)@}}{2  \sqrt{{@ev(bb*cc,simp)@}}} = {@ta@}\).</p><br><p></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>simp:false;
aa:ev(rand(9)+2,simp);
bb:ev((rand(4)+2)^2,simp);
cc:ev((rand(3)+1)^2,simp);
fx:aa*(sqrt(bb*x));
ta:ev(aa*bb/(2*sqrt(bb*cc)),simp);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>For \(f(x)={@fx@}, f'({@cc@})={@ta@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 2: Find a limit of the function at a given point*/
mistake1:at(fx,[x=cc]);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text>You determined the limit of the function using the formula \( \displaystyle \lim_{x \rightarrow {@cc@}} f(x) = f({@cc@}) \), which is incorrect. Check the worked solution!</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1300864414</deployedseed>
    <deployedseed>2060643544</deployedseed>
    <deployedseed>667617234</deployedseed>
    <deployedseed>944308148</deployedseed>
    <deployedseed>1149432654</deployedseed>
    <deployedseed>699133713</deployedseed>
    <deployedseed>1694776685</deployedseed>
    <deployedseed>24951117</deployedseed>
    <deployedseed>383748777</deployedseed>
    <deployedseed>472270856</deployedseed>
    <deployedseed>19007883</deployedseed>
    <deployedseed>496529704</deployedseed>
    <deployedseed>37527046</deployedseed>
    <deployedseed>1320693415</deployedseed>
    <deployedseed>2033603644</deployedseed>
    <deployedseed>1398858443</deployedseed>
    <deployedseed>223760223</deployedseed>
    <deployedseed>797189388</deployedseed>
    <deployedseed>1811939698</deployedseed>
    <deployedseed>894207429</deployedseed>
    <deployedseed>753481308</deployedseed>
    <deployedseed>2062482741</deployedseed>
    <deployedseed>1579675315</deployedseed>
    <deployedseed>1409560086</deployedseed>
    <deployedseed>488869930</deployedseed>
    <deployedseed>1501650478</deployedseed>
    <deployedseed>1691234168</deployedseed>
    <deployedseed>1271947741</deployedseed>
    <deployedseed>748340052</deployedseed>
    <deployedseed>1093430210</deployedseed>
    <deployedseed>198874610</deployedseed>
    <deployedseed>337086904</deployedseed>
    <deployedseed>150982426</deployedseed>
    <deployedseed>407163870</deployedseed>
    <deployedseed>643935489</deployedseed>
    <deployedseed>1875193465</deployedseed>
    <deployedseed>76376452</deployedseed>
    <deployedseed>1303372980</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
