<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Matrix Rank Determination</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Determine the rank of the matrix \( {\bf A} = {@r1@}\)</p>
<p style="margin-left: 30px">\(\text{rank}({\bf A})\) = [[input:ans1]] [[validation:ans1]]</p>
<p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked solution</h3>
<p>Given matrix \(A\),</p>
<p style="margin-left: 30px">\(A = {@r1@}\)</p>

<p>To find the rank of matrix \(A\), we can perform row reduction to convert the matrix to its row echelon form (REF) or reduced row echelon form (RREF).
<br>The rank of \(A\) is then the number of non-zero rows in either of these matrices.</p>

<p>Using row reduction, we obtain the reduced row echelon form of matrix \(A\) as being,</p>
<p style="margin-left: 30px">\( \text{rref}(A) = {@ech@} \)

<p>The rank of matrix \(A\) is equal to the number of non-zero rows in its row echelon form. In this case, there are {@correctanswer@} non-zero rows. 
<br>Hence, the rank of matrix \(A\) is {@correctanswer@}.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>r1:matrix([1,rand_with_step(-9,9,1),rand_with_step(-9,9,1)],[rand_with_step(-9,9,1),rand_with_step(-9,9,1),rand_with_step(-9,9,1)],[rand_with_step(-9,9,1),rand_with_step(-9,9,1),rand_with_step(-9,9,1)]);
ech:echelon(r1);
correctanswer:rank(r1);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>matrix \( {\bf A} = {@r1@}\)</p>
<p> rank\(={@correctanswer@}\) </p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>correctanswer</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>rank</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>correctanswer</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>918983769</deployedseed>
    <deployedseed>599605838</deployedseed>
    <deployedseed>117062141</deployedseed>
    <deployedseed>1453395443</deployedseed>
    <deployedseed>48906896</deployedseed>
    <deployedseed>1396405064</deployedseed>
    <deployedseed>926349852</deployedseed>
    <deployedseed>585470460</deployedseed>
    <deployedseed>1055620925</deployedseed>
    <deployedseed>1509806083</deployedseed>
    <deployedseed>43821266</deployedseed>
    <deployedseed>1073070437</deployedseed>
    <deployedseed>802597202</deployedseed>
    <deployedseed>649110905</deployedseed>
    <deployedseed>1718268075</deployedseed>
    <deployedseed>1658870748</deployedseed>
    <deployedseed>938074242</deployedseed>
    <deployedseed>568397174</deployedseed>
    <deployedseed>1980055658</deployedseed>
    <deployedseed>2035602027</deployedseed>
    <deployedseed>1861962756</deployedseed>
    <deployedseed>738615624</deployedseed>
    <deployedseed>2060273782</deployedseed>
    <deployedseed>715807693</deployedseed>
    <deployedseed>1838912215</deployedseed>
    <deployedseed>1269231951</deployedseed>
    <deployedseed>1772363977</deployedseed>
    <deployedseed>1294059867</deployedseed>
    <deployedseed>644039626</deployedseed>
    <deployedseed>1714662355</deployedseed>
    <deployedseed>1683528706</deployedseed>
    <deployedseed>364456476</deployedseed>
    <deployedseed>349299424</deployedseed>
    <deployedseed>1279516768</deployedseed>
    <deployedseed>840309441</deployedseed>
    <deployedseed>151861501</deployedseed>
    <deployedseed>1073722406</deployedseed>
    <deployedseed>1470158570</deployedseed>
    <deployedseed>1876488144</deployedseed>
    <deployedseed>1600439887</deployedseed>
    <deployedseed>1212791061</deployedseed>
    <deployedseed>503907473</deployedseed>
    <deployedseed>983464436</deployedseed>
    <deployedseed>276202177</deployedseed>
    <deployedseed>1781582033</deployedseed>
    <deployedseed>634737505</deployedseed>
    <deployedseed>797385793</deployedseed>
    <deployedseed>1925021350</deployedseed>
    <deployedseed>2029742286</deployedseed>
    <deployedseed>2125025269</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>correctanswer</value>
      </testinput>
    </qtest>
  </question>
</quiz>
