<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Truth tables for logical operators</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>(a) Complete the truth table below for {@operator[1]@}.</p>

<br>

<table style="border: 1px solid black; border-collapse: collapse; text-align: center; margin-left: 30px;">
  <tr>
    <th style="width:100px; border: 1px solid black">\(A\)</th>
    <th style="width:100px; border: 1px solid black">\(B\)</th>
    <th style="width:200px; border: 1px solid black">\(A {@operator[2]@} B\)</th>
  </tr>
  <tr>
    <td style="border: 1px solid black">true</td>
    <td style="border: 1px solid black">true</td>
    <td style="border: 1px solid black">[[input:ans1]]</td>
  </tr>
  <tr>
    <td style="border: 1px solid black">true</td>
    <td style="border: 1px solid black">false</td>
    <td style="border: 1px solid black">[[input:ans2]]</td>
  </tr>
  <tr>
    <td style="border: 1px solid black">false</td>
    <td style="border: 1px solid black">true</td>
    <td style="border: 1px solid black">[[input:ans3]]</td>
  </tr>
  <tr>
    <td style="border: 1px solid black">false</td>
    <td style="border: 1px solid black">false</td>
    <td style="border: 1px solid black">[[input:ans4]]</td>
  </tr>
</table>

<br>

<p>[[validation:ans1]] [[validation:ans2]] [[validation:ans3]] [[validation:ans4]]</p>
<p>[[feedback:prt1]]</p>

<p>(b) On a day where it <b>is raining</b>, I claim that "today {@statement_A[1]@} {@operator[3]@} {@statement_B[1]@}".
<br>Which of the following options describes the truth of this statement.</p>

<p>[[input:ans5]] [[validation:ans5]]</p>
<p>[[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked solution:</h3>

<p>(a) The correct truth table for the implication operator is the following,</p>

<table style="border: 1px solid black; border-collapse: collapse; text-align: center; margin-left: 30px;">
  <tr>
    <th style="width:100px; border: 1px solid black">\(A\)</th>
    <th style="width:100px; border: 1px solid black">\(B\)</th>
    <th style="width:200px; border: 1px solid black">\(A {@operator[2]@} B\)</th>
  </tr>
  <tr>
    <td style="border: 1px solid black">true</td>
    <td style="border: 1px solid black">true</td>
    <td style="border: 1px solid black"><b>{@t_ans1@}</b></td>
  </tr>
  <tr>
    <td style="border: 1px solid black">true</td>
    <td style="border: 1px solid black">false</td>
    <td style="border: 1px solid black"><b>{@t_ans2@}</b></td>
  </tr>
  <tr>
    <td style="border: 1px solid black">false</td>
    <td style="border: 1px solid black">true</td>
    <td style="border: 1px solid black"><b>{@t_ans3@}</b></td>
  </tr>
  <tr>
    <td style="border: 1px solid black">false</td>
    <td style="border: 1px solid black">false</td>
    <td style="border: 1px solid black"><b>{@t_ans4@}</b></td>
  </tr>
</table>

<br>

<p>(b) To evaluate the given statement, we can separate the two claims as,</p>
<p style="margin-left: 30px">
  \(A\): "{@statement_A[1]@}"
  <br>\(B\): "{@statement_B[1]@}"
</p>
<p>We can now use the truth table in part (a) to evaluate the full statement.</p>

<br>

<p>We know immediately from the information in the question that statement \(A\) is <b>{@if statement_A[2]=1 then "true" else "false"@}</b>, so we have a case of either the {@if statement_A[2]=1 then "first or second" else "third or fourth"@} row in the table.</p>

<table style="border: 1px solid black; border-collapse: collapse; text-align: center; margin-left: 30px;">
  <tr>
    <th style="width:100px; border: 1px solid black">\(A\)</th>
    <th style="width:100px; border: 1px solid black">\(B\)</th>
    <th style="width:200px; border: 1px solid black">\(A {@operator[2]@} B\)</th>
  </tr>
  <tr>
    <td style="border: 1px solid black">{@if statement_A[2]=1 then "true" else "false"@}</td>
    <td style="border: 1px solid black">true</td>
    <td style="border: 1px solid black"><b>{@worked_sol_opt1: if statement_A[2]=1 then t_ans1 else t_ans3@}</b></td>
  </tr>
  <tr>
    <td style="border: 1px solid black">{@if statement_A[2]=1 then "true" else "false"@}</td>
    <td style="border: 1px solid black">false</td>
    <td style="border: 1px solid black"><b>{@worked_sol_opt2: if statement_A[2]=1 then t_ans2 else t_ans4@}</b></td>
  </tr>
</table>

<br>

<p>{@
if worked_sol_opt1=worked_sol_opt2 then 

  sconcat("Then, from the table we see that <b>as long as there is not a contradiction</b>, the statement must be <b>", 
  worked_sol_opt1, 
  "</b>.<br>Supposing there is no contradiction and the statement is then ", 
  worked_sol_opt1, 
  ", I am therefore <b>",
  if worked_sol_opt1="true" then "telling the truth" else "not telling the truth",
  "</b>, so the claim \\(B\\) is <b>", 
  worked_sol_evB_1: if (worked_sol_opt1="true" and statement_B[2]=1) or (worked_sol_opt1="false" and statement_B[2]=2) then "true" else "false", 
  "</b>.<br> This corresponds to the ",
  if worked_sol_evB_1="true" then "first" else "second",
  " row in the shortened table above. <br>Hence, this case is permissible (there is no contradiction), and the statement is <b>",
  worked_sol_opt1,
  "</b>.")

else

  sconcat("Then, from the table we see that the <b>as long as there is not a contradiction</b>, the statement could be either true or false.</p>",

  "<p>Let us suppose first that the statement is <b>true</b>, so we are in the ",
  if worked_sol_opt1="true" then "first" else "second", 
  " row of the shortened truth table above. Since I am therefore telling the truth, we know statement \\(B\\) is <b>",
  worked_sol_evB_2: if statement_B[2]=1 then "true" else "false",
  "</b>.<br>Now looking at the relevant row in the shortened truth table, this ",
  if worked_sol_opt1=worked_sol_evB_2 then "<b>correctly</b> corresponds to the evaluation of \\(B\\), so this case is permissible and the statement <b>could be true</b>." else "<b>does not</b> correspond to the evaluation of \\(B\\), so we have a contradiction and the statement <b>cannot be true</b>.",
 "</p>",

  "<p>Let us suppose now that the statement is <b>false</b>, so we are in the ",
  if worked_sol_opt1="false" then "first" else "second", 
  " row of the shortened truth table above. Since I am therefore not telling the truth, we know statement \\(B\\) is <b>",
  worked_sol_evB_3: if statement_B[2]=2 then "true" else "false",
  "</b>.<br>Now looking at the relevant row in the shortened truth table, this ",
  if worked_sol_opt2=worked_sol_evB_3 then "<b>correctly</b> corresponds to the evaluation of \\(B\\), so this case is permissible and the statement <b>could be false</b>." else "<b>does not</b> correspond to the evaluation of \\(B\\), so we have a contradiction and the statement <b>cannot be false</b>.",
  "</p>",

  "<p>Then, considering both cases, ",
  if t_ans5=1 then 
    "the statement is <b>true</b>."
  else if t_ans5=2 then 
    "the statement is <b>false</b>."
  else if t_ans5=3 then 
    "the statement is <b>neither true nor false</b> since each option gives a contradiction."
  else if t_ans5=4 then 
    "the statement could be <b>either true or false</b>.")
@}</p>]]></text>
    </generalfeedback>
    <defaultgrade>4</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Defining the question information as a matrix */
/* [*operator name*, *latex symbol*, *word*, [*truth table*], *position in statement matrix lists*] */

operator: rand([
  ["a conjunction", "\\land", "and", ["true", "false", "false", "false"], 1], 
  ["a disjunction", "\\lor", "or", ["true", "true", "true", "false"], 2],
  ["an implication", "\\Rightarrow", "implies that", ["true", "false", "true", "true"], 3],
  ["a bi-implication", "\\iff", "if and only if", ["true", "false", "false", "true"], 4]
]);

statement_A: rand({["it is raining", 1], ["it is not raining", 2]});
statement_B: rand({["I am telling the truth", 1], ["I am not telling the truth", 2]});

/* Truth table inputs */

m_ans1: [
  ["true", is(operator[4][1]="true")],
  ["false", is(operator[4][1]="false")]
];
t_ans1: operator[4][1];

m_ans2: [
  ["true", is(operator[4][2]="true")],
  ["false", is(operator[4][2]="false")]
];
t_ans2: operator[4][2];

m_ans3: [
  ["true", is(operator[4][3]="true")],
  ["false", is(operator[4][3]="false")]
];
t_ans3: operator[4][3];

m_ans4: [
  ["true", is(operator[4][4]="true")],
  ["false", is(operator[4][4]="false")]
];
t_ans4: operator[4][4];

/* Statement inputs */
/* 
We store the truth evaluations of each possible statement as a matrix of lists, each position corresponds to a possible statement and each entry a list of the truth for [*conjunction*, *disjunction*, *implication*, *bi-implication*]

                     |  "raining"  | "not raining"
-----------------------------------------------------
"telling truth"      |    [***]    |    [***]   
-----------------------------------------------------
"not telling truth"  |    [***]    |    [***]   


*/

statement_matrix: [
  [[4, 1, 4, 4], [2, 4, 1, 3]],
  [[3, 1, 3, 3], [2, 3, 1, 4]]
];

m_ans5: [
  [1, is(statement_matrix[statement_B[2]][statement_A[2]][operator[5]]=1), "This statement is true"],
  [2, is(statement_matrix[statement_B[2]][statement_A[2]][operator[5]]=2), "This statement is false"],
  [3, is(statement_matrix[statement_B[2]][statement_A[2]][operator[5]]=3), "This statement is neither true nor false since each option gives a contradiction"],
  [4, is(statement_matrix[statement_B[2]][statement_A[2]][operator[5]]=4), "This statement could be either true or false"]
];

t_ans5: if m_ans5[1][2] then 1
  else if m_ans5[2][2] then 2
  else if m_ans5[3][2] then 3
  else if m_ans5[4][2] then 4;

/* Variables for question tests */

/* (a) */

/* Node 3: Check whether table fits a different connective (test using a conjunction for disj. or impl., and a disjunction for conj. or bi-impl. to ensure more than one row is incorrect) */
part_a_n3_1: if operator[5]=1 or operator[5]=4 then ["true", "true", "true", "false"] else ["true", "false", "false", "false"];]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Statement: {@statement_A[1]@} {@operator[3]@} {@statement_B[1]@}</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>m_ans1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>dropdown</type>
      <tans>m_ans2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>dropdown</type>
      <tans>m_ans3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>dropdown</type>
      <tans>m_ans4</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>radio</type>
      <tans>m_ans5</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[/* Node 1: Check truth table entirely correct */
node1_test: if [ans1,ans2,ans3,ans4]=[t_ans1,t_ans2,t_ans3,t_ans4] then 1 else 0;

/* Node 2: Three of four rows correct */
node2_test: if 
    [ans1,ans2,ans3]=[t_ans1,t_ans2,t_ans3]
    or [ans1,ans2,ans4]=[t_ans1,t_ans2,t_ans4]
    or [ans1,ans3,ans4]=[t_ans1,t_ans3,t_ans4]
    or [ans2,ans3,ans4]=[t_ans2,t_ans3,t_ans4]
  then 1 else 0;

/* Node 3: Table fits different connective */
node3_test: if elementp(makelist(i[4], i, operators)) then 1 else 0;
node3_feedback: if 
    [ans1,ans2,ans3,ans4]=["true", "false", "false", "false"]
  then "a conjunction" else if 
    [ans1,ans2,ans3,ans4]=["true", "true", "true", "false"]
  then "a disjunction" else if 
    [ans1,ans2,ans3,ans4]=["true", "false", "true", "true"]
  then "an implication" else if 
    [ans1,ans2,ans3,ans4]=["true", "false", "false", "true"]
  then "a bi-implication" else "";]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>node1_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>node2_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>node3_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>3</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><br>You have not given the correct truth table for {@operator[1]@}. 
<br>The truth table you have given is actually the truth table for {@node3_feedback@}.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>t_ans1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>4</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p><br>The first row of your truth table is incorrect.
<br>If \(A\) is true and \(B\) is true, then "\(A\) {@operator[3]@} \(B\)" is <b>{@t_ans1@}</b>.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>t_ans2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>5</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>5</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p><br>The second row of your truth table is incorrect.
<br>If \(A\) is true and \(B\) is false, then "\(A\) {@operator[3]@} \(B\)" is <b>{@t_ans2@}</b>.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>5</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>t_ans3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>6</truenextnode>
        <trueanswernote>prt1-6-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>6</falsenextnode>
        <falseanswernote>prt1-6-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p><br>The third row of your truth table is incorrect.
<br>If \(A\) is false and \(B\) is true, then "\(A\) {@operator[3]@} \(B\)" is <b>{@t_ans3@}</b>.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>6</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>t_ans4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-7-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-7-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p><br>The fourth row of your truth table is incorrect.
<br>If \(A\) is false and \(B\) is false, then "\(A\) {@operator[3]@} \(B\)" is <b>{@t_ans4@}</b>.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[/* Node 2: Claims the statement is true when it cannot be true */
node2_test: if ans5=1 and (t_ans5=2 or t_ans5=3) then 1 else 0;

/* Node 3: Claims the statement is true when it can be true or false */
node3_test: if ans5=1 and t_ans5=4 then 1 else 0;

/* Node 4: Claims the statement is false when it cannot be false */
node4_test: if ans5=2 and (t_ans5=1 or t_ans5=3) then 1 else 0;

/* Node 5: Claims the statement is false when it can be true or false */
node5_test: if ans5=2 and t_ans5=4 then 1 else 0;

/* Node 6: Incorrectly claims the statement is neither true nor false */
node6_test: if ans5=3 then 1 else 0;
node6_feedback: if t_ans5=1 then "true" else "false";

/* Node 7: Incorrectly claims the statement could be either true or false */
node7_test: if ans5=4 then 1 else 0;
node7_feedback: if t_ans5=1 then "false" else "true";]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>t_ans5</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>node2_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><br>The statement given cannot be true because this will lead to a <b>contradiction</b>.
<br>Writing the two claims in the statement as \(A\) and \(B\), for "{@statement_A[1]@}" and "{@statement_B[1]@}" respectively, we can use the truth table in part (a) to evaluate the whole statement.</p>

<p>From the information in the question, we have that statement \(A\) is <b>{@if statement_A[2]=1 then "true" else "false"@}</b>.
<br>We also have that if the overall statement is true, then I am telling the truth and statement \(B\) is <b>{@if statement_B[2]=1 then "true" else "false"@}</b>.</p>

<p>These two observations cannot be simultaneously true in any of the rows in the truth table for {@operator[1]@}.
<br>So, the statement "\(A\) {@operator[3]@} \(B\) " <b>cannot be true</b>.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>node3_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><br>You are correct that the statement given can be true. 
<br>However, it can also be false, so the correct answer is that the statement could be <b>either true or false</b>.
<br>Writing the two claims in the statement as \(A\) and \(B\), for "{@statement_A[1]@}" and "{@statement_B[1]@}" respectively, we can use the truth table in part (a) to evaluate the whole statement.</p>

<p>From the information in the question, we have that statement \(A\) is <b>{@if statement_A[2]=1 then "true" else "false"@}</b>.
<br>We also have that if the overall statement is false, then I am not telling the truth and statement \(B\) is <b>{@if statement_B[2]=2 then "true" else "false"@}</b>.</p>

<p>You can observe in the table for {@operator[1]@} that these two observations can be simultaneously true.
<br>So, the statement "\(A\) {@operator[3]@} \(B\) " <b>can be false</b>.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt2-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>node4_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><br>The statement given cannot be false because this will lead to a <b>contradiction</b>.
<br>Writing the two claims in the statement as \(A\) and \(B\), for "{@statement_A[1]@}" and "{@statement_B[1]@}" respectively, we can use the truth table in part (a) to evaluate the whole statement.</p>

<p>From the information in the question, we have that statement \(A\) is <b>{@if statement_A[2]=1 then "true" else "false"@}</b>.
<br>We also have that if the overall statement is false, then I am not telling the truth and statement \(B\) is <b>{@if statement_B[2]=2 then "true" else "false"@}</b>.</p>

<p>These two observations cannot be simultaneously true in any of the rows in the truth table for {@operator[1]@}.
<br>So, the statement "\(A\) {@operator[3]@} \(B\) " <b>cannot be false</b>.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt2-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>node5_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><br>You are correct that the statement given can be false. 
<br>However, it can also be true, so the correct answer is that the statement could be <b>either true or false</b>.
<br>Writing the two claims in the statement as \(A\) and \(B\), for "{@statement_A[1]@}" and "{@statement_B[1]@}" respectively, we can use the truth table in part (a) to evaluate the whole statement.</p>

<p>From the information in the question, we have that statement \(A\) is <b>{@if statement_A[2]=1 then "true" else "false"@}</b>.
<br>We also have that if the overall statement is true, then I am telling the truth and statement \(B\) is <b>{@if statement_B[2]=1 then "true" else "false"@}</b>.</p>

<p>You can observe in the table for {@operator[1]@} that these two observations can be simultaneously true.
<br>So, the statement "\(A\) {@operator[3]@} \(B\) " <b>can be true</b>.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>5</falsenextnode>
        <falseanswernote>prt2-5-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>5</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>node6_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-6-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><br>The statement given can in fact be <b>{@if t_ans5=1 then "true" else if t_ans5=2 then "false" else "either true or false"@}</b> without leading to a contradiction.
<br>Writing the two claims in the statement as \(A\) and \(B\), for "{@statement_A[1]@}" and "{@statement_B[1]@}" respectively, we can use the truth table in part (a) to evaluate the whole statement.</p>

<p>From the information in the question, we have that statement \(A\) is <b>{@if statement_A[2]=1 then "true" else "false"@}</b>.
<br>We also have that if the overall statement is {@node6_feedback@}, then I am {@if node6_feedback="false" then "not" else ""@} telling the truth and statement \(B\) is <b>{@if statement_B[2]=1 then node6_feedback else (if node6_feedback="true" then "false" else "true")@}</b>.</p>

<p>You can observe in the table for {@operator[1]@} that these two observations can be simultaneously true.
<br>So, the statement "\(A\) {@operator[3]@} \(B\) " <b>can be {@node6_feedback@}</b>.
{@if t_ans5=4 then "<br>A similar argument shows that the statement <b>can also be true</b>." else ""@}</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>6</falsenextnode>
        <falseanswernote>prt2-6-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>6</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>node7_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-7-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><br>The statement given in fact <b>cannot be {@if t_ans5=1 then "false" else if t_ans5=2 then "true" else "either true or false"@}</b> without leading to a contradiction.
<br>Writing the two claims in the statement as \(A\) and \(B\), for "{@statement_A[1]@}" and "{@statement_B[1]@}" respectively, we can use the truth table in part (a) to evaluate the whole statement.</p>

<p>From the information in the question, we have that statement \(A\) is <b>{@if statement_A[2]=1 then "true" else "false"@}</b>.
<br>We also have that if the overall statement is {@node7_feedback@}, then I am {@if node7_feedback="false" then "not" else ""@} telling the truth and statement \(B\) is <b>{@if statement_B[2]=1 then node7_feedback else (if node7_feedback="true" then "false" else "true")@}</b>.</p>

<p>These two observations cannot be simultaneously true in any of the rows in the truth table for {@operator[1]@}.
<br>So, the statement "\(A\) {@operator[3]@} \(B\) " <b>cannot be {@node7_feedback@}</b>.
{@if t_ans5=3 then "<br>A similar argument shows that the statement <b>also cannot be false</b>." else ""@}</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-7-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1077413899</deployedseed>
    <deployedseed>2109569196</deployedseed>
    <deployedseed>1855499228</deployedseed>
    <deployedseed>1514432889</deployedseed>
    <deployedseed>1815852674</deployedseed>
    <deployedseed>633413267</deployedseed>
    <deployedseed>1208291299</deployedseed>
    <deployedseed>787571236</deployedseed>
    <deployedseed>1952993703</deployedseed>
    <deployedseed>523260439</deployedseed>
    <deployedseed>708142620</deployedseed>
    <deployedseed>2013483403</deployedseed>
    <deployedseed>2075017288</deployedseed>
    <deployedseed>886113624</deployedseed>
    <deployedseed>352595944</deployedseed>
    <deployedseed>287708815</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>t_ans1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>t_ans2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>t_ans3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>t_ans4</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>t_ans5</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value><![CDATA[if t_ans1="true" then "false" else "true"]]></value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>t_ans2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>t_ans3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>t_ans4</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value/>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-7-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>t_ans1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value><![CDATA[if t_ans2="true" then "false" else "true"]]></value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>t_ans3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>t_ans4</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value/>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-7-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>t_ans1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>t_ans2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value><![CDATA[if t_ans3="true" then "false" else "true"]]></value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>t_ans4</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value/>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-7-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>5</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>t_ans1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>t_ans2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>t_ans3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value><![CDATA[if t_ans4="true" then "false" else "true"]]></value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value/>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-7-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>6</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>part_a_n3_1[1]</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>part_a_n3_1[2]</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>part_a_n3_1[3]</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>part_a_n3_1[4]</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value/>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-7-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
