<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Evaluating Statements and Truth Values</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>A statement is a sentence (or a mathematical expression) that has a well-defined truth value, <br>either <strong>True</strong> or <strong>False</strong>.</p>
<p>Use the definition of statement to say whether or not the following sentences are statements and justify your answer.</p>

<br>

<table>
<tr>

<td width="300px" style="padding:10px">(a) "\(8\) is not perfect square."</td>
<td style="padding:10px">[[input:ans1]] [[validation:ans1]] because [[input:ans2]] [[validation:ans2]]
<br>[[feedback:prt1]]<td>

</tr>
<tr>

<td style="padding:10px">(b) "\(15 - 6\)"</td>
<td style="padding:10px">[[input:ans3]] [[validation:ans3]] because [[input:ans4]] [[validation:ans4]]
<br>[[feedback:prt2]]</td>

</tr>
<tr>

<td style="padding:10px">(c) "\(15 - 6 = 0\)"</td>
<td style="padding:10px">[[input:anstC]] [[validation:anstC]] because [[input:ansC]] [[validation:ansC]]
<br>[[feedback:prtC]]</td>

</tr>
<tr>

<td style="padding:10px">(d) "\(3x - 2 = 10\)"</td>
<td style="padding:10px">[[input:ans5]] [[validation:ans5]] because [[input:ans6]] [[validation:ans6]]
<br>[[feedback:prt3]]</td>

</tr>
<tr>

<td style="padding:10px">(e) "if \(x = 4\), then \(3x - 2 = 10\)"</td>
<td style="padding:10px">[[input:anstE]] [[validation:anstE]] because [[input:ansE]] [[validation:ansE]]
<br>[[feedback:prtE]]</td>

</tr>
<tr>

<td style="padding:10px">(f) "\(8,000,000\) is a very big number."</td>
<td style="padding:10px">[[input:ans7]] [[validation:ans7]] because [[input:ans8]] [[validation:ans8]]
<br>[[feedback:prt4]]</td>

</tr>
<tr>

<td style="padding:10px">(g) "It is the highest peak in the world."</td>
<td style="padding:10px">[[input:ans9]] [[validation:ans9]] because [[input:ans10]] [[validation:ans10]]
<br>[[feedback:prt5]]</td>

</tr>
</table>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked solution</h3>

<p>(a) "\(8\) is not a perfect square."</p>
<p>This <b>is</b> a statement because it makes a clear assertion about the number 8, which can be objectively verified. In this case, the statement is <b>true</b> because \(8\) is not a perfect square (as it cannot be expressed as the square of an integer).</p>

<br>

<p>(b) "\(15 - 6\)"</p>
<p>This <b>is not</b> a statement because it is an incomplete mathematical expression. It lacks an operation, such as addition or subtraction, to form a valid mathematical statement. As a standalone expression, it doesn't have a truth value.</p>

<br>

<p>(c) "\(15 - 6 = 0\)"</p>
<p>This <b>is</b> a statement because it makes a clear assertion about a mathematical relationship between \(15\) and \(6\). However, the statement is false because \(15 - 6\) does not equal \(0\); it equals \(9\).</p>

<br>

<p>(d) "\(3x - 2 = 10\)"</p>
<p>This <b>is not</b> a statement because it is an incomplete mathematical equation. The presence of the variable \(x\) means that the truth value depends on the value assigned to \(x\). Without knowing the value of \(x\), we cannot determine the truth value of the equation.</p>

<br>

<p>(e) "if \(x=4\), then \(3x-2=10\)"</p>
<p>This <b>is</b> a statement because it presents a conditional relationship between two mathematical expressions. The statement is true because if \(x\) is equal to \(4\), then \(3x - 2\) equals \(10\).</p>

<br>

<p>(f) "\(8,000,000\) is a very big number."</p>
<p>This <b>is not</b> a statement because the term "very big" is subjective and lacks a well-defined truth value. What is considered "very big" can vary from person to person, making the statement unclear and subjective.</p>

<br>

<p>(g) "It is the highest peak in the world."</p>
<p>This <b>is not</b> a statement because it uses the pronoun "it" without specifying what "it" refers to. Without knowing the subject being referred to, the statement lacks clarity and a definite truth value.</p>]]></text>
    </generalfeedback>
    <defaultgrade>5</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[ta1: [
    ["(a) is a statement", true],
    ["(a) is not a statement", false]
];

ta2: random_permutation([
    ["it has only one truth value which is true", true],
    ["it has only one truth value which is false", false], 
    ["it is both true and false", false],
    ["it's impossible to say whether it is true or false", false]
]);

ta3: [
    ["(b) is a statement", false],
    ["(b) is not a statement", true]
];

ta4: random_permutation([
    ["it is a number, which has no truth value", true],
    ["it has only one truth value which is true", false],
    ["it has only one truth value which is false", false],
    ["it is both true and false", false]
]);

tatC: [
    ["(c) is a statement", true],
    ["(c) is not a statement", false]
];

taC: random_permutation([
    ["it is a number, which has no truth value", false],
    ["it has only one truth value which is true", false],
    ["it has only one truth value which is false", true],
    ["it is both true and false", false]
]);

ta5: [
    ["(d) is a statement", false],
    ["(d) is not a statement", true]
];

ta6: random_permutation([
    ["has only one truth value", false],
    ["we don't know the value of x", true]
]);

tatE: [
    ["(e) is a statement", true],
    ["(e) is not a statement", false]
];

taE: random_permutation([
    ["has only one truth value", true],
    ["we don't know the value of x", false]
]);

ta7: [
    ["(f) is a statement", false],
    ["(f) is not a statement", true]
];

ta8: random_permutation([
    ["it has truth value true.", false],
    ["everybody will agree that it is true", false],
    ["it is both true and false", false],
    ["'very big' is subjective and not formally defined", true]
]);

ta9: [
    ["(g) is a statement", false],
    ["(g) is not a statement", true]
];

ta10: random_permutation([
    ["it has only one truth value which is false", false],
    ["it is both true and false", false],
    ["it is true because it must refer to Mount Everest", false],
    ["we don't know what 'it' refers to", true]
]);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Model Answers</p>
<p><b>Note:</b> There is no randomization</p>
<p>{@mcq_correct(ta1)[1]@}</p>
<p>{@mcq_correct(ta2)[1]@}</p>
<p>{@mcq_correct(ta3)[1]@}</p>
<p>{@mcq_correct(ta4)[1]@}</p>
<p>{@mcq_correct(ta5)[1]@}</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> 
Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> 
Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans10</name>
      <type>dropdown</type>
      <tans>ta10</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>dropdown</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>dropdown</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>dropdown</type>
      <tans>ta4</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>dropdown</type>
      <tans>ta5</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6</name>
      <type>dropdown</type>
      <tans>ta6</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans7</name>
      <type>dropdown</type>
      <tans>ta7</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans8</name>
      <type>dropdown</type>
      <tans>ta8</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans9</name>
      <type>dropdown</type>
      <tans>ta9</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ansC</name>
      <type>dropdown</type>
      <tans>taC</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ansE</name>
      <type>dropdown</type>
      <tans>taE</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>anstC</name>
      <type>dropdown</type>
      <tans>tatC</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>anstE</name>
      <type>dropdown</type>
      <tans>tatE</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/
/*Node 2: Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(ta1)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>mcq_correct(ta2)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/
/*Node 2: Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>mcq_correct(ta3)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>mcq_correct(ta4)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/
/*Node 2: Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>mcq_correct(ta5)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans6</sans>
        <tans>mcq_correct(ta6)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/
/*Node 2: Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans7</sans>
        <tans>mcq_correct(ta7)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans8</sans>
        <tans>mcq_correct(ta8)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/
/*Node 2: Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans9</sans>
        <tans>mcq_correct(ta9)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans10</sans>
        <tans>mcq_correct(ta10)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtC</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/
/*Node 2: Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>anstC</sans>
        <tans>mcq_correct(tatC)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prtC-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prtC-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ansC</sans>
        <tans>mcq_correct(taC)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtC-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtC-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtE</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/
/*Node 2: Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>anstE</sans>
        <tans>mcq_correct(tatE)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prtE-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prtE-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ansE</sans>
        <tans>mcq_correct(taE)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtE-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtE-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1966326540</deployedseed>
    <deployedseed>1723392621</deployedseed>
    <deployedseed>932575373</deployedseed>
    <deployedseed>959267870</deployedseed>
    <deployedseed>1005613358</deployedseed>
    <deployedseed>1281569436</deployedseed>
    <deployedseed>170577516</deployedseed>
    <deployedseed>2043773255</deployedseed>
    <deployedseed>110367984</deployedseed>
    <deployedseed>506490342</deployedseed>
    <deployedseed>290793937</deployedseed>
    <deployedseed>1570593909</deployedseed>
    <deployedseed>1937587715</deployedseed>
    <deployedseed>1350113944</deployedseed>
    <deployedseed>1910460950</deployedseed>
    <deployedseed>118916425</deployedseed>
    <deployedseed>1622613861</deployedseed>
    <deployedseed>1067137774</deployedseed>
    <deployedseed>1233889400</deployedseed>
    <deployedseed>1330981779</deployedseed>
    <deployedseed>593929362</deployedseed>
    <deployedseed>1182250530</deployedseed>
    <deployedseed>783016535</deployedseed>
    <deployedseed>1012219305</deployedseed>
    <deployedseed>287722311</deployedseed>
    <deployedseed>1880246396</deployedseed>
    <deployedseed>1745321322</deployedseed>
    <deployedseed>1036789335</deployedseed>
    <deployedseed>671433442</deployedseed>
    <deployedseed>1242252043</deployedseed>
    <deployedseed>1233022978</deployedseed>
    <deployedseed>1434780189</deployedseed>
    <deployedseed>656704998</deployedseed>
    <deployedseed>608991713</deployedseed>
    <deployedseed>950456112</deployedseed>
    <deployedseed>1121311362</deployedseed>
    <deployedseed>305596032</deployedseed>
    <deployedseed>1800623712</deployedseed>
    <deployedseed>303574155</deployedseed>
    <deployedseed>1595177587</deployedseed>
    <deployedseed>1496242881</deployedseed>
    <deployedseed>2049496427</deployedseed>
    <deployedseed>1538721209</deployedseed>
    <deployedseed>488646906</deployedseed>
    <deployedseed>1785765543</deployedseed>
    <deployedseed>403906946</deployedseed>
    <deployedseed>1694940147</deployedseed>
    <deployedseed>1863298238</deployedseed>
    <deployedseed>589415093</deployedseed>
    <deployedseed>1012898919</deployedseed>
    <deployedseed>1698141789</deployedseed>
    <deployedseed>1382540122</deployedseed>
    <deployedseed>9975661</deployedseed>
    <deployedseed>1331479580</deployedseed>
    <deployedseed>279312927</deployedseed>
    <deployedseed>584989314</deployedseed>
    <deployedseed>1410772628</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>first(mcq_correct(ta1))</value>
      </testinput>
      <testinput>
        <name>ans10</name>
        <value>first(mcq_correct(ta10))</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>first(mcq_correct(ta2))</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>first(mcq_correct(ta3))</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>first(mcq_correct(ta4))</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>first(mcq_correct(ta5))</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>first(mcq_correct(ta6))</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>first(mcq_correct(ta7))</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>first(mcq_correct(ta8))</value>
      </testinput>
      <testinput>
        <name>ans9</name>
        <value>first(mcq_correct(ta9))</value>
      </testinput>
      <testinput>
        <name>ansC</name>
        <value>first(mcq_correct(taC))</value>
      </testinput>
      <testinput>
        <name>ansE</name>
        <value>first(mcq_correct(taE))</value>
      </testinput>
      <testinput>
        <name>anstC</name>
        <value>first(mcq_correct(tatC))</value>
      </testinput>
      <testinput>
        <name>anstE</name>
        <value>first(mcq_correct(tatE))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt5-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prtC</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtC-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prtE</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtE-2-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
