<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Equation of Tangent Line to Implicit Curve</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Find the equation of the tangent line to \(\displaystyle {@func_x@}+{@func_y@} ={@lhs_question_func@}\) at ({@x_point@},{@y_point@}).</p>
<p>Follow the outlined steps below. Fill in the missing blanks:</p>

<p><b>Step 1:</b> Differentiate \(\displaystyle {@func_x@}+{@func_y@} ={@lhs_question_func@}\) implicitly with respect to \(x\). </p>
<p><b>Hint:</b> \(a \cdot x^2+ b \cdot y^4 \left( \frac{dy}{dx} \right)=c\) is typed as \(a*x^2  +b*y^4*diff(y,x)=c\) </p>
<br>  [[input:ans4]] [[validation:ans4]] [[feedback:prt4]]

<p><b>Step 2:</b> Make  \(\displaystyle y^{\prime} \) which is \( \frac{dy}{dx} \) the subject.</p>
<br> \(y'\)= [[input:ans1]] [[validation:ans1]][[feedback:prt1]] 

<p><b>Step 3:</b>  Evaluate the derivative \(y'\) at ({@x_point@},{@y_point@}) to get the slope \(m\) of the tangent line. </p>
<br> \(m=\) [[input:ans2]] [[validation:ans2]][[feedback:prt2]] </p>

<p><b>Step 4:</b> Write down the equation of the tangent line. </p> 
<br> [[input:ans3]] [[validation:ans3]][[feedback:prt3]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p><b>Worked solution</b></p>
<p><b>Step 1:</b> To differentiate  \(\displaystyle {@func_x@}+{@func_y@} ={@lhs_question_func@}\) implicitly with respect to \(x\), we apply the chain rule to the terms involving \(y\). The chain rule states that \( \frac{{d}}{{dx}}(f(y)) = f'(y) \cdot \frac{{dy}}{{dx}}\).</p>
<p>Differentiating each term with respect to \(x\):</p>
<p>\( \displaystyle \frac{{d}}{{dx}}({@func_x@}) + \frac{{d}}{{dx}}({@func_y@}) = \displaystyle \frac{{d}}{{dx}}({@lhs_question_func@})\)</p>
<p>The derivative of \({@func_x@}\) with respect to \(x\) is \({@aa*(coef_x)*x^(aa-1)@}\).
The derivative of \({@func_y@}\) with respect to \(y\) is \({@bb*(coef_y)*y^(bb-1)@}\), and we multiply by \(\displaystyle \frac{{dy}}{{dx}}\) to apply the chain rule.</p>
<p>So, the equation becomes:</p>
<p>\(\displaystyle  {@aa*(coef_x)*x^(aa-1)@}+ {@bb*(coef_y)*y^(bb-1)@} \cdot \frac{{dy}}{{dx}} = 0\).</p>
<p><b>Step 2:</b>To make \(\displaystyle \frac{{dy}}{{dx}}\) which is (\(y^{\prime} \)) the subject, we isolate it on one side of the equation:</p>
<p>\( \displaystyle  {@bb*(coef_y)*y^(bb-1)@} \cdot \frac{{dy}}{{dx}} = \displaystyle  -{@aa*(coef_x)*x^(aa-1)@}\)</p>
<p>Now, divide both sides by \({@bb*(coef_y)*y^(bb-1)@}\):</p>
<p>\(\displaystyle  \frac{{dy}}{{dx}} = \displaystyle  \frac{{-{@aa*(coef_x)*x^(aa-1)@}}}{{{@bb*(coef_y)*y^(bb-1)@}}}\)</p>
<p>Therefore, \(y^{\prime} = \displaystyle  {@ta1@}\)</p>
<p>Step 3: The slope \(m\) is obtained by evaluating \(y^{ \prime } \) at \( ({@x_point@},{@y_point@})\). We substitute \(x ={@x_point@}\) and \(y ={@y_point@}\) into the expression for \(y^{\prime}\);</p>
<p>\(
\begin{align}
\displaystyle \frac{{dy}}{{dx}} & = \displaystyle  \frac{{-{@aa*(coef_x)*x^(aa-1)@}}}{{{@bb*(coef_y)*y^(bb-1)@}}} \\
m & =  \displaystyle \frac{{-{@aa*(coef_x)@} \times ({@x_point@})^{@(aa-1)@}}}{{{@bb*(coef_y)@} \times ({@y_point@})^{@(bb-1)@}}}\\
m &= \displaystyle {@ratsimp((-aa*(coef_x)*x_point^(aa-1))/(bb*(coef_y)*y_point^(bb-1)))@} 
\end{align}\)</p>
<p>The slope \( m \) of the tangent line at the point \( ({@x_point@},{@y_point@})\) is given by \(m=\displaystyle 
 {@ta2@}\).</p>
<p><b>Step 4:</b> Write down the equation of the tangent line.</p> 
<p>\(y-y_1=m(x-x_1) \implies y=m(x-x_1)+y_1  \)</p>
</p>By substituting the values of \(m\), \(x_{1}\) and \(y_{1}\) in the above expression, we get;<p>
<br><p>\(\displaystyle  {@ta3@}\implies  \displaystyle  {@expand(ta3)@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*Question func*/
coef_x:2+rand_with_step(-1,3,1);
coef_y:2+rand_with_step(-1,3,1);
aa:2+rand(3);
bb:2+rand(3)
func_x:(coef_x)*x^(aa);
func_y:(coef_y)*y^(bb);

question_func: func_x+func_y;

lhs_question_func:2+rand(4);
x_point:2+rand(4);
y_point:-sqrt(2+rand(4));

/*Worked solution*/
diff_func1:diff(question_func,x);
diff_func2:diff(question_func,y);

tta:'diff(y,x); /*depends(y,x)*/
ta4:  diff_func1 + diff_func2*tta=0; /*input ans 4*/

ta1:-diff_func1/diff_func2; /*input ans1: Gradient function*/
ta21:subst(x_point,x,ta1);
ta2:subst(y_point,y,ta21); /*input ans2: Gradient*/

ta3: y=ta2*(x- x_point)+y_point;  /*input ans3: Equation of the tangent line*/</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>The given equation: \(\displaystyle {@func_x@}+{@func_y@} ={@lhs_question_func@}\) at ({@x_point@},{@y_point@})</p>
<p> Differentiating with respect to \(x\) we get </p>
<p> \({@ta4@}\)</p>
<p> \(y'\)= \({@ta1@}\)</p>
<p>The slope \(m\) of the tangent line \(={@ta2@}\)</p>
<p> Writing down the equation of the tangent line we get</p> 
<p> \({@ta3@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*node1: Checking for the correct answer*/
/*node2: Making y the subject of the formula instead of y'*/
simp:false;
node2:solve(question_func=lhs_question_func,y)[1];
node2_prt1:rhs(node2);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>EqualComAss</answertest>
        <sans>ans1</sans>
        <tans>node2_prt1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You just made \( y \) the subject of expression {@question_func@}, you need to make {@tta@} the subject of the expression. </p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*node1: checking for correct answer*/
/*node2: interchanging the points during substitution*/
var1:subst(x_point,y,ta1);
mistake1a:subst(y_point,x,var1);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake1a</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>It looks like you have interchanged the points during substitution, you need to substitute \(x\) with {@x_point@} and \(y\) with {@y_point@} in the equation to find \( y' \)</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*node1:checking for correct answer*/
/*node2: using the wrong gradient in ans2 to compute the equation of the tangent*/
mistake1b:ans2*(x- x_point)-ans2;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>mistake1b</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.3</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>It looks like you have correctly computed the equation of the tangent using the incorrect answer in the first input answer box above. However, this is not the model answer which was expected here.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*node1: Checking for correct answer*/

/*node2: Checking if the student misses out the dy/dx in their input*/
ans_node2:diff_func1 + diff_func2=0;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ans_node2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Note that you need to include {@tta@} as part of the derivative of {@func_y@}. Your input would be marked correct if you had {@diff_func2*tta@} as part of your answer.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>405727304</deployedseed>
    <deployedseed>514687066</deployedseed>
    <deployedseed>638460690</deployedseed>
    <deployedseed>1967234033</deployedseed>
    <deployedseed>439963560</deployedseed>
    <deployedseed>1029366489</deployedseed>
    <deployedseed>1034793950</deployedseed>
    <deployedseed>2142721889</deployedseed>
    <deployedseed>507177769</deployedseed>
    <deployedseed>1642848252</deployedseed>
    <deployedseed>886578622</deployedseed>
    <deployedseed>40119891</deployedseed>
    <deployedseed>1917449938</deployedseed>
    <deployedseed>1675779445</deployedseed>
    <deployedseed>569165240</deployedseed>
    <deployedseed>728348844</deployedseed>
    <deployedseed>97919094</deployedseed>
    <deployedseed>1552843324</deployedseed>
    <deployedseed>1451320964</deployedseed>
    <deployedseed>140782634</deployedseed>
    <deployedseed>834694105</deployedseed>
    <deployedseed>553517688</deployedseed>
    <deployedseed>77988781</deployedseed>
    <deployedseed>401914422</deployedseed>
    <deployedseed>1688080242</deployedseed>
    <deployedseed>1593223715</deployedseed>
    <deployedseed>227857168</deployedseed>
    <deployedseed>1744218257</deployedseed>
    <deployedseed>1589922362</deployedseed>
    <deployedseed>617357512</deployedseed>
    <deployedseed>1393232572</deployedseed>
    <deployedseed>1557925369</deployedseed>
    <deployedseed>1975409235</deployedseed>
    <deployedseed>1644675872</deployedseed>
    <deployedseed>1096671001</deployedseed>
    <deployedseed>1295389710</deployedseed>
    <deployedseed>909645543</deployedseed>
    <deployedseed>401448046</deployedseed>
    <deployedseed>564126546</deployedseed>
    <deployedseed>1479162846</deployedseed>
    <deployedseed>39684252</deployedseed>
    <deployedseed>374969975</deployedseed>
    <deployedseed>1541986306</deployedseed>
    <deployedseed>525390488</deployedseed>
    <deployedseed>1570549656</deployedseed>
    <deployedseed>928846578</deployedseed>
    <deployedseed>992537092</deployedseed>
    <deployedseed>1457520005</deployedseed>
    <deployedseed>2123545381</deployedseed>
    <deployedseed>1155146102</deployedseed>
    <deployedseed>2116725299</deployedseed>
    <deployedseed>425714350</deployedseed>
    <deployedseed>1723983305</deployedseed>
    <deployedseed>1611686408</deployedseed>
    <deployedseed>1150812434</deployedseed>
    <deployedseed>1998123553</deployedseed>
    <deployedseed>929719091</deployedseed>
    <deployedseed>961741829</deployedseed>
    <deployedseed>2079898055</deployedseed>
    <deployedseed>766518228</deployedseed>
    <deployedseed>742236484</deployedseed>
    <deployedseed>1578339387</deployedseed>
    <deployedseed>1868858389</deployedseed>
    <deployedseed>1191385774</deployedseed>
    <deployedseed>643856421</deployedseed>
    <deployedseed>1015917032</deployedseed>
    <deployedseed>125745505</deployedseed>
    <deployedseed>48042762</deployedseed>
    <deployedseed>538973577</deployedseed>
    <deployedseed>1635585426</deployedseed>
    <deployedseed>108144000</deployedseed>
    <deployedseed>1937557922</deployedseed>
    <deployedseed>905972815</deployedseed>
    <deployedseed>1030098302</deployedseed>
    <deployedseed>1348030557</deployedseed>
    <deployedseed>98857534</deployedseed>
    <deployedseed>1448016270</deployedseed>
    <deployedseed>909670340</deployedseed>
    <deployedseed>1590751030</deployedseed>
    <deployedseed>1494972302</deployedseed>
    <deployedseed>911153874</deployedseed>
    <deployedseed>180423830</deployedseed>
    <deployedseed>1699289806</deployedseed>
    <deployedseed>1777265752</deployedseed>
    <deployedseed>1601877155</deployedseed>
    <deployedseed>1612186495</deployedseed>
    <deployedseed>1778176280</deployedseed>
    <deployedseed>1842912049</deployedseed>
    <deployedseed>1386636983</deployedseed>
    <deployedseed>759136996</deployedseed>
    <deployedseed>496110568</deployedseed>
    <deployedseed>1966527688</deployedseed>
    <deployedseed>1892172973</deployedseed>
    <deployedseed>338817729</deployedseed>
    <deployedseed>1539869168</deployedseed>
    <deployedseed>1968924892</deployedseed>
    <deployedseed>1708999077</deployedseed>
    <deployedseed>1988461209</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
