<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Introduction to sets and  set notation</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>a) Which one of the following statements defines a set  and explain why set theory is important in mathematics?
[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>
<p>b) Consider a set \(A\) of even numbers less than \({@selected_number@}\): </p>
<p>Select statements which hold for elements in set \(A\).</p>
<p>[[input:ans2]] [[validation:ans2]][[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p>A set is a well-defined collection of objects that share certain properties or characteristics. </p>
<p>In set theory, sets serve as the fundamental building blocks, and they are used to organize and classify various mathematical objects.</p>
<p>Set theory is vital in mathematics because it provides a rigorous foundation for other mathematical concepts and operations. </p>
<p>It helps define relationships between different elements, aids in understanding mathematical structures and enables the development of formal mathematical proofs.</p>
<p>Example of a well-defined set:</p>
<p>Let's consider the set \(A\) of even numbers less than \({@selected_number@}\):</p>
<p>\(A = {@even_numbers@}\)</p>
<p>In this example, \(A\) is a well-defined set because it clearly states the property that all its elements are even numbers less than \({@selected_number@}\).</p>
<p>Illustration using notation:</p>
<p>To represent the elements of set \(A\) using notation:</p>
<p>\({@ta_2_a@} ({@element_to_check@}\) is an element of \(A)\)
<p>\({@ta_2_b@} ({@element_to_check@}\) is not an element of \(A)\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[random_numbers: [10, 20, 30, 40, 50]; /* List of random numbers */

random_index: rand(length(random_numbers)) + 1; /* Generate a random index */
selected_number: random_numbers[random_index]; /* Select the number at the random index */
even_numbers: {};

for num: 2 step 2 while num < selected_number do (
    even_numbers: adjoin(num, even_numbers)
);

var_even:even_numbers;
position:rand([first,third,fourth,second]);
element_to_check: position(even_numbers); 
var_member:if member(element_to_check, even_numbers) then "\\(element_to_check ∈ A\\)"
else "\\(element_to_check ∉ A\\)";


ta_1:random_permutation([["A set is a collection of numbers arranged in ascending order. Set theory simplifies arithmetic calculations.",false],["A set is a well-defined collection of objects that share specific properties or characteristics. Set theory provides a rigorous foundation for other concepts, defines relationships between elements, aids in understanding structures, and enables formal proofs.",true],["A set is a grouping of unrelated objects.
Set theory helps in creating art and design.",false],["A set is a sequence of letters that form a word. Set theory is used in language studies.",false]]);
ta_2_a: sconcat(element_to_check, " ∈ A ");
ta_2_b: sconcat(element_to_check, " ∉ A ");
ta_2:random_permutation([[ta_2_a,true],[ta_2_b,false],["Elements in set A are odd numbers",false]]);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>{@ta_2@}</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>radio</type>
      <tans>ta_1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <input>
      <name>ans2</name>
      <type>radio</type>
      <tans>ta_2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(ta_1)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>mcq_correct(ta_2)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>first(mcq_correct(ta_1))</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>first(mcq_correct(ta_2))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
