<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Relationship between elements and sets  notation</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Answer the following questions to get a better understanding of the set theory concept.</p>
 <p>1. If \(A\) is a set and \(x\) is an object, how do we use notation to signify whether \(x\) is an element of set \(A\) or not?</p>
<p> a) To show \(x\) is an element of \(A\) [[input:ans1]] [[validation:ans1]]</p>
 <p>b) To show  \(x\) is not an element of \(A\) [[input:ans2]][[validation:ans2]]</p>
<p>[[feedback:prt1]]</p>
<p>2. Consider the sets \(A\) and \(B\):</p>
<p>Set \(A\) contains prime numbers less than \({@selected_number@}\) .</p>
<p>Set \(B\) contains the numbers \({@var_even@}\).</p>
<p>For each statement below, choose the suitable set notation that best describes the relationship between the sets.</p>
[[input:ans3]][[validation:ans3]][[feedback:prt3]]]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p>In set theory, elements are individual objects that belong to a particular set. If \(A\) is a set and \(x\) is an object, we use the notation "\(x \in A\)" to indicate that \(x\) is an element of set \(A\). Conversely, if \(x\) does not belong to set \(A\), we use the notation "\(x \notin A\)".</p>
<p>Example of a set and elements relationship:</p>
<p>Consider the set \(B\) of prime numbers less than \({@selected_number@}\):</p>
<p>\(B = {@var_prime@}\)</p>
<p>Using notation, we can express the relationship between elements and set \(B\) as follows:
<p>\({@element_to_check@} \in  B  ({@element_to_check@}\) is a prime number and is an element of \(B)\)</p>
<p>\({@element_to_check_1@} \notin  B  ({@element_to_check_1@}\) is not a prime number and is not an element of \(B)\)</p>

<p>In summary, the notation"\(x \in A\)"  indicates that \(x\) is an element of set \(A,\) while  "\(x \notin A\)" indicates that \(x\) is not an element of set \(A\). Set theory helps us understand the relationships between elements and sets, enabling us to reason about various mathematical structures effectively.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/*variables for prime numbers*/
prime_numbers: [2, 3, 5, 7, 11, 13, 17, 19, 23, 29, 31, 37, 41, 43, 47]; /* Your list of prime numbers */

random_index: rand(length(prime_numbers)) + 1; /* Generate a random index */
selected_number: prime_numbers[random_index]; /* Select the number at the random index */
prime_numbers_less_than_selected: {};

for num in prime_numbers do (
    if num < selected_number then (
        prime_numbers_less_than_selected: adjoin(num, prime_numbers_less_than_selected)
    )
);

var_prime: prime_numbers_less_than_selected;
position:rand([first,third,fourth,second]);
element_to_check: position(prime_numbers_less_than_selected); 
var_member:if member(element_to_check, prime_numbers_less_than_selected) then "\\(element_to_check ∈ A\\)"
else "\\(element_to_check ∉ A\\)";
/*variables for even numbers*/
random_numbers: [10, 20, 30, 40, 50]; /* Your list of random numbers */

random_index: rand(length(random_numbers)) + 1; /* Generate a random index */
selected_number_1: random_numbers[random_index]; /* Select the number at the random index */
even_numbers: {};

for num: 2 step 2 while num < selected_number_1 do (
    even_numbers: adjoin(num, even_numbers)
);
var_even:even_numbers;
position:rand([first,third,fourth,second]);
element_to_check_1: position(even_numbers); 
var_member:if member(element_to_check_1, even_numbers) then "\\(element_to_check_1 ∈ A\\)"
else "\\(element_to_check_1 ∉ A\\)";


ta_1:random_permutation ([["x ∉ A",false],["x ∈ A",true],["x ㄷ A",false]]);
ta_2:random_permutation ([["x ∉ A",true],["x ∈ A",false],["x ㄷ A",false]]);

ta_3_a: sconcat(element_to_check, " ∈ A ");
ta_3_b: sconcat(element_to_check_1, " ∉ A ");
ta_3_c: sconcat("B ㄷ A");
ta_3_d: sconcat("A ㄷ B");
ta_3:random_permutation([[ta_3_a,true],[ta_3_b,true],[ta_3_c,false],[ta_3_d,false]]);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\({@ta_1@}\)</p>
<p>\({@ta_2@}\)</p>
<p>\({@ta_3@}\)</p>
<p>\({@selected_number@}\)</p>
<p> \({@var_even@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>radio</type>
      <tans>ta_1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <input>
      <name>ans2</name>
      <type>radio</type>
      <tans>ta_2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <input>
      <name>ans3</name>
      <type>checkbox</type>
      <tans>ta_3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(ta_1)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Remember, to express that an object \( x \) is an element of set \( A \), you should use the notation \( x \in A \).</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>mcq_correct(ta_2)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Remember, to express that an object \( x \) is not an element of set \( A \), you should use the notation \( x \notin A \).  </p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>mcq_correct(ta_3)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>595633336</deployedseed>
    <deployedseed>1035473293</deployedseed>
    <deployedseed>537065522</deployedseed>
    <deployedseed>1412378921</deployedseed>
    <deployedseed>598650226</deployedseed>
    <deployedseed>365033703</deployedseed>
    <deployedseed>97484396</deployedseed>
    <deployedseed>552798541</deployedseed>
    <deployedseed>321431416</deployedseed>
    <deployedseed>1107362640</deployedseed>
    <deployedseed>1959451871</deployedseed>
    <deployedseed>1041390328</deployedseed>
    <deployedseed>220700190</deployedseed>
    <deployedseed>1566315287</deployedseed>
    <deployedseed>493966058</deployedseed>
    <deployedseed>2003844552</deployedseed>
    <deployedseed>1918575727</deployedseed>
    <deployedseed>966999930</deployedseed>
    <deployedseed>907343728</deployedseed>
    <deployedseed>1198585579</deployedseed>
    <deployedseed>2031347994</deployedseed>
    <deployedseed>478908022</deployedseed>
    <deployedseed>901822323</deployedseed>
    <deployedseed>1208481382</deployedseed>
    <deployedseed>426058151</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>first(mcq_correct(ta_1))</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>first(mcq_correct(ta_2))</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>mcq_correct(ta_3)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
