<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Exercise 1 and 2 Matrix Operations Quiz</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Let \({\bf A}\), \({\bf B}\) and \({\bf C}\) be the following matrices: \[{\bf A} = {@A@}\]&nbsp;<br>\[{\bf B} = {@B@}\]&nbsp;\[{\bf C} = {@C@}\]</p>
<p><br></p>
<p>1. Calculate&nbsp;\({\bf 2 \cdot B+ 3 \cdot C}\)</p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>
<p></p>
<p>2. Calculate&nbsp;\({\bf A \cdot B}\)</p>
<p>[[input:ans2]] [[validation:ans2]]</p>
<p></p>
<p>3. Calculate&nbsp;\({\bf 2 \cdot B \cdot A}\)</p>
<p>[[input:ans3]] [[validation:ans3]]</p>
<p>4. Calculate&nbsp;\({\bf C \cdot A - B \cdot A}\)</p>
<p>[[input:ans4]] [[validation:ans4]]</p>
<p></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text/>
    </generalfeedback>
    <defaultgrade>2</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023010400</text>
    </stackversion>
    <questionvariables>
      <text>A:matrix([rand_with_step(-5,5,1),rand_with_step(-5,5,1)],[rand_with_step(-5,5,1),rand_with_step(-5,5,1)], [rand_with_step(-5,5,1),rand_with_step(-5,5,1)]);
B:matrix([rand_with_step(-5,5,1),rand_with_step(-5,5,1),rand_with_step(-5,5,1)],[rand_with_step(-5,5,1),rand_with_step(-5,5,1),rand_with_step(-5,5,1)]);
C:matrix([rand_with_step(-5,5,1),rand_with_step(-5,5,1),rand_with_step(-5,5,1)],[rand_with_step(-5,5,1),rand_with_step(-5,5,1),rand_with_step(-5,5,1)]);

ta1: 2*B + 3*C;
ta2: A.B;
ta3: 2*B.A;
ta4: C.A - B.A;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\({\bf A} = {@A@}\)</p>
<p>\({\bf B} = {@B@}\) </p>
<p>\({\bf C} = {@C@}\)</p>
<p>\({\bf 2 \cdot B+ 3 \cdot C}={@ta1@}\)</p>
<p>\({\bf A \cdot B}={@ta2@}\)</p>
<p>\({\bf 2 \cdot B \cdot A}={@ta3@}\)</p>
<p>\({\bf C \cdot A - B \cdot A}={@ta4@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>matrix</type>
      <tans>ta1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>matrix</type>
      <tans>ta2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>matrix</type>
      <tans>ta3</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>matrix</type>
      <tans>ta4</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the first correct answer */
/*Node 2: Checks for the second correct answer */
/*Node 3: Checks for the third correct answer */
/*Node 4: Checks for the fourth correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>3</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1670954098</deployedseed>
    <deployedseed>981269327</deployedseed>
    <deployedseed>1292773660</deployedseed>
    <deployedseed>20610338</deployedseed>
    <deployedseed>1868475277</deployedseed>
    <deployedseed>837478814</deployedseed>
    <deployedseed>307690569</deployedseed>
    <deployedseed>472438594</deployedseed>
    <deployedseed>1655838993</deployedseed>
    <deployedseed>1514022509</deployedseed>
    <deployedseed>269955639</deployedseed>
    <deployedseed>1551135852</deployedseed>
    <deployedseed>1558921234</deployedseed>
    <deployedseed>2006705585</deployedseed>
    <deployedseed>1097040424</deployedseed>
    <deployedseed>1283419824</deployedseed>
    <deployedseed>1465219917</deployedseed>
    <deployedseed>1006557946</deployedseed>
    <deployedseed>1644725539</deployedseed>
    <deployedseed>1375101758</deployedseed>
    <deployedseed>336084010</deployedseed>
    <deployedseed>1753261024</deployedseed>
    <deployedseed>1171289478</deployedseed>
    <deployedseed>1069287496</deployedseed>
    <deployedseed>43863097</deployedseed>
    <deployedseed>2058986126</deployedseed>
    <deployedseed>580017623</deployedseed>
    <deployedseed>452156242</deployedseed>
    <deployedseed>251438677</deployedseed>
    <deployedseed>559034848</deployedseed>
    <deployedseed>883759164</deployedseed>
    <deployedseed>298199970</deployedseed>
    <deployedseed>1548646141</deployedseed>
    <deployedseed>1590808095</deployedseed>
    <deployedseed>1202744466</deployedseed>
    <deployedseed>2094841244</deployedseed>
    <deployedseed>1207467617</deployedseed>
    <deployedseed>799005316</deployedseed>
    <deployedseed>1012126934</deployedseed>
    <deployedseed>599805221</deployedseed>
    <deployedseed>1061307226</deployedseed>
    <deployedseed>1255883082</deployedseed>
    <deployedseed>73566593</deployedseed>
    <deployedseed>477968352</deployedseed>
    <deployedseed>2017433111</deployedseed>
    <deployedseed>444979425</deployedseed>
    <deployedseed>546010306</deployedseed>
    <deployedseed>1924146795</deployedseed>
    <deployedseed>330644719</deployedseed>
    <deployedseed>1499821964</deployedseed>
    <deployedseed>1283152934</deployedseed>
    <deployedseed>789562062</deployedseed>
    <deployedseed>579629086</deployedseed>
    <deployedseed>326089245</deployedseed>
    <deployedseed>367035992</deployedseed>
    <deployedseed>2075823667</deployedseed>
    <deployedseed>1695540898</deployedseed>
    <deployedseed>1954715813</deployedseed>
    <deployedseed>1486123303</deployedseed>
    <deployedseed>950902885</deployedseed>
    <deployedseed>2035576754</deployedseed>
    <deployedseed>2131621818</deployedseed>
    <deployedseed>792208172</deployedseed>
    <deployedseed>1869750526</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-4-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
