<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Inverse Function Computation</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p></p><p>Find the inverse of \( \displaystyle R(x)={@rx@}\).</p><p>\(R^{-1}(x)=\) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]<br></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>To find the inverse of the function \( \displaystyle R(x) = {@rx@}\), we'll follow these steps:</p>
<p>1. Replace \(R(x)\) with \(y\): \( \displaystyle {@x_eq@}\).</p>
<p>2. Swap the roles of \(x\) and \(y\): \( \displaystyle {@y_eq@}\).</p>
<p>3. We solve the resulting equation for \(y\) to get \(R^{-1}(x)\).</p>
<p>\(\displaystyle {@y_eq@} \\
{@(denom(rhs(y_eq)))*x@}= {@num(rhs(y_eq))@} \\
{@expand((denom(rhs(y_eq)))*x)@}= {@num(rhs(y_eq))@} \\
 \displaystyle {@solutions[1]@}\)</p>
<p>Now, we have \(\displaystyle R^{-1}(x) = {@y_solved@}\).</p>
<p>Please note that for the inverse to be valid, the original function \(R(x)\) must be one-to-one (bijective), meaning that each value of \(x\) corresponds to a unique value of \(y\), and vice versa. </p><p>If the original function is not one-to-one, its inverse might not exist.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>aa:rand(3)+2; 
pm1:rand_with_prohib(-1,1,[0]); 
bb:pm1*(rand(15)+1); 
pm2:rand_with_prohib(-1,1,[0]); 
dd:pm2*(rand(15)+1); 
/* Define the function R(x) */ 
rx : (aa*x + bb)/(aa*x - dd); 
/*Repalce R(x) with y*/
x_eq: y = rx;  
/* Step 2: Swap x and y */ 
y_eq:x=subst(y,x,rhs(x_eq));
/* Solve for y */ 
solutions: solve(y_eq, y); 
y_solved: rhs(solutions[1]);
ta:y_solved;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Find the inverse of \(R(x)={@rx@}\).</p><p>\(R^{-1}(x)={@ta@}\) </p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>x</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p></p><p>When calculating the inverse of a function \(R(x)\) we are finding a new&nbsp;function that will tell us which \(x\) value will give as a particular&nbsp;\(y\) value when we input it in \(R\). Therefore, in the case of \(R(x)={@(aa*x+bb)/(cc*x+dd)@}\), we write it as \(y={@(aa*x+bb)/(cc*x+dd)@}\) and invert the variables \(x\) and \(y\) to get \(x={@(aa*y+bb)/(cc*y+dd)@}\) and rearrange as follows:</p><p>\[x={@(aa*y+bb)/(cc*y+dd)@}\]<br></p><p></p><p>\[x\cdot ({@cc*y+dd@})={@aa*y+bb@}\]<br></p><p></p><p>\[{@cc*y*x+dd*x@}={@aa*y+bb@}\]<br></p>\[{@cc*y*x-aa*y@}={@bb-dd*x@}\]<br>\[{@y*(cc*x-aa)@}={@bb-dd*x@}\]<br>\[{@y@}={@(bb-dd*x)/(cc*x-aa)@}\]<br>The inverse of \(R(x)\) is therefore \(R^{-1}(x) = {@(bb-dd*x)/(cc*x-aa)@}\).<br><p></p><p></p><p></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>549809041</deployedseed>
    <deployedseed>198522732</deployedseed>
    <deployedseed>469291720</deployedseed>
    <deployedseed>1143818981</deployedseed>
    <deployedseed>1596754657</deployedseed>
    <deployedseed>1106685697</deployedseed>
    <deployedseed>1180960958</deployedseed>
    <deployedseed>1734344533</deployedseed>
    <deployedseed>2135008834</deployedseed>
    <deployedseed>1077488481</deployedseed>
    <deployedseed>1920444725</deployedseed>
    <deployedseed>214634065</deployedseed>
    <deployedseed>909033721</deployedseed>
    <deployedseed>880475114</deployedseed>
    <deployedseed>437056888</deployedseed>
    <deployedseed>1936229858</deployedseed>
    <deployedseed>893711907</deployedseed>
    <deployedseed>768983537</deployedseed>
    <deployedseed>828786200</deployedseed>
    <deployedseed>1182303882</deployedseed>
    <deployedseed>1479019933</deployedseed>
    <deployedseed>153908564</deployedseed>
    <deployedseed>1447076419</deployedseed>
    <deployedseed>576823771</deployedseed>
    <deployedseed>142261119</deployedseed>
    <deployedseed>897892767</deployedseed>
    <deployedseed>430414058</deployedseed>
    <deployedseed>1188054616</deployedseed>
    <deployedseed>897901296</deployedseed>
    <deployedseed>1367653196</deployedseed>
    <deployedseed>2030085126</deployedseed>
    <deployedseed>858744092</deployedseed>
    <deployedseed>269040225</deployedseed>
    <deployedseed>585030876</deployedseed>
    <deployedseed>1497879970</deployedseed>
    <deployedseed>1259916931</deployedseed>
    <deployedseed>1434357526</deployedseed>
    <deployedseed>1849629733</deployedseed>
    <deployedseed>8375119</deployedseed>
    <deployedseed>462601020</deployedseed>
    <deployedseed>1387437718</deployedseed>
    <deployedseed>818196365</deployedseed>
    <deployedseed>1113920478</deployedseed>
    <deployedseed>2032243166</deployedseed>
    <deployedseed>475294874</deployedseed>
    <deployedseed>329348937</deployedseed>
    <deployedseed>1403960286</deployedseed>
    <deployedseed>1586191274</deployedseed>
    <deployedseed>405074036</deployedseed>
    <deployedseed>1801431905</deployedseed>
    <deployedseed>2059590169</deployedseed>
    <deployedseed>514062794</deployedseed>
    <deployedseed>193628788</deployedseed>
    <deployedseed>1383223979</deployedseed>
    <deployedseed>266442134</deployedseed>
    <deployedseed>18106957</deployedseed>
    <deployedseed>1192309024</deployedseed>
    <deployedseed>320796304</deployedseed>
    <deployedseed>51588127</deployedseed>
    <deployedseed>1185603524</deployedseed>
    <deployedseed>256803901</deployedseed>
    <deployedseed>1817582188</deployedseed>
    <deployedseed>1285602656</deployedseed>
    <deployedseed>1407483754</deployedseed>
    <deployedseed>2095501561</deployedseed>
    <deployedseed>1400566567</deployedseed>
    <deployedseed>33141834</deployedseed>
    <deployedseed>298015277</deployedseed>
    <deployedseed>1060265939</deployedseed>
    <deployedseed>1027382580</deployedseed>
    <deployedseed>1410835163</deployedseed>
    <deployedseed>396343047</deployedseed>
    <deployedseed>1967345488</deployedseed>
    <deployedseed>380960472</deployedseed>
    <deployedseed>1554778056</deployedseed>
    <deployedseed>475604931</deployedseed>
    <deployedseed>1023059219</deployedseed>
    <deployedseed>1498671086</deployedseed>
    <deployedseed>834920156</deployedseed>
    <deployedseed>269393928</deployedseed>
    <deployedseed>1742238468</deployedseed>
    <deployedseed>354340678</deployedseed>
    <deployedseed>705416118</deployedseed>
    <deployedseed>764017622</deployedseed>
    <deployedseed>848711856</deployedseed>
    <deployedseed>1611020963</deployedseed>
    <deployedseed>2141713536</deployedseed>
    <deployedseed>1167660087</deployedseed>
    <deployedseed>601683694</deployedseed>
    <deployedseed>203941690</deployedseed>
    <deployedseed>1564621038</deployedseed>
    <deployedseed>453700080</deployedseed>
    <deployedseed>1817622678</deployedseed>
    <deployedseed>1737066711</deployedseed>
    <deployedseed>1440981451</deployedseed>
    <deployedseed>1824040523</deployedseed>
    <deployedseed>117069603</deployedseed>
    <deployedseed>599388066</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
