<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Calculating the (Identity_matrix -matrix )^2 of a given matrix</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Let \( A= {@A@} \)</p>
<p>Calculate \( (I_2-A)^2 \)</p>

<table style="margin-left: 30px">
  <tr>
    <td>\( (I_2-A)^2 = \)</td>
    <td>&nbsp; [[input:ans1]] [[validation:ans1]]</td>
  </tr>
</table>
<p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked solution</h3>
<p>First, calculate the difference,</p>
<p style="margin-left: 30px">\(\begin{align*} 
I_2-A &= {@I@}-{@A@} \\
&= {@ta0@} 
\end{align*}\)</p>
<p>Then, we square the resulting matrix,</p>
<p style="margin-left: 30px">\( (I_2-A)^2 = {@ta@} \)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>A:rand(matrix([5,5],[5,5]))
I:matrix([1,0],[0,1])
ta0:I-A
ta:ta0.ta0</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>matrix \(A={@A@}\)</p>
<p>\((I_2-A)^2 = {@ta@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>matrix</type>
      <tans>ta</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>5</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>EqualComAss</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>3055763</deployedseed>
    <deployedseed>1075223874</deployedseed>
    <deployedseed>238130655</deployedseed>
    <deployedseed>484076339</deployedseed>
    <deployedseed>52888154</deployedseed>
    <deployedseed>216576934</deployedseed>
    <deployedseed>1094537557</deployedseed>
    <deployedseed>1822529685</deployedseed>
    <deployedseed>880904605</deployedseed>
    <deployedseed>254438331</deployedseed>
    <deployedseed>1032408186</deployedseed>
    <deployedseed>333673187</deployedseed>
    <deployedseed>1545686185</deployedseed>
    <deployedseed>1988192402</deployedseed>
    <deployedseed>849343655</deployedseed>
    <deployedseed>347907876</deployedseed>
    <deployedseed>1484945132</deployedseed>
    <deployedseed>1514634100</deployedseed>
    <deployedseed>1416682007</deployedseed>
    <deployedseed>238561487</deployedseed>
    <deployedseed>591303286</deployedseed>
    <deployedseed>812587934</deployedseed>
    <deployedseed>364840286</deployedseed>
    <deployedseed>1890111539</deployedseed>
    <deployedseed>1470443663</deployedseed>
    <deployedseed>1860196435</deployedseed>
    <deployedseed>1582848538</deployedseed>
    <deployedseed>929074343</deployedseed>
    <deployedseed>1088618988</deployedseed>
    <deployedseed>186947685</deployedseed>
    <deployedseed>2145483891</deployedseed>
    <deployedseed>681436516</deployedseed>
    <deployedseed>2092808486</deployedseed>
    <deployedseed>2672519</deployedseed>
    <deployedseed>1085491637</deployedseed>
    <deployedseed>1735468543</deployedseed>
    <deployedseed>1578591261</deployedseed>
    <deployedseed>840193156</deployedseed>
    <deployedseed>1428017679</deployedseed>
    <deployedseed>938048709</deployedseed>
    <deployedseed>101961336</deployedseed>
    <deployedseed>1449565867</deployedseed>
    <deployedseed>281835055</deployedseed>
    <deployedseed>1523963837</deployedseed>
    <deployedseed>1552278988</deployedseed>
    <deployedseed>1204592678</deployedseed>
    <deployedseed>254989771</deployedseed>
    <deployedseed>510463257</deployedseed>
    <deployedseed>1436267295</deployedseed>
    <deployedseed>1796494240</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
