<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Matrix Multiplication</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Let \({\bf A}\) and \({\bf B}\) be the following matrices: \[{\bf A} = {@matrixa@}\] <br>\[{\bf B} = {@matrixb@}\] Find the matrix multiplication \({\bf A} {\bf B}\): [[input:ans1]] [[validation:ans1]][[feedback:Result]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3><b>Worked Solution</b></h3>
<p style="margin-left: 30px">\({\bf A} = {@matrixa@}\) </p>
<p style="margin-left: 30px"> \({\bf B} = {@matrixb@}\) </p>
<p style="margin-left: 30px">\( {\bf A} {\bf B}={@C@}\) </p>
<p style="margin-left: 30px">\(\qquad= {@D@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>2</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[matrixa:matrix([rand_with_step(-5,5,1),rand_with_step(-5,5,1)],[rand_with_step(-5,5,1),rand_with_step(-5,5,1)]);
matrixa:ev(matrixa,simp);
 matrixb:matrix([rand_with_step(-5,5,1),rand_with_step(-5,5,1),rand_with_step(-5,5,1)],[rand_with_step(-5,5,1),rand_with_step(-5,5,1),rand_with_step(-5,5,1)]);
matrixb:ev(matrixb,simp);
correctanswer:ev(matrixa . matrixb,simp);
A:matrixa;
BT:transpose(matrixb);
C:zeromatrix (first(matrix_size(A)), first(matrix_size(BT)));
S:for a:1 thru first(matrix_size(A)) do for b:1 thru first(matrix_size(BT)) do C[ev(a,simp),ev(b,simp)]:apply("+",zip_with("*",A[ev(a,simp)],BT[ev(b,simp)]));
D:ev(C,simp);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\({\bf A} = {@matrixa@}, {\bf B} = {@matrixb@}\)</p>
<p>The matrix product \({\bf A} {\bf B}\):\({@correctanswer@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>0</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>matrix</type>
      <tans>correctanswer</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>determinant</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>correctanswer</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty>0</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty>0</falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1863165057</deployedseed>
    <deployedseed>1199226619</deployedseed>
    <deployedseed>1251369678</deployedseed>
    <deployedseed>1780500090</deployedseed>
    <deployedseed>164481967</deployedseed>
    <deployedseed>418637178</deployedseed>
    <deployedseed>384739331</deployedseed>
    <deployedseed>770904891</deployedseed>
    <deployedseed>754668842</deployedseed>
    <deployedseed>1685964336</deployedseed>
    <deployedseed>1029126614</deployedseed>
    <deployedseed>2032869379</deployedseed>
    <deployedseed>952327111</deployedseed>
    <deployedseed>631356423</deployedseed>
    <deployedseed>1580380660</deployedseed>
    <deployedseed>301248469</deployedseed>
    <deployedseed>307590759</deployedseed>
    <deployedseed>1371323179</deployedseed>
    <deployedseed>2081048832</deployedseed>
    <deployedseed>1245313824</deployedseed>
    <deployedseed>1327498914</deployedseed>
    <deployedseed>1531701411</deployedseed>
    <deployedseed>281411885</deployedseed>
    <deployedseed>344053870</deployedseed>
    <deployedseed>432082406</deployedseed>
    <deployedseed>991653335</deployedseed>
    <deployedseed>1124678713</deployedseed>
    <deployedseed>914175893</deployedseed>
    <deployedseed>2104626773</deployedseed>
    <deployedseed>994272912</deployedseed>
    <deployedseed>1380705517</deployedseed>
    <deployedseed>645864520</deployedseed>
    <deployedseed>2010536458</deployedseed>
    <deployedseed>1216966490</deployedseed>
    <deployedseed>1624336089</deployedseed>
    <deployedseed>172001610</deployedseed>
    <deployedseed>117761905</deployedseed>
    <deployedseed>1868693057</deployedseed>
    <deployedseed>868148979</deployedseed>
    <deployedseed>749642472</deployedseed>
    <deployedseed>1988834272</deployedseed>
    <deployedseed>1509894914</deployedseed>
    <deployedseed>270189668</deployedseed>
    <deployedseed>1595688681</deployedseed>
    <deployedseed>1211818789</deployedseed>
    <deployedseed>377480868</deployedseed>
    <deployedseed>485843232</deployedseed>
    <deployedseed>459621725</deployedseed>
    <deployedseed>415025577</deployedseed>
    <deployedseed>28929510</deployedseed>
    <deployedseed>987787472</deployedseed>
    <deployedseed>1674914365</deployedseed>
    <deployedseed>1776947871</deployedseed>
    <deployedseed>1787884005</deployedseed>
    <deployedseed>440531999</deployedseed>
    <deployedseed>1863868661</deployedseed>
    <deployedseed>630865637</deployedseed>
    <deployedseed>1209180347</deployedseed>
    <deployedseed>2114390343</deployedseed>
    <deployedseed>583080148</deployedseed>
    <deployedseed>1442207498</deployedseed>
    <deployedseed>1177166133</deployedseed>
    <deployedseed>1005332643</deployedseed>
    <deployedseed>383849065</deployedseed>
    <deployedseed>809327464</deployedseed>
    <deployedseed>907510524</deployedseed>
    <deployedseed>2127669311</deployedseed>
    <deployedseed>1123956575</deployedseed>
    <deployedseed>178939565</deployedseed>
    <deployedseed>2063059198</deployedseed>
    <deployedseed>956395392</deployedseed>
    <deployedseed>237607145</deployedseed>
    <deployedseed>2045673218</deployedseed>
    <deployedseed>1202822049</deployedseed>
    <deployedseed>660652034</deployedseed>
    <deployedseed>662658250</deployedseed>
    <deployedseed>2102560182</deployedseed>
    <deployedseed>881658243</deployedseed>
    <deployedseed>125201840</deployedseed>
    <deployedseed>149002760</deployedseed>
    <deployedseed>493898269</deployedseed>
    <deployedseed>1262879292</deployedseed>
    <deployedseed>3362450</deployedseed>
    <deployedseed>1517739240</deployedseed>
    <deployedseed>1293457423</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>correctanswer</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
