<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Area Between Curves</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Determine the area of the region bounded by the following equations,</p>
<p style="margin-left:30px">\( \displaystyle {@eq1@}, \hspace{10px} {@eq2@}, \hspace{10px} x= \displaystyle{@x_1@}, \hspace{10px} x={@x_2@}\)</p>

<br>

<p>[[input:ans1]] [[validation:ans1]]</p>
<p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>Here is a sketch of the bounded region we want to find the area of,</p>

[[jsxgraph]]

 const board = JXG.JSXGraph.initBoard(divid, {
            boundingbox: [-5, 40, 5, -3], axis: true,showNavigation: false,  showCopyright: false
        });

        // Plot functions only up to the mutual intersections
        var function1 = board.create('functiongraph', ['{#eq_2#}', {#bb#}, {#aa#}], {name: '{#eq2#}', fontSize: 20, withLabel:true, fixed:true});
        var function2 = board.create('functiongraph', ['{#eq_1#}', {#aa#}, {#bb#}], {name: '{#eq1#}',  fontSize: 20, withLabel:true, fixed:true});
        var function3 = board.create('functiongraph', ['{#eq_1#}', {#cc#}, {#aa#}]);
        var function4 = board.create('functiongraph', ['{#eq_2#}',  {#aa#}, {#cc#}]);
       
        // properties is a dictionary of properties for the generated function.
        function make_filled_curve(board, element_list, properties) {
            var curve = board.create('curve', [[], []], properties);
            curve.updateDataArray = function () {
                this.dataX = [];
                this.dataY = [];
                for (const element of element_list) {
                    if (Array.isArray(element)) {
                        // Add the point to the curve
                        this.dataX.push(element[0]);
                        this.dataY.push(element[1]);
                    } else {
                        // Copy all points from function
                        this.dataX = this.dataX.concat(element.points.map(p => p.usrCoords[1]));
                        this.dataY = this.dataY.concat(element.points.map(p => p.usrCoords[2]));
                    }
                }
                // Close the loop
                this.dataX.push(this.dataX[0]);
                this.dataY.push(this.dataY[0]);
            };
            board.update();
        }

        // Plot different filled functions
        make_filled_curve(board, [function1, function2], { strokeWidth: 3, fillColor: 'green', fillOpacity: 0.3 });
        make_filled_curve(board, [function3, function4], { strokeWidth: 3, fillColor: 'green', fillOpacity: 0.3 });

[[/jsxgraph]]

<br>

<p>First we note that when we graph the region, the curves cross each other. Therefore, we need to find the intersection point directly by setting the two equations equal and solving them.</p>
<p>\( \displaystyle  {@expr@}\) , we get</p>
<p>\(\displaystyle  {@(x+const)*func_2@} =1\)</p>
<p>\(\displaystyle  x+{@const@}=\sqrt[3]{1}\)</p>
<p>\( \displaystyle  {@value_x@}\)</p>
<p>Hence the graphs of the functions intersect at \(\displaystyle  {@value_x@}\).</p>
<p>For \( \displaystyle {@x_1@} \leq x \leq {@value_x_1@}, {@func_1@} \geq {@func_2@}\), so</p>
<p>\(\displaystyle  |{@func_1@}-{@func_2@}|={@func_1@}-{@func_2@}\)</p>
<p>On the other hand, for \( \displaystyle {@value_x_1@} \leq x \leq {@x_2@}, {@func_2@} \geq {@func_1@}\), so</p>
<p>\(\displaystyle  |{@func_1@}-{@func_2@}|={@func_2@}-{@func_1@}\)</p>
<p>Then</p>
<p>\( \begin{align} 
  A &amp;= \displaystyle \int_{{@x_1@}}^{{@x_2@}}|{@func_1@}-{@func_2@}|dx \\ 
  A &amp;= \displaystyle \int_{{@x_1@}}^{{@value_x_1@}} \left( {@func_a@} \right) dx+\int_{{@value_x_1@}}^{{@x_2@}} \left( {@func_b@} \right) dx \\ &amp;= \displaystyle \left[{@int_func_a1@}\right]_{@x_1@}^{@value_x_1@}+\left[{@int_func_b1@}\right]_{@value_x_1@}^{@x_2@} \\ 
&amp;= \displaystyle{@int_func_a@}+{@int_func_b@} \\
&amp;= \displaystyle {@ta@} \\
\end{align} \)</p>
<p>
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*Defining the question variables*/ 
const:rand(3)+1;/*constant*/
eq1:y=1/(x+const);/*first equation*/
eq_1: 1/(x+const);/*first equation*/
eq2:y=(x+const)^2;/*second equation*/
eq_2: (x+const)^2;/*second equation*/
expr:rhs(eq1)=rhs(eq2);
expr1:solve(expr,x);
value_x:expr1[3];
value_x_1:rhs(value_x);
num_x_1:rand(3)+3;
den_x_1:rand(2)+2;
x_1: -num_x_1/den_x_1;/*minimum value of x*/
x_2:rand(3)+1;/*maximum value of x*/
y_1:rhs(ev(eq1,y,x=x_1));
y_2:rhs(ev(eq2,y,x=x_2));
func_1:rhs(eq1);
func_2:rhs(eq2);
func_a:func_1-func_2;
int_func_a1:int(func_a,x);
int_func_a:int(func_a,x,rhs(x_1),rhs(value_x));
func_b:func_2-func_1;
int_func_b1:int(func_b,x);
int_func_b:int(func_b,x,rhs(value_x),rhs(x_2));
Area:int_func_a+int_func_b;
/*model answer*/
ta: Area;

/*worked out solution*/
graph:plot([rhs(eq1),rhs(eq2)],[x,x_1,x_2]);


/*Data about jsxgraph*/
aa: value_x_1;
bb: x_1;
cc: x_2;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p><p>The area of the region bounded by \({@eq1@}, \qquad {@eq2@}, \, x={@x_1@} ,  x={@x_2@}\) is:     \( \qquad {@ta@} \)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1 checks for correct answer*/
/*Node 2 mistake:Integrating the indefinite integral without the limits*/
mistake_2:int(area,x);
/*Node 3 mistake:Interchanging the limits of integration, however, with the correct definite integral.*/
mistake_3:int(area,x,x_2,x_1);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>367546284</deployedseed>
    <deployedseed>1985210531</deployedseed>
    <deployedseed>448463015</deployedseed>
    <deployedseed>1946822610</deployedseed>
    <deployedseed>1555566116</deployedseed>
    <deployedseed>3779471</deployedseed>
    <deployedseed>1905168569</deployedseed>
    <deployedseed>1423531415</deployedseed>
    <deployedseed>509654270</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
