<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Basic integration by substitution</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Integrate the following function by substitution, \(f(x)= \displaystyle {@f_x@}\).</p>
<p>\( \displaystyle \int f(x) dx=\)[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p><b>Worked solution</b></p>
<p>Let's rewrite the integral using \(u\) as the variable of substitution.</p>
<p> We observe that the factor \( {@x^(n-1)@}\) within the integrand is a multiple of the derivative inside the power of the second one: \({@x^(n-1)@} = k \cdot ({@a*x^n+b@})^{\prime} \). this means we can find this integral by substitution as follows.</p>
<p>Let \(u={@a*x^n+b@}\)</p>
<p> \(u={@a*x^n+b@} \implies  \displaystyle \frac{du}{dx} = \frac{d}{dx}({@a*x^n+b@}) = {@diff(a*x^n+b,x)@}\).</p>
<p> So  \(\displaystyle \frac{1}{@n*a@} du = \displaystyle   {@(1/(n*a))*diff(a*x^n+b,x)@} dx\) </p>
<p> \( \int {@f_x@} dx = \displaystyle \int \frac{1}{@n*a@}  u^{@m@} du =  \displaystyle \frac{1}{@n*a@}  \left( \frac{@u^(m+1)@}{@m+1@} \right) + c \)</p>
<p>Now, we replace the expresion corresponding to \(u\), and we get</p>
<p> \( \int {@f_x@} dx = \displaystyle  {@ta@}\), \(c\) is a constant of integration.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>n: rand_with_step(3,6,1);
m: rand_with_step(5,9,1);
a: rand_with_prohib(-5,5,[0]);
b: rand_with_prohib(-5,5,[0]);
f_x: x^(n-1)*(a*x^n+b)^(m);
ta: integrate(f_x,x)+c;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\( \displaystyle \int {@f_x@} dx={@ta@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>11</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>Int</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>x</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2037223012</deployedseed>
    <deployedseed>984904260</deployedseed>
    <deployedseed>1941682119</deployedseed>
    <deployedseed>1566941596</deployedseed>
    <deployedseed>1538184357</deployedseed>
    <deployedseed>1684265238</deployedseed>
    <deployedseed>1766764637</deployedseed>
    <deployedseed>877729978</deployedseed>
    <deployedseed>309032618</deployedseed>
    <deployedseed>449258618</deployedseed>
    <deployedseed>868615998</deployedseed>
    <deployedseed>1862575354</deployedseed>
    <deployedseed>1166775218</deployedseed>
    <deployedseed>2020381073</deployedseed>
    <deployedseed>1744314182</deployedseed>
    <deployedseed>1586013124</deployedseed>
    <deployedseed>1434014395</deployedseed>
    <deployedseed>1585408624</deployedseed>
    <deployedseed>597732565</deployedseed>
    <deployedseed>128144792</deployedseed>
    <deployedseed>700383446</deployedseed>
    <deployedseed>707812663</deployedseed>
    <deployedseed>977919136</deployedseed>
    <deployedseed>146618028</deployedseed>
    <deployedseed>1362759986</deployedseed>
    <deployedseed>827735971</deployedseed>
    <deployedseed>1804339835</deployedseed>
    <deployedseed>1294583378</deployedseed>
    <deployedseed>1636329830</deployedseed>
    <deployedseed>1616665419</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
