<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Integration by parts cyclic</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Find \(\displaystyle \int {@quest_expr@}\,dx\)</p>
<p></p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p>Recall,</p>
<p>Integration by parts is given by the formula, \(\displaystyle \int udv=uv-\int vdu \)</p>
<p>Let \(u={@u@}\) and \(dv={@dv@}\,dx\)</p>
<p>Then, \(du={@du@}\,dx\) and \(v={@v@}\).</p>
<p>\(\displaystyle \int {@quest_expr@}\,dx={@uv_1@}-\displaystyle \int {@vdu_1@}\,dx \,\,\,\,(1)\)</p>
<p>Hence applying integration by parts for the second time to \((1)\)</p>
<p>\(u={@u_2@} \Rightarrow  du={@du_2@}\,dx\) and \(dv={@dv_2@}\,dx\), then  \(v={@v_2@}\)</p>
<p>\(\displaystyle \int {@u_2@}{@dv_2@}\,dx={@u2_v2@}-\int {@v_du_2@}\,dx\)</p>
<p>But \(\displaystyle \int {@quest_expr@}\,dx={@uv_1@}-\displaystyle \int {@vdu_1@}\,dx={@uv_1@}+{@u2_v2@}-\int {@v_du_2@}\,dx\)</p>
<p>By collecting like terms, you obtain</p>
<p>\( \displaystyle \int  {@quest_expr@}\,dx+\int {@v_du_2@}\,dx={@u2_v2@}+{@uv_1@}\)</p>
<p>\(\displaystyle \int  {@quest_expr@}\,dx= {@ta1@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*integration by parts cyclic*/
trig:rand([sin,cos]);
power: rand(3)+1;
quest_expr:%e^(power*x)*trig(x);
/*worked solution variables*/
/*choosing the values of u and dv*/
u:%e^(power*x);
du:diff(u,x);
dv:trig(x);
v:int(dv,x);
uv_1:u*v;
vdu_1:v*du;
/*applying integration by parts for the second time*/
u_2:diff(u,x);
du_2:diff(du,x);
v_2:diff(v,x);
dv_2:int(dv,x);
u2_v2:u_2*v_2;
v_du_2:v_2*du_2;
/*model ansswer*/
ta1:int(quest_expr,x)+c;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\( \displaystyle \int {@quest_expr@}\,dx={@ta1@} \)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */
/*Node 2: Forgetting the constant of integration*/
ta1_cste:int(quest_expr,x);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>Int</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>x</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1_cste</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You answer is partially correct because you forgot the constant of integration</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1869899534</deployedseed>
    <deployedseed>774224863</deployedseed>
    <deployedseed>1378516124</deployedseed>
    <deployedseed>267086876</deployedseed>
    <deployedseed>169000529</deployedseed>
    <deployedseed>170829202</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
