<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the submatrices of a matrix</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>For the matrix \(A={@matrix_A@}\), give any submatrix of \( 2 \times 2 \).</p>
<p>Submatrix \( =\) [[input:ans1]] [[validation:ans1]]</p>
<p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p> Considering the matrix \(A={@matrix_A@}\), one of the \( 2 \times 2 \) submatrix is \({@sol@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>aa:rand(3)+1;
bb:rand_with_prohib(4,8,[aa]);
cc:rand_with_prohib(-5,5,[aa,bb,0]);
dd:rand_with_step(1,2,8);
ee:rand([2,4,6]);
ff:-1*(rand_with_prohib(-5, 5, [0]));
matrix_A:matrix([cc,aa,bb],[aa,bb,dd],[ff,dd,ee]);
submatrix_1:submatrix(3,matrix_A,3);
submatrix_2:submatrix(1,matrix_A,3);
submatrix_3:submatrix(3,matrix_A,1);
submatrix_4:submatrix(1,matrix_A,2);
submatrix_5:submatrix(3,matrix_A,2);
submatrix_6:submatrix(1,matrix_A,1);
submatrix_7:submatrix(2,matrix_A,1);
submatrix_8:submatrix(2,matrix_A,2);
submatrix_9:submatrix(2,matrix_A,3);
sol: rand([submatrix_1,submatrix_2,submatrix_3,submatrix_4,submatrix_5,submatrix_6,submatrix_7,submatrix_8,submatrix_9]);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Matrix: \({@matrix_A@}\)</p>
<p>Submatrix \(={@sol@}\) </p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>matrix</type>
      <tans> submatrix_8</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>ta:[submatrix_1,submatrix_2,submatrix_3,submatrix_4,submatrix_5,submatrix_6,submatrix_7,submatrix_8,submatrix_9];
cor_input: member(ans1, ta);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>cor_input</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You got the correct submatrix!!</p>]]></text>
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2100645992</deployedseed>
    <deployedseed>1246305319</deployedseed>
    <deployedseed>1529376489</deployedseed>
    <deployedseed>144486622</deployedseed>
    <deployedseed>773732602</deployedseed>
    <deployedseed>600244292</deployedseed>
    <deployedseed>397764289</deployedseed>
    <deployedseed>10784279</deployedseed>
    <deployedseed>1027335114</deployedseed>
    <deployedseed>1913014814</deployedseed>
    <deployedseed>763921289</deployedseed>
    <deployedseed>1335264890</deployedseed>
    <deployedseed>1577499692</deployedseed>
    <deployedseed>1936485448</deployedseed>
    <deployedseed>606924893</deployedseed>
    <deployedseed>1668527794</deployedseed>
    <deployedseed>2003053831</deployedseed>
    <deployedseed>1168655740</deployedseed>
    <deployedseed>1327567360</deployedseed>
    <deployedseed>562589806</deployedseed>
    <deployedseed>2063944002</deployedseed>
    <deployedseed>269910561</deployedseed>
    <deployedseed>1567788588</deployedseed>
    <deployedseed>631657823</deployedseed>
    <deployedseed>1223680300</deployedseed>
    <deployedseed>556033135</deployedseed>
    <deployedseed>2114558087</deployedseed>
    <deployedseed>617701467</deployedseed>
    <deployedseed>1378307847</deployedseed>
    <deployedseed>1389420491</deployedseed>
    <deployedseed>1167371556</deployedseed>
    <deployedseed>2039388606</deployedseed>
    <deployedseed>2066343742</deployedseed>
    <deployedseed>1595186024</deployedseed>
    <deployedseed>1191466797</deployedseed>
    <deployedseed>1933959186</deployedseed>
    <deployedseed>1021105994</deployedseed>
    <deployedseed>991471921</deployedseed>
    <deployedseed>1982096998</deployedseed>
    <deployedseed>1549548967</deployedseed>
    <deployedseed>124695393</deployedseed>
    <deployedseed>1726993018</deployedseed>
    <deployedseed>904805620</deployedseed>
    <deployedseed>90043941</deployedseed>
    <deployedseed>1727096611</deployedseed>
    <deployedseed>2042413991</deployedseed>
    <deployedseed>856970110</deployedseed>
    <deployedseed>726658832</deployedseed>
    <deployedseed>686101530</deployedseed>
    <deployedseed>1661072043</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>submatrix_1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
