<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the the distance between point A and point B on a number line.</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>On the number line below, determine the distance between point \(A\) and point \(B\)</p>
<p>[[jsxgraph height='100px']]
JXG.Options.text.useMathJax=true;

var brd = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [-{#b-2#}, {#1/a#}, {#b+2#}, -{#1/a#}],
    showCopyright:false,
keepaspectratio:true, 
    grid:true,
    showNavigation:false});
  var axisx = brd.create('axis', [[0, 0], [1, 0]], {
      firstArrow: true,
      lastArrow: true,
      ticks: {
        drawZero: true,
        ticksDistance: 1,
        majorHeight: 10,
        tickEndings: [1, 1],
        minorTicks: 0
      }
    });
    
    //var p = brd.create('glider', [2.2, 0, axisx], {});
    // Points
    var p1 = brd.create('point', [{#a#}, 0], {name: 'A', size: 4, strokeColor: 'blue', fixed:true});
    var p2 = brd.create('point', [{#b#}, 0], {name: 'B', size: 4, strokeColor: 'black', fixed:true});
    
    // Line segment
    var line = brd.create('segment', [p1, p2], {strokeColor: 'red', strokeWidth: 2,fixed:true});
    
    // Distance label
    //var distance = brd.create('text', [(p1.X() + p2.X()) / 2, -0.15, () => `Distance: ${Math.abs(p1.X() - p2.X())}`], {anchorX: 'middle', anchorY: 'top',fixed:true});


    brd.update();

[[/jsxgraph]]</p>
<p>Distance between \(A\) and \(B\): [[input:ans1]] [[validation:ans1]] units [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>The distance between two points \(A\) and \(B\) on a number line can be calculated using the absolute value of the difference between their coordinates. </p>
<p>In this case, point \(A\) is \({@a@}\) and point \(B\) is {@b@}. </p>
<p>Distance  \(=|B - A|\)</p>
<p>Distance  \(=|{@b@} - {@a@}|\)</p>
<p>Distance \(=|{@ta@}|\)</p>
<p>Distance \(= {@ta@}\) units.</p>
<p>Therefore, the distance between \(A\) and \(B\) on the number line is {@ta@} units as shown in the real number line below:</p>
<p>[[jsxgraph height='100px']]
JXG.Options.text.useMathJax=true;

var brd = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [-{#b-2#}, {#1/a#}, {#b+2#}, -{#1/a#}],
    showCopyright:false,
keepaspectratio:true, 
    grid:true,
    showNavigation:false});
  var axisx = brd.create('axis', [[0, 0], [1, 0]], {
      firstArrow: true,
      lastArrow: true,
      ticks: {
        drawZero: true,
        ticksDistance: 1,
        majorHeight: 10,
        tickEndings: [1, 1],
        minorTicks: 0
      }
    });
    
    //var p = brd.create('glider', [2.2, 0, axisx], {});
    // Points
    var p1 = brd.create('point', [{#a#}, 0], {name: 'A={#a#}', size: 4, strokeColor: 'blue',fixed:true});
    var p2 = brd.create('point', [{#b#}, 0], {name: 'B={#b#}', size: 4, strokeColor: 'black',fixed:true});
    
    // Line segment
    var line = brd.create('segment', [p1, p2], {strokeColor: 'red', strokeWidth: 2,fixed:true});
    
    // Distance label
    var distance = brd.create('text', [(p1.X() + p2.X()) / 2, -0.15, () => `Distance: ${Math.abs(p1.X() - p2.X())}`], {anchorX: 'bottom', anchorY: 'top',fixed:true});


    brd.update();

[[/jsxgraph]]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>a:rand(7)+1;
b:a+(rand(8)+1);
ta:b-a;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>On the number line below, determine the distance between point \(A={@a@}\) and point \(B={@b@}\) \(AB={@ta@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>369312014</deployedseed>
    <deployedseed>1781965065</deployedseed>
    <deployedseed>956142585</deployedseed>
    <deployedseed>834287738</deployedseed>
    <deployedseed>965469664</deployedseed>
    <deployedseed>320456048</deployedseed>
    <deployedseed>884889911</deployedseed>
    <deployedseed>601411994</deployedseed>
    <deployedseed>1727952518</deployedseed>
    <deployedseed>2006042252</deployedseed>
    <deployedseed>1530074264</deployedseed>
    <deployedseed>837758585</deployedseed>
    <deployedseed>908856458</deployedseed>
    <deployedseed>1534205444</deployedseed>
    <deployedseed>1005543969</deployedseed>
    <deployedseed>577054275</deployedseed>
    <deployedseed>1554589290</deployedseed>
    <deployedseed>940574324</deployedseed>
    <deployedseed>1997351003</deployedseed>
    <deployedseed>1550297098</deployedseed>
    <deployedseed>447609531</deployedseed>
    <deployedseed>1734431930</deployedseed>
    <deployedseed>962273940</deployedseed>
    <deployedseed>4514625</deployedseed>
    <deployedseed>1429247733</deployedseed>
    <deployedseed>905805803</deployedseed>
    <deployedseed>1945989984</deployedseed>
    <deployedseed>1019098913</deployedseed>
    <deployedseed>21258477</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
