<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding Focus and Directrix of a Parabola (Direct)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find the focus and directrix of the parabola, \({@parabola@}\)</p>

<br>

<p style="margin-left: 30px"> Focus: &emsp; \( F \big( \) [[input:ans2]] [[validation:ans2]] \(,\) [[input:ans3]] [[validation:ans3]] \( \big) \)</p>
<p>[[feedback:prtA]]</p>

<p style="margin-left: 30px">Directrix: &emsp; [[input:ans4]] [[validation:ans4]]</p>
<p>[[feedback:prt4]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked solution</h3>

<p>We want to use the fact that a general parabola in the form, \((x-x_{0})^{2} = 4 p (y - y_{0}) \) has focus \(F(x_{0}, y_{0}+p)\) and directrix \(y = y_{0}-p\)</p>

<br>

<p>Firstly, we should rewrite the given parabola in the same form.
<br>To do this, we should "complete the square" with the \(x\) terms and then collect the \(y\) term and constants on the right hand side.</p>
<p style="margin-left: 30px">\(\begin{align*}
{@parabola@} 
&\Rightarrow ({@expr1@}) + ({@expr2@})= 0 \\ 
&\Rightarrow {@(x-aa)^2@} = {@ev(-expr2,simp)@} \\ 
&\Rightarrow {@(x-aa)^2@} = {@4*pp@}({@y-cc@}) 
\end{align*}\)</p>

<br>

<p>Now, relating this parabola to the general form given, we have that,</p>
<p style="margin-left: 30px">\(x_0 = {@aa@}, &emsp; y_0 = {@cc@} &emsp; p = \frac{{@4*pp@}}{4} = {@pp@}\)</p>

<p>Therefore, we can use the general formulae given to find the focus and directrix,</p>
<p style="margin-left: 30px">Focus: &emsp; \( F(x_{0}, y_{0}+p) = ({@aa@},{@cc+pp@}) \)</p>
<p style="margin-left: 30px">Directrix: &emsp; \( y = y_{0}-p = {@cc-pp@} \)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*Defining parabola equation*/
aa: rand_with_prohib(-5,5,[0]); /*value of y0*/
pp: rand_with_prohib(-4,4,[0]); /*value of pp*/
cc: rand_with_prohib(-5,5,[0]); /*value of x0*/
parabola: expand((x-aa)^2-4*pp*(y-cc))=0;

/*Coordinates of the focus*/
ta2:  aa; /*abscissa*/
ta3: cc+pp; /*ordinate*/
/*Directrix*/
ta4: y=cc-pp;

/*Worked solution*/
expr1: expand((x-aa)^2);
expr2: expand((x-aa)^2-4*pp*(y-cc))-expr1;
ta1: (x-aa)^2=4*pp*(y-cc);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>The focus and directrix of the parabola \({@parabola@}\)</p>
<p>Focus: \(F ({@ta2@}, {@ta3@})\) 
<br>Directrix: \({@ta4@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node1: Check correct answer*/
/*Node 2: Mistake using variable y*/
mistake: x=cc-pp;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>mistake</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer is almost correct, you have determined the correct constant of the directrix equation.
<br>However, the directrix should be an equation for \(y\), whereas you have used the variable \(x\).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtA</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks if the first part of the answer is correct */
/*Node 2: Checks if the second part of the answer is correct */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>[ans2, ans3]</sans>
        <tans>[ta2, ta3]</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtA-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prtA-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>prtA-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prtA-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The \(x\)-coordinate of your answer is not correct.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtA-3-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtA-3-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The \(y\)-coordinate of your answer is not correct.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2008759323</deployedseed>
    <deployedseed>954704372</deployedseed>
    <deployedseed>1916745194</deployedseed>
    <deployedseed>202382068</deployedseed>
    <deployedseed>1848714297</deployedseed>
    <deployedseed>1133111847</deployedseed>
    <deployedseed>1587487074</deployedseed>
    <deployedseed>92479984</deployedseed>
    <deployedseed>353923391</deployedseed>
    <deployedseed>2109247063</deployedseed>
    <deployedseed>1967495686</deployedseed>
    <deployedseed>1466223837</deployedseed>
    <deployedseed>1313260495</deployedseed>
    <deployedseed>2064462699</deployedseed>
    <deployedseed>1894054891</deployedseed>
    <deployedseed>900840987</deployedseed>
    <deployedseed>1829769658</deployedseed>
    <deployedseed>2051102775</deployedseed>
    <deployedseed>2004061844</deployedseed>
    <deployedseed>1184611421</deployedseed>
    <deployedseed>567100986</deployedseed>
    <deployedseed>1273761075</deployedseed>
    <deployedseed>991718753</deployedseed>
    <deployedseed>1043897030</deployedseed>
    <deployedseed>1147251201</deployedseed>
    <deployedseed>896122386</deployedseed>
    <deployedseed>158016682</deployedseed>
    <deployedseed>787902869</deployedseed>
    <deployedseed>1271374877</deployedseed>
    <deployedseed>483436421</deployedseed>
    <deployedseed>1690305208</deployedseed>
    <deployedseed>1684440328</deployedseed>
    <deployedseed>2145986372</deployedseed>
    <deployedseed>1288310696</deployedseed>
    <deployedseed>404764114</deployedseed>
    <deployedseed>861675625</deployedseed>
    <deployedseed>1627013519</deployedseed>
    <deployedseed>442573609</deployedseed>
    <deployedseed>1069375012</deployedseed>
    <deployedseed>2025633801</deployedseed>
    <deployedseed>658614063</deployedseed>
    <deployedseed>1074238352</deployedseed>
    <deployedseed>269121983</deployedseed>
    <deployedseed>1491194802</deployedseed>
    <deployedseed>51083442</deployedseed>
    <deployedseed>1625895340</deployedseed>
    <deployedseed>639844614</deployedseed>
    <deployedseed>2140300235</deployedseed>
    <deployedseed>145740745</deployedseed>
    <deployedseed>779402079</deployedseed>
    <deployedseed>952102565</deployedseed>
    <deployedseed>327261328</deployedseed>
    <deployedseed>510384749</deployedseed>
    <deployedseed>1001364767</deployedseed>
    <deployedseed>1706331964</deployedseed>
    <deployedseed>20662089</deployedseed>
    <deployedseed>1656824248</deployedseed>
    <deployedseed>165909864</deployedseed>
    <deployedseed>364599485</deployedseed>
    <deployedseed>1716493704</deployedseed>
    <deployedseed>158126374</deployedseed>
    <deployedseed>79051166</deployedseed>
    <deployedseed>422701110</deployedseed>
    <deployedseed>1166837708</deployedseed>
    <deployedseed>1527960512</deployedseed>
    <deployedseed>516634763</deployedseed>
    <deployedseed>423304192</deployedseed>
    <deployedseed>876232220</deployedseed>
    <deployedseed>1770199682</deployedseed>
    <deployedseed>1138222566</deployedseed>
    <deployedseed>928114137</deployedseed>
    <deployedseed>1508284814</deployedseed>
    <deployedseed>1297059298</deployedseed>
    <deployedseed>975305201</deployedseed>
    <deployedseed>18640873</deployedseed>
    <deployedseed>1507006663</deployedseed>
    <deployedseed>1676373258</deployedseed>
    <deployedseed>840979616</deployedseed>
    <deployedseed>1949414418</deployedseed>
    <deployedseed>688188627</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prtA</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtA-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
