<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Checking definition of relation</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Consider the sets \(A={@A_set@}\) and \(B={@B_set@}.\)</p> 
<p>Which of the following are relations? Select all the relations from \(A\) to \(B\).</p>

<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p><b><h3>Worked Solution</h3></b></p>
<p>Recall: For a relation to be valid, each ordered pair must contain one element from the first set and one element from the second set.</p>
<p>In this case we have;</p>
<p>\({@option1@}\)</p>
<p>\({@option4@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*Generating the domain and codomain sets*/
a_1:rand_with_prohib(-4,8,[]);
a_2:rand_with_prohib(-4,8,[a_1]);
a_3:rand_with_prohib(-4,8,[a_1,a_2]);
a_4:rand_with_prohib(-4,8,[a_1,a_2,a_3]);
a_5:rand_with_prohib(-4,8,[a_1,a_2,a_3,a_4]);
A_set:{a_1,a_2,a_3,a_4,a_5};
b_1:rand_with_prohib(-4,8,[]);
b_2:rand_with_prohib(-4,8,[b_1]);
b_3:rand_with_prohib(-4,8,[b_1,b_2]);
b_4:rand_with_prohib(-4,8,[b_1,b_2,b_3]);
B_set:{b_1,b_2,b_3,b_4};

/*Generate R_1, which is a relation*/
A_list:listify(A_set);
r_1_x:A_list[rand(5)+1];
B_list:listify(B_set);
B_list:random_permutation(B_list);
r_1_y_1:B_list[1];
r_1_y_2:B_list[2];
r_1_y_3:B_list[3];
r_1_y_4:B_list[4];
option1:R_1={[r_1_x,r_1_y_1],[r_1_x,r_1_y_2]};
/* For the question text we will want: \(R_1 = \{ ({@r_1_x@},{@r_1_y_1@}),({@r_1_x@},{@r_1_y_2@}) \} \) */

/*Generate R_2, which is not a relation, as there is a y value not in the codomain*/
b_5:rand_with_prohib(-4,8,[b_1,b_2,b_3,b_4]);
r_1_x_1:A_list[1];
r_1_x_2:A_list[2];
r_1_x_3:A_list[3];
r_1_x_4:A_list[4];
r_1_x_5:A_list[5];
option2:R_2={[r_1_x_1,r_1_y_1],[r_1_x_1,b_5],[r_1_x,r_1_y_2]};
/*Generate R_3, which is not a relation*/
option3:R_3={[r_1_x_1,r_1_y_1],[r_1_x,r_1_y_2],[r_1_x_3,r_1_y_3],[r_1_x_4,r_1_x_5,r_1_y_4]};

/*Generate R_4, which is a relation*/
option4:R_4={[r_1_x_1,r_1_y_1],[r_1_x_1,r_1_y_2],[r_1_x_1,r_1_y_3],[r_1_x_1,r_1_y_4],[r_1_x_2,r_1_y_1],[r_1_x,r_1_y_2],[r_1_x,r_1_y_3],[r_1_x_5,r_1_y_4]};

ta: random_permutation([[option4,true],[option1,true],[option2,false],[option3,false]]);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Consider the sets \(A={@A_set@}\) and \(B={@B_set@}.\)</p> 
<p>The relations formed from \(A\) to \(B\) are \({@option1@}\) and \({@option4@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>checkbox</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(ta)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1272883968</deployedseed>
    <deployedseed>882908917</deployedseed>
    <deployedseed>1820739695</deployedseed>
    <deployedseed>501443303</deployedseed>
    <deployedseed>83279051</deployedseed>
    <deployedseed>1737582322</deployedseed>
    <deployedseed>2100029362</deployedseed>
    <deployedseed>1933204500</deployedseed>
    <deployedseed>750511226</deployedseed>
    <deployedseed>197377306</deployedseed>
    <deployedseed>1870988655</deployedseed>
    <deployedseed>1498055117</deployedseed>
    <deployedseed>1678131693</deployedseed>
    <deployedseed>888231635</deployedseed>
    <deployedseed>594700035</deployedseed>
    <deployedseed>170013083</deployedseed>
    <deployedseed>241540374</deployedseed>
    <deployedseed>705670471</deployedseed>
    <deployedseed>849216777</deployedseed>
    <deployedseed>640799810</deployedseed>
    <deployedseed>386463944</deployedseed>
    <deployedseed>375773118</deployedseed>
    <deployedseed>87908837</deployedseed>
    <deployedseed>1584374602</deployedseed>
    <deployedseed>126841264</deployedseed>
    <deployedseed>504136767</deployedseed>
    <deployedseed>371017126</deployedseed>
    <deployedseed>190648254</deployedseed>
    <deployedseed>1573623254</deployedseed>
    <deployedseed>1026266187</deployedseed>
    <deployedseed>373916344</deployedseed>
    <deployedseed>1966269715</deployedseed>
    <deployedseed>704451322</deployedseed>
    <deployedseed>62774226</deployedseed>
    <deployedseed>227698103</deployedseed>
    <deployedseed>345370911</deployedseed>
    <deployedseed>1999085645</deployedseed>
    <deployedseed>1807014946</deployedseed>
    <deployedseed>685994429</deployedseed>
    <deployedseed>1917214469</deployedseed>
    <deployedseed>951212315</deployedseed>
    <deployedseed>1858090641</deployedseed>
    <deployedseed>828422674</deployedseed>
    <deployedseed>1914831505</deployedseed>
    <deployedseed>1611929572</deployedseed>
    <deployedseed>606015835</deployedseed>
    <deployedseed>730699633</deployedseed>
    <deployedseed>979059300</deployedseed>
    <deployedseed>508293826</deployedseed>
    <deployedseed>368203789</deployedseed>
    <deployedseed>1096709014</deployedseed>
    <deployedseed>2147467064</deployedseed>
    <deployedseed>1390768734</deployedseed>
    <deployedseed>1045324930</deployedseed>
    <deployedseed>2129221658</deployedseed>
    <deployedseed>1289526102</deployedseed>
    <deployedseed>1953220144</deployedseed>
    <deployedseed>1262257302</deployedseed>
    <deployedseed>958943284</deployedseed>
    <deployedseed>1360019874</deployedseed>
    <deployedseed>716293124</deployedseed>
    <deployedseed>441521100</deployedseed>
    <deployedseed>1153395500</deployedseed>
    <deployedseed>1758714884</deployedseed>
    <deployedseed>1262352110</deployedseed>
    <deployedseed>3601216</deployedseed>
    <deployedseed>604780799</deployedseed>
    <deployedseed>1741863166</deployedseed>
    <deployedseed>1979514819</deployedseed>
    <deployedseed>489268385</deployedseed>
    <deployedseed>90471718</deployedseed>
    <deployedseed>254634037</deployedseed>
    <deployedseed>124741046</deployedseed>
    <deployedseed>234715542</deployedseed>
    <deployedseed>255998581</deployedseed>
    <deployedseed>1479617296</deployedseed>
    <deployedseed>1316744247</deployedseed>
    <deployedseed>1642132642</deployedseed>
    <deployedseed>586110182</deployedseed>
    <deployedseed>51891356</deployedseed>
    <deployedseed>1638240332</deployedseed>
    <deployedseed>743779127</deployedseed>
    <deployedseed>33741760</deployedseed>
    <deployedseed>886683226</deployedseed>
    <deployedseed>1264851550</deployedseed>
    <deployedseed>394611649</deployedseed>
    <deployedseed>173075576</deployedseed>
    <deployedseed>129921668</deployedseed>
    <deployedseed>1741977697</deployedseed>
    <deployedseed>1285437124</deployedseed>
    <deployedseed>222613635</deployedseed>
    <deployedseed>355780174</deployedseed>
    <deployedseed>491233583</deployedseed>
    <deployedseed>351098951</deployedseed>
    <deployedseed>1600893383</deployedseed>
    <deployedseed>2110881167</deployedseed>
    <deployedseed>939353390</deployedseed>
    <deployedseed>1456003443</deployedseed>
    <deployedseed>2135339711</deployedseed>
    <deployedseed>1691872333</deployedseed>
    <deployedseed>1994438833</deployedseed>
    <deployedseed>931247173</deployedseed>
    <deployedseed>295964038</deployedseed>
    <deployedseed>69512240</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>mcq_correct(ta)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
