<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Evaluating the rational functions at the given point</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Consider the function on the interval \( {@-bb/aa@} < x < {@cc@} \) given by,</p>
<p style="margin-left: 30px">\( h(x)={@Exp4@} \)</p>

<br>

<p>(a) Find \(a\) such that \(h(a)=0\)</p>
<p style="margin-left: 30px">\(a=\) [[input:ans1]] [[validation:ans1]]</p>
<p>[[feedback:prt1]]</p>

<br>

<p>(b) Find \(b\) such that \(h^{-1}(b)={@ff@}\)</p>
<p style="margin-left: 30px">\(b=\) [[input:ans2]] [[validation:ans2]]</p>
<p>[[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked solution</h3></p>
<p>(a) We want to solve the equation,</p>
<p style="margin-left: 30px">\({@Exp4@}=0\)</p>
<p>This is true if the numerator is equal to \(0\) and the denominator is non-zero. First we solve the linear equation given by the numerator,</p>

<p style="margin-left: 30px">\({@Exp1@}=0\)</p>
<p style="margin-left: 30px">\({@dd*x@}={@-ee@}\)</p>
<p style="margin-left: 30px">\(x={@-ee/dd@}\)</p>

<p>Now we check that the denominator is non-zero at this point,</p>
<p hidden>{@simp:false;@}
<p style="margin-left: 30px">\(\begin{align*}
{@line1@} &= {@line2@} \\
&= {@line3@} \\
& \neq 0
\end{align*}\)</p>
<p hidden>{@simp:true;@}

<p>So, we may take \(a = {@-ee/dd@}\)</p>

<br>

<p>(b) You might have been tempted to find the inverse of \(h(x)\) and then solve \(h^{-1}(b)={@ff@}\).
<br>However, it is much easier to notice that if \(h^{-1}(b)={@ff@}\) then \(b = h({@ff@})\)
<br>So, we need only to substitute {@ff@} into \(h(x)\).</p>

<p style="margin-left: 30px">\(h(x) = {@Exp4@}\)</p>

<p style="margin-left: 30px">\(
h({@ff@}) = \frac{{@dd@} \cdot {@ff@} + {@ee@}}{{@aa@} \cdot {@ff@}^2 {@S1(-aa*cc+bb)@} \cdot {@ff@} - {@cc*bb@}}
= {@(dd*ff + ee) / (aa*ff^2 + (-aa*cc+bb)*ff - cc*bb)@}
\)</p>

<p>So, we have \(b={@Tans2@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[aa:rand(3)+2;
bb:rand(5)+2;
cc:rand(5)+3;
dd:rand(8)+2;
ee:rand_with_prohib(1,8,[dd*bb/aa, -dd*cc]);
ff:rand_with_prohib(1,3,[cc]);

Exp1:dd*x+ee;
Exp2:(aa*x+bb)*(x-cc);
Exp3:aa*x^2-aa*cc*x+bb*x-cc*bb;
Exp4:Exp1/Exp3;

Tans1:-ee/dd;
Tans2:subst(ff,x,Exp4);


/*-------------------------------------------------------------------------------------------*/
/*Define function S(n) that given a number n, returns the string "+n" if positive (or zero) and "-n" if negative, if n=+1 or -1 it returns + or - respectively*/
S(n):= block([], if (n=1) then return("+") elseif (n=-1) then return("-") elseif (n>=0) then return(vconcat("+", abs(n)))  else  return(vconcat("-", abs(n))));

/*Define function S1(n) that given a number n, returns the string "+n" if positive (or zero) and "-n" if negative (1 and -1 are treated as the other numbers)*/
S1(n):= block([], if (n>=0) then return(vconcat("+", abs(n)))  else  return(vconcat("-", abs(n))));

/*Define function SP(n) that given a number n, returns the string "n" if positive (or zero) and "(-n)" if negative*/
SP(n):= block([], if (n>=0) then return(n)  else  return(vconcat("(-", abs(n), ")")));

/*-------------------------------------------------------------------------------------------*/

/*Worked solution: simplified lines*/
simp:false;

line1: aa*Tans1^2 - ev(aa*cc - bb, simp)*Tans1 - ev(cc*bb, simp);
line2: ev(aa*Tans1^2,simp) - ev((aa*cc - bb)*Tans1,simp) - ev(cc*bb, simp);
line3: ev(aa*Tans1^2 - (aa*cc - bb)*Tans1 - cc*bb, simp);

simp:true;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>function: \(h(x)={@Exp4@}\) defined for \( {@-bb/aa@} < x < {@cc@} \).</p>
<p>When \(h(a)=0, a={@Tans1@}\)</p>
<p>When \(h^{-1}(b)={@ff@}, b={@Tans2@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>Tans1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>Tans2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>mistake12:-ee/(cc*bb);
mistake13:cc;
mistake14:-bb/aa;

sol1:solve(dd*x+ee=aa*x^2-aa*cc*x+bb*x-cc*bb,x);
mistake15:rhs(sol1[1]);
mistake16:rhs(sol1[2]);

mistake17:ee/dd;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>Tans1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake12</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You need to solve \(h(x)=0\) not substitute \(x=0\)</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake13</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have solved the denominator equal to 0. This would not find when \(h(x)=0\) but rather when the function was not defined (because you cannot divide by 0.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake14</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have solved the denominator equal to 0. This would not find when \(h(x)=0\) but rather when the function was not defined (because you cannot divide by 0).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake15</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>I think you have solved the numerator equal to the denominator. Perhaps you thought it was to be equal to 1 and not 0?</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>5</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>5</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake16</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-6-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>I think you have solved the numerator equal to the denominator. Perhaps you thought it was to be equal to 1 and not 0?</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>6</falsenextnode>
        <falseanswernote>prt1-6-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>6</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake17</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-7-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>I think you have made a mistake with the sign </p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-7-F</falseanswernote>
        <falsefeedback format="html">
          <text>You must have made a mistake somewhere.</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>sol2:solve(Exp4=ff,x);
mistake22:rhs(sol2[1]);
mistake23:rhs(sol2[2]);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>Tans2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>mistake22</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>I think you have solved \(h(b)={@ff@}\) not \(h^{-1}({@ff@})=b\)</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>mistake23</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>I think you have solved \(h(b)={@ff@}\) not \(h^{-1}({@ff@})=b\)</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-3-F</falseanswernote>
        <falsefeedback format="html">
          <text>I think you have made a mistake somewhere.</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>523399364</deployedseed>
    <deployedseed>1026335313</deployedseed>
    <deployedseed>19754215</deployedseed>
    <deployedseed>233853006</deployedseed>
    <deployedseed>62967986</deployedseed>
    <deployedseed>649191430</deployedseed>
    <deployedseed>163139578</deployedseed>
    <deployedseed>826921803</deployedseed>
    <deployedseed>1372053024</deployedseed>
    <deployedseed>1511180538</deployedseed>
    <deployedseed>1757596551</deployedseed>
    <deployedseed>1630340295</deployedseed>
    <deployedseed>1156647825</deployedseed>
    <deployedseed>687670075</deployedseed>
    <deployedseed>1864754257</deployedseed>
    <deployedseed>20253961</deployedseed>
    <deployedseed>922674072</deployedseed>
    <deployedseed>3753077</deployedseed>
    <deployedseed>285716372</deployedseed>
    <deployedseed>1496412103</deployedseed>
    <deployedseed>740855957</deployedseed>
    <deployedseed>1775766777</deployedseed>
    <deployedseed>228632099</deployedseed>
    <deployedseed>1030051445</deployedseed>
    <deployedseed>151261010</deployedseed>
    <deployedseed>1270666789</deployedseed>
    <deployedseed>1416306763</deployedseed>
    <deployedseed>57636830</deployedseed>
    <deployedseed>695017069</deployedseed>
    <deployedseed>1282223134</deployedseed>
    <deployedseed>2019174479</deployedseed>
    <deployedseed>80549212</deployedseed>
    <deployedseed>1927079788</deployedseed>
    <deployedseed>94140104</deployedseed>
    <deployedseed>536780722</deployedseed>
    <deployedseed>1118645124</deployedseed>
    <deployedseed>303549440</deployedseed>
    <deployedseed>595929026</deployedseed>
    <deployedseed>826214009</deployedseed>
    <deployedseed>142910073</deployedseed>
    <deployedseed>562786114</deployedseed>
    <deployedseed>415386879</deployedseed>
    <deployedseed>1738327986</deployedseed>
    <deployedseed>347371211</deployedseed>
    <deployedseed>888745476</deployedseed>
    <deployedseed>1866100839</deployedseed>
    <deployedseed>142686597</deployedseed>
    <deployedseed>1952952290</deployedseed>
    <deployedseed>1008366787</deployedseed>
    <deployedseed>654237424</deployedseed>
    <deployedseed>164484389</deployedseed>
    <deployedseed>1240685700</deployedseed>
    <deployedseed>442826771</deployedseed>
    <deployedseed>843111385</deployedseed>
    <deployedseed>1464809411</deployedseed>
    <deployedseed>1608872162</deployedseed>
    <deployedseed>1981670867</deployedseed>
    <deployedseed>386580943</deployedseed>
    <deployedseed>148918924</deployedseed>
    <deployedseed>160595694</deployedseed>
    <deployedseed>651302339</deployedseed>
    <deployedseed>259063256</deployedseed>
    <deployedseed>367089737</deployedseed>
    <deployedseed>2106227927</deployedseed>
    <deployedseed>1983233536</deployedseed>
    <deployedseed>232438260</deployedseed>
    <deployedseed>1233742474</deployedseed>
    <deployedseed>1407293741</deployedseed>
    <deployedseed>787782740</deployedseed>
    <deployedseed>79358024</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>Tans1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>Tans2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
