<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the reduced row echelon form of a matrix</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Find the reduced row echelon form of the matrix \(A\).</p>
<p>\(A={@matrix_A@}\)</p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p><b>Step1: </b>Here, the left most nonzero entry is in the second column.</p>
<p><b>Step2: </b>By swapping row \(1\) with row \(3\) , we can obtain the pivot position as follows:</p>
<p>\({@matrix_A@} R_{1}\leftrightarrow R_{3} {@matrix_1@}\)</p>
<p><b>Step 3:</b>  Now the leading entry is \(1\), and to create zeros in all positions below the pivot, we use the replacement rule:</p>
<p>\(R_{2}\rightarrow R_{2}+(-{@aa@})R_{1} {@matrix_2@}\)</p>
<p><b>Step 4:</b> Now we proceed to the second row. Here, the leading entry is \(−2\). Using a scaling rule we obtain a leading \(1\):</p>
<p>\(R_{2}\rightarrow ({@ww@})R_{2} {@matrix_3@}\)</p>
<p>And applying row replacement rule:</p>
<p>\(R_{3}\rightarrow R_{3}+({@ww_1@})R_{2} {@matrix_4@}\)</p>
<p>And scaling \(R_{3},\)</p>
<p>\(R_{3}\leftrightarrow ({@ww_4@})R_{3}{@matrix_5@}\)</p>
<p><b>Step 5:</b> Beginning with the rightmost pivot column, we create zeros above each pivot
element. That is, we start from the fourth column:</p>
<p>\(R_{1}\rightarrow R_{1}+({@-ee@})R_{3},R_{2}\rightarrow R_{2}+({@ww_7@})R_{3}{@matrix_6@}\)</p>
<p>And using row replacement (to create zeros above the pivot element in the third column),
</p>
<p>\(R_{1}\rightarrow R_{1}+({@ww_8@})R_{2}{@matrix_7@}\)</p>
<p>Thus, the required matrix in reduced row echelon form is given by
</p>
<p>\({@ta@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>aa:rand(3)+1;
bb:rand_with_prohib(4,8,[aa]);
cc:rand_with_prohib(-5,5,[aa,bb,0]);
dd:rand_with_step(1,2,8);
ee:rand([2,4,6]);
matrix_A:matrix([0,0,aa,bb],[0,aa,0,dd],[0,dd,dd,ee]);/*matrix A*/
ta:rref(matrix_A);/*reduced row echelon form*/


/*Worked solution*/
matrix_1:matrix([0,dd,dd,ee],[0,aa,0,dd],[0,0,aa,bb]);/*Interchanging row 1 and 3*/
matrix_2:matrix([0,dd,dd,ee],[0+(-aa*0),aa+(-aa*dd),0+(-aa*dd),dd+(-aa*ee)],[0,0,aa,bb]);
ww:1/(0+(-aa*dd));
ww_1:0+(-aa*dd);
ww_2:ww*(0+(-aa*dd));
ww_3:ww*(dd+(-aa*ee));
ww_4:1/(bb+(ww_1*ww_3));
ww_5:aa+(ww_1*ww_2);
ww_6:ww_4*(bb+(ww_1*ww_3));
ww_7:-ww_3;
ww_8:-1*(dd+(-ee*ww_5));
ww_9:ww*(0+(-aa*0))+ww_7*0;
ww_10:ww*(aa+(-aa*dd))+ww_7*0;
ww_11:ww*(0+(-aa*dd))+ww_7*ww_5;
w_12:ww*(dd+(-aa*ee))+ww_7*ww_6;
matrix_3:matrix([0,dd,dd,ee],[ww*(0+(-aa*0)),ww*(aa+(-aa*dd)),ww*(0+(-aa*dd)),ww*(dd+(-aa*ee))],[0,0,aa,bb]);
matrix_4:matrix([0,dd,dd,ee],[ww*(0+(-aa*0)),ww*(aa+(-aa*dd)),ww*(0+(-aa*dd)),ww*(dd+(-aa*ee))],[0,0,aa+(ww_1*ww_2),bb+(ww_1*ww_3)]);
matrix_5:matrix([0,dd,dd,ee],[ww*(0+(-aa*0)),ww*(aa+(-aa*dd)),ww*(0+(-aa*dd)),ww*(dd+(-aa*ee))],[0,0,(aa+ww_1)*ww_2,ww_4*(bb+(ww_1*ww_3))]);
matrix_6:matrix([0+(-ee*0),dd+(-ee*0),dd+(-ee*ww_5),ee+(-ee*ww_6)],[ww*(0+(-aa*0))+ww_7*0,ww*(aa+(-aa*dd))+ww_7*0,ww*(0+(-aa*dd))+ww_7*ww_5,ww*(dd+(-aa*ee))+ww_7*ww_6],[0,0,(aa+ww_1)*ww_2,ww_4*(bb+(ww_1*ww_3))]);
matrix_7::matrix([0+(-ee*0)+ww_8*ww_9,dd+(-ee*0)+ww_8*ww_10,dd+(-ee*ww_5)+ww_8*ww_11,ee+(-ee*ww_6)+ww_8*w_12],[ww*(0+(-aa*0))+ww_7*0,ww*(aa+(-aa*dd))+ww_7*0,ww*(0+(-aa*dd))+ww_7*ww_5,ww*(dd+(-aa*ee))+ww_7*ww_6],[0,0,(aa+ww_1)*ww_2,ww_4*(bb+(ww_1*ww_3))]);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>matrix \(A={@matrix_A@}\)</p>
<p>RREF\(={@ta@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>matrix</type>
      <tans>ta</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1800584193</deployedseed>
    <deployedseed>1556720942</deployedseed>
    <deployedseed>1799677040</deployedseed>
    <deployedseed>2024746885</deployedseed>
    <deployedseed>1998989111</deployedseed>
    <deployedseed>248588166</deployedseed>
    <deployedseed>820832314</deployedseed>
    <deployedseed>505401396</deployedseed>
    <deployedseed>2039635210</deployedseed>
    <deployedseed>599296126</deployedseed>
    <deployedseed>925265658</deployedseed>
    <deployedseed>1765276480</deployedseed>
    <deployedseed>928594230</deployedseed>
    <deployedseed>1610658871</deployedseed>
    <deployedseed>141044003</deployedseed>
    <deployedseed>1794719336</deployedseed>
    <deployedseed>572713011</deployedseed>
    <deployedseed>1739558762</deployedseed>
    <deployedseed>1218933729</deployedseed>
    <deployedseed>846889083</deployedseed>
    <deployedseed>2094474788</deployedseed>
    <deployedseed>1685626404</deployedseed>
    <deployedseed>582007984</deployedseed>
    <deployedseed>821093564</deployedseed>
    <deployedseed>1221139015</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
