<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Reduced Row Echelon Form of a Matrix</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Write the following matrix in Reduced Row Echelon form.</p> 
<p style="margin-left: 30px">\(R = {@R@}\)</p>

<br>

<table>
  <tr>
    <td>\(\text{rref}(R) = \)</td> 
    <td>&nbsp; [[input:ans1]] [[validation:ans1]]</td>
  </tr>
</table>

<p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>We need to perform row operations so that the matrix is in reduced row echelon form. 
<br>A matrix is in reduced row echelon form if</p>
<ul>
  <li>All non-zero rows are above any rows of all zeroes</li>
  <li>The leading coefficient of a nonzero row is always strictly to the right of the leading coefficient of the row above it</li>
  <li>The leading coefficient of a nonzero row is \(1\)</li>
  <li>The leading coefficient is the only non-zero entry in its column.</li>
</ul>

<p>We are given the matrix,</p>
<p style="margin-left: 30px">\( R = {@matrix([1,2,3],[0,2,aa])@} \)</p>

<br>

<p><b>Step 1:</b> It is already the case that the leading coefficient in the first row is \(1\), and that this is the only non-zero entry in the first column.
<br>Therefore, no changes are required to the first column of the matrix.</p>

<p style="margin-left: 30px">\( R = {@matrix([1,2,3],[0,2,aa])@} \)</p>

<br>

<p><b>Step 2:</b> We need to make the first non-zero entry in the second row equal to \(1\). 
<br>To do this, we'll perform the following row operation,</p>

<p style="margin-left: 30px">\( \text{Row(2)} \rightarrow \frac{1}{2} \times \text{Row(2)} \)</p>
<p style="margin-left: 30px">
\( {@matrix([1,2,3],[0,2,aa])@} \rightarrow {@matrix([1,2,3],[0,1,aa/2])@} \)
</p>

<br>

<p><b>Step 3:</b> Since the second column now contains a leading \(a\), we must make all the other entries in this column zero. In particular, the second entry in the first row must be zero.
<p>To do this, we'll perform the following row operation,</p>
<p style="margin-left: 30px">\( \text{Row(1)} \rightarrow \text{Row(1)} - 2 \times \text{Row(2)} \).</p>
<p style="margin-left: 30px">
\( {@matrix([1,2,3],[0,1,aa/2])@} \rightarrow {@matrix([1,0,(3-2*(aa/2))],[0,1,aa/2])@} \)
</p>

<br>

<p>The resulting matrix is in satisfies all of the required condition and is therefore in Reduced Row Echelon Form (RREF).</p>
<p style="margin-left: 30px">\( \text{rref}(R) = {@ matrix([1,0,(3-2*(aa/2))],[0,1,aa/2])@} \)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>aa:rand([2,4,6,8]);
R:matrix([1,2,3],[0,2,aa]);
RR:rref(R);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Matrix  \[{@R@}\]</p>
<p>Reduced Row Echelon form: \({@RR@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>matrix</type>
      <tans>RR</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>RR</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>632419799</deployedseed>
    <deployedseed>1048331305</deployedseed>
    <deployedseed>695195288</deployedseed>
    <deployedseed>449928778</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>RR</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>matrix([1,0,0],[0,1,0])</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
