<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Proving the set of natural numbers using mathematical induction</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Consider the following proposition, then fill in the blanks to complete the proof outlined below it.</p>

<br>

<p><strong>Proposition:</strong> Show that by mathematical induction that &nbsp; \(\displaystyle  \sum_{k=1}^{n} k^2 = \frac{n(n+1)(2n+1)}{6}\)</p>
<p><em>Note: "LHS" means Left Hand Side and "RHS" means Right Hand Side.</em></p>
<p><em>We also use the notation, &nbsp; \( \sum_{k=1}^{n} k^2 = 1^2 + 2^2 + 3^2 + \cdots + n^2 \)</em></p>

<br>

<p><strong>Proof by induction:</strong></p>
<p>Define \(P(n)\) to be the statement of the proposition, that: \(\displaystyle  \sum_{k=1}^{n} k^2 = \frac{n(n+1)(2n+1)}{6} \)</p>

<p><b>Step 1:</b> Base case.</p>
<p>For \(n = 1\), we have</p>

<p style="margin-left:30px">LHS: \(\displaystyle \sum_{k=1}^{n} k^2 =\) [[input:ans1]] [[validation:ans1]]</p>
<p>[[feedback:prt1]]</p> 

<p style="margin-left:30px">RHS: \( \frac{n(n+1)(2n+1)}{6}=\) [[input:ans2]] [[validation:ans2]]</p>
<p>[[feedback:prt2]]</p>

<p>Since these are equal, \( P(1) \) is [[input:ans3]] [[validation:ans3]]</p>
<p>[[feedback:prt3]]</p>

<br>

<p><b>Step 2:</b> Induction step.</p>
<p>Assume \(P(n)\) is true, which means that \( \sum_{k=1}^{n} k^2 = \frac{n(n+1)(2n+1)}{6} \)
<br>This is called the <b>induction hypothesis</b>.</p>
<p>We then need to show \( P(n+1) \) is true.</p>

<br>

<p>\(P(n+1): \displaystyle \sum_{k=1}^{n+1} k^2 = \sum_{k=1}^{n} k^2 \) + [[input:ans4]] [[validation:ans4]]</p>
<p style="margin-left:131px">\(=\) [[input:ans5]] [[validation:ans5]] &nbsp; (using the induction hypothesis)</p>
<p style="margin-left:131px">\(= (n+1) \times ( \) [[input:ans5b]] \()\) [[validation:ans5b]] &nbsp; (factorizing out the common term)</p>
<p style="margin-left:131px">\(= (n+1)(n+2) \times ( \) [[input:ans5c]] \()\) [[validation:ans5c]] &nbsp; (factorising the quadratic)</p>
<p style="margin-left:131px">\(=\) [[input:ans5d]] [[validation:ans5d]] &nbsp; (writing in the form of the RHS of \(P(n+1)\))</p>
<p>[[feedback:prt4]]</p>

<p>So, \(P(n+1)\) is [[input:ans6]] [[validation:ans6]] [[feedback:prt5]]</p>

<br>

<p>Therefore, by the Principle of Mathematical Induction, we have \( \displaystyle  \sum_{k=1}^{n} k^2 = \frac{n(n+1)(2n+1)}{6}\)</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p><strong>Proposition:</strong> Show that by mathematical induction that &nbsp; \(\displaystyle  \sum_{k=1}^{n} k^2 = \frac{n(n+1)(2n+1)}{6}\)</p>
<p><em>Note: "LHS" means Left Hand Side and "RHS" means Right Hand Side.</em></p>
<p><em>We also use the notation, &nbsp; \( \sum_{k=1}^{n} k^2 = 1^2 + 2^2 + 3^2 + \cdots + n^2 \)</em></p>

<br>

<p><strong>Proof by induction:</strong></p>
<p>Define \(P(n)\) to be the statement of the proposition, that: \(\displaystyle  \sum_{k=1}^{n} k^2 = \frac{n(n+1)(2n+1)}{6} \)</p>

<p><b>Step 1:</b> Base case.</p>
<p>For \(n = 1\), we have</p>

<p style="margin-left:30px">LHS: \(\displaystyle \sum_{k=1}^{n} k^2 = 1^2 = 1\)</p>

<p style="margin-left:30px">RHS: \( \frac{n(n+1)(2n+1)}{6} = \frac{1 \cdot (1+1) \cdot (2+1)}{6} = 1\)</p>

<p>Since these are equal, \( P(1) \) is <b>true</b>.</p>

<br>

<p><b>Step 2:</b> Induction step.</p>
<p>Assume \(P(n)\) is true, which means that \( \sum_{k=1}^{n} k^2 = \frac{n(n+1)(2n+1)}{6} \)
<br>This is called the <b>induction hypothesis</b>.</p>
<p>We then need to show \( P(n+1) \) is true.</p>

<br>

<p>\(P(n+1): \displaystyle \sum_{k=1}^{n+1} k^2 = \sum_{k=1}^{n} k^2 \) + (n+1)^2</p>
<p style="margin-left:131px">\(=\) {@ta5@} &nbsp; (using the induction hypothesis)</p>
<p style="margin-left:131px">\(= (n+1) \times ({@ta5b@})\) &nbsp; (factorizing out the common term)</p>
<p style="margin-left:131px">\(= (n+1)(n+2) \times ({@ta5c@})\) &nbsp; (factorising the quadratic)</p>
<p style="margin-left:131px">\(= {@ta5d@}\) &nbsp; (writing in the form of the RHS of \(P(n+1)\))</p>

<p>So, \(P(n+1)\) is <b>true</b>.</p>

<br>

<p>Therefore, by the Principle of Mathematical Induction, we have \( \displaystyle  \sum_{k=1}^{n} k^2 = \frac{n(n+1)(2n+1)}{6}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>ta1:1;
ta2:1;
ta3:true;
ta4:sum(k^2,k,n+1,n+1);
ta5:n*(n+1)*(2*n+1)/6 + (n+1)^2;
ta5b:(n*(2*n+1)/6 + (n+1));
ta5c:(2*n+3)/6;
ta5d:(n+1)*(n+2)*(2*(n+1)+1)/6;
ta6:true;

/*
res1:n*(n+1)*(2*n+1)/6+(n+1)^2;
res2:(n+1)*(n*(2*n+1)/6+(n+1));
res3:(n+1)*(2*n^2+7*n+6)/6;
res4:(n+1)*(n+2)*(2*n+3)/6;
res5:(n+1)*(n+2)*(2*(n+1)+1)/6;
ta5:[n*(n+1)*(2*n+1)/6+(n+1)^2,(n+1)*(n*(2*n+1)/6+(n+1)),(n+1)*(2*n^2+7*n+6)/6,(n+1)*(n+2)*(2*n+3)/6,(n+1)*(n+2)*(2*(n+1)+1)/6];
*/</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text/>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>boolean</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>ta5</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5b</name>
      <type>algebraic</type>
      <tans>ta5b</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5c</name>
      <type>algebraic</type>
      <tans>ta5c</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5d</name>
      <type>algebraic</type>
      <tans>ta5d</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6</name>
      <type>boolean</type>
      <tans>ta6</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Since \(n=1\), we only have the first term of the series.
<br>So, the series evaluates to \(1^2 = 1\) for the base case.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Evaluating the expression at \(n=1\) gives, \( \frac{n(n+1)(2n+1)}{6}=\frac{1 \cdot (1+1) \cdot (2+1)}{6} = 1\)</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Since the LHS and RHS of the equation in \(P(1)\) are equal to \(1\), we have that the statement of \(P(1)\) is <b>true</b>.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/* node 1: correct first line */
/* node 2: correct second line */
/* node 3: second line equivalent but in the wrong form */
/* node 4: correct third line */
/* node 5: third line equivalent but in the wrong form */
/* node 6: correct fourth line */
/* node 7: fourth line equivalent but in the wrong form */
/* node 8: correct fifth line */
/* node 9: fifth line equivalent but in the wrong form */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1/5</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>In the <b>first line</b>, we should extract the <b>last</b> term in the summation to give, \( \sum_{k=1}^{n+1} k^2 = \sum_{k=1}^{n} k^2 + (n+1)^2 \)</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>EqualComAss</answertest>
        <sans>ans5</sans>
        <tans>ta5</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1/5</truescore>
        <truepenalty/>
        <truenextnode>3</truenextnode>
        <trueanswernote>prt4-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt4-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[{@if ans4=ta4 then "" else "<br>"@}]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>ta5</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>3</truenextnode>
        <trueanswernote>prt4-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>The answer you have given for the <b>second line</b> is equal to the series given.
<br>However, your answer is not of the required form.
<br>You should only have used the induction hypothesis to replace the summation in the line above,</p>

<p style="margin-left:30px">\(\sum_{k=1}^{n} k^2 + (n+1)^2 = {@ta5@}\)</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt4-3-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The answer you have given for the <b>second line</b> is not equal to the series given.
<br>You should have used the induction hypothesis to replace the summation in the line above,</p>
<p style="margin-left:30px">\(\sum_{k=1}^{n} k^2 + (n+1)^2 = {@ta5@}\)</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description/>
        <answertest>EqualComAss</answertest>
        <sans>ans5b</sans>
        <tans>ta5b</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1/5</truescore>
        <truepenalty/>
        <truenextnode>5</truenextnode>
        <trueanswernote>prt4-4-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt4-4-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[{@if ans4=ta4 and ans5=ta5 then "" else "<br>"@}]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans5b</sans>
        <tans>ta5b</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>5</truenextnode>
        <trueanswernote>prt4-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>The answer you have given for the <b>third line</b> is equal to the series given.
<br>However, your answer is not of the required form.
<br>You should only have factorized out the common term \((n+1)\) from the line above,</p>
<p style="margin-left:30px">\({@ta5@} = (n+1) \times ({@ta5b@})\)</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>5</falsenextnode>
        <falseanswernote>prt4-5-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The answer you have given for the <b>third line</b> is not equal to the series given.
<br>You should have factorized out the common term \((n+1)\) from the line above,</p>
<p style="margin-left:30px">\({@ta5@} = (n+1) \times ({@ta5b@})\)</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>5</name>
        <description/>
        <answertest>EqualComAss</answertest>
        <sans>ans5c</sans>
        <tans>ta5c</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1/5</truescore>
        <truepenalty/>
        <truenextnode>7</truenextnode>
        <trueanswernote>prt4-6-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>6</falsenextnode>
        <falseanswernote>prt4-6-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[{@if ans4=ta4 and ans5=ta5 and ans5b=ta5b then "" else "<br>"@}]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>6</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans5c</sans>
        <tans>ta5c</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>7</truenextnode>
        <trueanswernote>prt4-7-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>The answer you have given for the <b>fourth line</b> is equal to the series given.
<br>However, your answer is not of the required form.
<br>You should only have factorized the quadratic obtained from expanding the line above,</p>
<p style="margin-left:30px">\((n+1) \times {@ta5b@} = (n+1)(n+2) \times ({@ta5c@})\)</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>7</falsenextnode>
        <falseanswernote>prt4-7-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The answer you have given for the <b>fourth line</b> is not equal to the series given.
<br>You should have factorized the quadratic obtained from expanding the line above,</p>
<p style="margin-left:30px">\((n+1) \times {@ta5b@} = (n+1)(n+2) \times ({@ta5c@})\)</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>7</name>
        <description/>
        <answertest>EqualComAss</answertest>
        <sans>ans5d</sans>
        <tans>ta5d</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1/5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-8-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>8</falsenextnode>
        <falseanswernote>prt4-8-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[{@if ans4=ta4 and ans5=ta5 and ans5b=ta5b and ans5c=ta5c then "" else "<br>"@}]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>8</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans5d</sans>
        <tans>ta5d</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-9-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>The answer you have given for the <b>last line</b> is equal to the series given.
<br>However, your answer is not of the required form.
<br>You should have simplified the line above to be equal to the RHS of the statement \(P(n+1)\),</p>
<p style="margin-left:30px">\((n+1)(n+2) \times {@ta5c@} = {@ta5d@}\)</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-9-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The answer you have given for the <b>last line</b> is not equal to the series given.
<br>You should have simplified the line above to be equal to the RHS of the statement \(P(n+1)\),</p>
<p style="margin-left:30px">\((n+1)(n+2) \times {@ta5c@} = {@ta5d@}\)</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans6</sans>
        <tans>ta6</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>\(P(n+1)\) is the statement that \(\sum_{k=1}^{n+1} k^2 = \frac{(n+1)(n+2)(2(n+1) + 1)}{6}\)</p>
<p>We have shown that this is true using the induction hypothesis in the previous part.
<br>Therefore, \(P(n+1)\) is <b>true</b>.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ta5</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>ta6</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt5-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
