<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Converting terminating decimal to fraction form</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Convert the  termination decimal \({@val_decim@}\) to fraction form.</p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p>Step 1: Identify the whole number and decimal parts.</p>
<p>The whole number part is {@whole_part@}. </p>
<p>The decimal part is {@decimal_part@}. </p>

<p>Step 2: Convert the decimal part to a fraction and simplify it by dividing both the numerator and the denominator by the GCD. {@decimal_part_frac@}</p>

<p>Step 3: Combine the whole number and  the fraction to get \({@ta@}\) .</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>aa: rand_with_step(27,399,5);
bb: rand_with_step(257,399,3);
val1_aa: aa/10;
val2_aa: aa/100;
val3_aa: aa/1000;

val1_bb: bb/10;
val2_bb: bb/100;
val3_bb: bb/1000;

val_decim: float(rand([val1_aa,val2_aa,val3_aa,val1_bb,val2_bb,val3_bb]));
whole_part:floor(val_decim);
decimal_part:val_decim-whole_part;
decimal_part_frac:rat(decimal_part)
ta: if val_decim=float(val1_aa) then val1_aa else (if val_decim=float(val2_aa) then val2_aa else (if val_decim=float(val3_aa) then val3_aa else (if val_decim=float(val1_bb) then val1_bb else (if val_decim=float(val2_bb) then val2_bb else val3_bb))));</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Terminating decimal: \({@val_decim@}\)</p>
<p>fraction form: \({@ta@}\)</p>
<p>To express \({@val_decim@}\) as a fraction:</p>
<p>The whole number part is \({@whole_part@}\).</p.
<p>The decimal part, \({@decimal_part@}\), can be written as \({@decimal_part_frac@}\).</p>
<P>Combine the whole number \({@whole_part@}\) and the fraction \({@decimal_part_frac@}\)to get \({@ta@}\) .</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>369126791</deployedseed>
    <deployedseed>1932363033</deployedseed>
    <deployedseed>1437358640</deployedseed>
    <deployedseed>680567859</deployedseed>
    <deployedseed>766989064</deployedseed>
    <deployedseed>1787379254</deployedseed>
    <deployedseed>2074718722</deployedseed>
    <deployedseed>445831357</deployedseed>
    <deployedseed>2110529022</deployedseed>
    <deployedseed>1304783736</deployedseed>
    <deployedseed>801936695</deployedseed>
    <deployedseed>1941460124</deployedseed>
    <deployedseed>636840516</deployedseed>
    <deployedseed>1123217547</deployedseed>
    <deployedseed>1640277215</deployedseed>
    <deployedseed>333626499</deployedseed>
    <deployedseed>731699463</deployedseed>
    <deployedseed>386909988</deployedseed>
    <deployedseed>685370035</deployedseed>
    <deployedseed>1896746435</deployedseed>
    <deployedseed>1294165834</deployedseed>
    <deployedseed>1878492766</deployedseed>
    <deployedseed>33055171</deployedseed>
    <deployedseed>88405955</deployedseed>
    <deployedseed>571514733</deployedseed>
    <deployedseed>80909005</deployedseed>
    <deployedseed>1488905233</deployedseed>
    <deployedseed>17767562</deployedseed>
    <deployedseed>454202225</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
