<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Locating Center, Foci, and Vertices of an Ellipse with Given Minor Axis and Eccentricity</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Suppose the endpoints of the minor axis of an ellipse are \(({@x_cord@}, \pm {@y_cord@})\) and its eccentricity is \({@decimalplaces(eccentricity,1)@}\).</p>
<p>Find the coordinates of:</p>

<p>(a) The center, &nbsp; \(C:\) [[input:ans1]] [[validation:ans1]]</p>
<p>[[feedback:prt1]] </p>

<br>

<p>(b) The foci, &nbsp; \(F_1:\) [[input:ans2]] [[validation:ans2]]</p>
<p style="margin-left:90px">\(F_2:\) [[input:ans3]] [[validation:ans3]]</p>
<p>[[feedback:prt2]]</p>

<br>

<p>(c) The vertices of the ellipse, &nbsp; \(V_1:\) [[input:ans4]] [[validation:ans4]]</p>
<p style="margin-left:202px">\(V_2:\) [[input:ans5]] [[validation:ans5]]</p>
<p>[[feedback:prt3]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>To find the coordinates of the center, foci, and vertices of the ellipse, we can use the standard form for an ellipse with centre \((x_0, y_0)\),</p>

<p style="margin-left:30px">\(\displaystyle 
\frac{(x - x_0)^2}{a^2} + \frac{(y - y_0)^2}{b^2} = 1
\)</p>
<p>where \(a\) is the semi-major axis length and \(b\) is the semi-minor axis length.</p>

<p>The eccentricity \(e\) of the ellipse is given by:</p>
<p style="margin-left:30px">\(\displaystyle 
e = \sqrt{1 - \frac{b^2}{a^2}}
\)</p>

<p>Given that the endpoints of the minor axis are \(({@x_cord@}, \pm {@y_cord@})\), we can use this information to find \(b\) and then calculate \(a\) using the eccentricity.</p>

<br>

<p>(a) To find the center of the ellipse, we can simply find the midpoint of the minor axis:</p>
<p style="margin-left:30px">\(C: \displaystyle  \left(\frac{{@minor_axis_endpoints[1][1]@} + {@minor_axis_endpoints[2][1]@}}{2}, \frac{{@minor_axis_endpoints[1][2]@}{@minor_axis_endpoints[2][2]@}}{2}\right) 
= ({@center[1]@}, {@center[2]@})\)</p>
<p>Then, we can write the equation of this ellipse as,</p>
<p style="margin-left:30px">\(\displaystyle 
\frac{(x - x_0)^2}{a^2} + \frac{y^2}{b^2} = 1
\)</p>

<br>

<p>(b) The foci of the ellipse can be found using the formula \(c = ae\), where \(c\) is the distance from the center to each focus. 
<br>Since the center is at \(({@center[1]@}, {@center[2]@})\) and the eccentricity \(e\) is given as {@float(eccentricity)@}, we can calculate \(a\) and then use it to find \(c\).</p>

<p>Given \(e = {@decimalplaces(eccentricity,1)@}\), we can rearrange the eccentricity formula to solve for \(a\):</p>
<p style="margin-left:30px">\( \displaystyle 
a = \frac{b}{\sqrt{1 - e^2}}
\)</p>

<p>Using the endpoints of the minor axis, \( \textstyle  b = \frac{{@minor_axis_endpoints[1][2]@}-({@minor_axis_endpoints[2][2]@})}{2} = {@b@}\) and \(e = {@float(eccentricity)@}\).</p>

<p>Then, for \(a\),</p>
<p style="margin-left:30px">\( \displaystyle 
a = \frac{{@b@}}{\sqrt{1 - {@float(eccentricity)@}^2}} 
= \frac{{@b@}}{\sqrt{{@float(1 - eccentricity^2)@}}} 
= {@a@}
\)</p>

<p>Now we can find \(c\) using \(c = ae\):</p>
<p style="margin-left:30px">\( \displaystyle 
c = {@a@} \cdot {@float(eccentricity)@} 
= {@c@}
\)</p>

<p>So, the coordinates of the foci are,</p>
<p style="margin-left:30px">\( \displaystyle 
\left({@center[1]@} \pm {@c@}, 0 \right) 
= \left({@focus1[1]@}, {@focus1[2]@} \right), \left({@focus2[1]@}, {@focus2[2]@}\right)\)
</p>

<br>

<p>(c) The vertices of the ellipse can be found using the values of \(a\) and \(b\), which we already calculated. Since the vertices occur when \(y=0\), we have.</p>

<p style="margin-left:30px">\((x-{@center[1]@})^2 = a^2\)</p>
<p style="margin-left:30px">\(x-{@center[1]@} = \pm a = \pm {@a@}\)</p>

<p>Therefore, the vertices along the major axis are at,</p>
<p style="margin-left:30px">\( \displaystyle  
\left( {@center[1]@} \pm {@a@}, 0 \right) 
= \left({@vertex1[1]@}, {@vertex1[2]@} \right), \left({@vertex2[1]@}, {@vertex2[2]@} \right)\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>x_cord:rand(5)+1;
y_cord:rand(7)+3;
minor_axis_endpoints: [[x_cord, y_cord], [x_cord, -y_cord]];
nn:8;
eccentricity:nn/10;
/* Calculate b from the endpoints of the minor axis */
b: (minor_axis_endpoints[1][2] - minor_axis_endpoints[2][2]) / 2;  

/* Calculate a using the eccentricity */
a: b / sqrt(1 - eccentricity^2);  

/* Calculate the distance from the center to the foci (c) */
c: a * eccentricity;

/* Calculate the center of the ellipse */
center: [(minor_axis_endpoints[1][1] + minor_axis_endpoints[2][1]) / 2, (minor_axis_endpoints[1][2]+minor_axis_endpoints[2][2])/2];

/* Calculate the coordinates of the foci */
focus1: [center[1] + c, 0];
focus2: [center[1] - c, 0];  

/* Calculate the coordinates of the vertices */
vertex1: [center[1] + a, 0];
vertex2: [center[1] - a, 0];</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>minor axes coordinates of an ellipse are \(({@x_cord@}, \pm {@y_cord@})\) and its eccentricity is \({@decimalplaces(eccentricity,1)@}\)</p>
<p>(a) the center {@center@}</p>
<p>(b) the foci \(F_1{@focus1@}\)  and \(F_2{@focus2@}\)</p>
<p>(c) the vertices of the ellipse: \(V_1{@vertex1@}\) \(V_2{@vertex2@}\) </p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>center</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>[?,?]</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>focus1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>[?,?]</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>focus2</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>[?,?]</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>vertex1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>[?,?]</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>vertex2</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>[?,?]</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>center</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: checks correct foci*/
sans_set:{ans2,ans3};
tans_set:{focus1,focus2};</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>sans_set</sans>
        <tans>tans_set</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: checks correct vertice4*/
sans_set:{ans4,ans5};
tans_set:{vertex1,vertex2};</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>sans_set</sans>
        <tans>tans_set</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1147489516</deployedseed>
    <deployedseed>686038862</deployedseed>
    <deployedseed>1802775089</deployedseed>
    <deployedseed>388187177</deployedseed>
    <deployedseed>2056669373</deployedseed>
    <deployedseed>1577266983</deployedseed>
    <deployedseed>1549143971</deployedseed>
    <deployedseed>183489531</deployedseed>
    <deployedseed>657949994</deployedseed>
    <deployedseed>1085458015</deployedseed>
    <deployedseed>1497435061</deployedseed>
    <deployedseed>774301874</deployedseed>
    <deployedseed>446790540</deployedseed>
    <deployedseed>1126036923</deployedseed>
    <deployedseed>1049974015</deployedseed>
    <deployedseed>1282853382</deployedseed>
    <deployedseed>293772092</deployedseed>
    <deployedseed>1087951044</deployedseed>
    <deployedseed>1842179773</deployedseed>
    <deployedseed>1462557517</deployedseed>
    <deployedseed>1263590827</deployedseed>
    <deployedseed>827694419</deployedseed>
    <deployedseed>1260334700</deployedseed>
    <deployedseed>676383476</deployedseed>
    <deployedseed>1536802423</deployedseed>
    <deployedseed>264004269</deployedseed>
    <deployedseed>471225457</deployedseed>
    <deployedseed>1995541988</deployedseed>
    <deployedseed>651166798</deployedseed>
    <deployedseed>452824517</deployedseed>
    <deployedseed>47291013</deployedseed>
    <deployedseed>1181405041</deployedseed>
    <deployedseed>746885910</deployedseed>
    <deployedseed>1754016677</deployedseed>
    <deployedseed>1604681658</deployedseed>
    <deployedseed>690978459</deployedseed>
    <deployedseed>584429346</deployedseed>
    <deployedseed>1013272118</deployedseed>
    <deployedseed>58951786</deployedseed>
    <deployedseed>1597886424</deployedseed>
    <deployedseed>888776065</deployedseed>
    <deployedseed>251489080</deployedseed>
    <deployedseed>30731669</deployedseed>
    <deployedseed>1081747735</deployedseed>
    <deployedseed>918672148</deployedseed>
    <deployedseed>447770797</deployedseed>
    <deployedseed>1389883220</deployedseed>
    <deployedseed>803671588</deployedseed>
    <deployedseed>136209473</deployedseed>
    <deployedseed>1499616039</deployedseed>
    <deployedseed>729494022</deployedseed>
    <deployedseed>301251186</deployedseed>
    <deployedseed>1335681698</deployedseed>
    <deployedseed>112952236</deployedseed>
    <deployedseed>993724067</deployedseed>
    <deployedseed>1230371175</deployedseed>
    <deployedseed>937112419</deployedseed>
    <deployedseed>1047889115</deployedseed>
    <deployedseed>640708090</deployedseed>
    <deployedseed>1854048920</deployedseed>
    <deployedseed>497178889</deployedseed>
    <deployedseed>143178652</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>center</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>focus1</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>focus2</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>vertex1</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>vertex2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
