<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Solving Rational Equations</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Solve the rational equation \({@expr@}\)</p>
<p>\(x=\) [[input:ans1]] [[validation:ans1]]</p>
<p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked solution</h3>
<p>To solve the equation\({@expr@},\) we can first combine the fractions on the right-hand side by finding a common denominator.</p>
<p style="margin-left: 30px">\({@frac_1@}={@sol_frac_2@}\)</p>
<p>Then, we can cross-multiply and simplify.</p>
<p style="margin-left: 30px">\({@expr_1@}\)</p>
<p style="margin-left: 30px">\(x={@ta_2@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>aa:rand(5)+1;
bb:rand(4)+2;
cc:rand_with_prohib(-5,5,[aa,bb,0]);
frac_1:aa/x^2;/*first fraction*/
frac_2:bb/x-(cc/x^2);/*second fraction*/

expr:frac_1=frac_2;/*rational equation*/
sol_frac_2:ratsimp(frac_2);
expr_1:num(frac_1)=num(sol_frac_2);


ta_2:rhs(solve(expr_1,x)[1]);
ta_1:{ta_2};



/*possible mistakes*/

frac_22:bb/x+(cc/x^2);/*second fraction*/

expr1:frac_1=frac_22;/*rational equation*/
sol_frac_2_1:ratsimp(frac_22);
expr_11:num(frac_1)=num(sol_frac_2_1);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>expression \(= {@expr@}\)</p>
<p>answer \(= {@ta_2@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta_2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1 checks for correct answer*/
/*Node 2 mistake: Adding instead of subtracting. */
mistake_2:rhs(solve(expr_11,x)[1]);
mistake_1:mistake_2;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta_2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_1</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Ensure you check the signs between the fractions to compute the value of \(x\) correctly.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1709793413</deployedseed>
    <deployedseed>1962740691</deployedseed>
    <deployedseed>1707283379</deployedseed>
    <deployedseed>634883163</deployedseed>
    <deployedseed>1452988861</deployedseed>
    <deployedseed>1098286248</deployedseed>
    <deployedseed>596010456</deployedseed>
    <deployedseed>1123027273</deployedseed>
    <deployedseed>75857322</deployedseed>
    <deployedseed>2058458114</deployedseed>
    <deployedseed>1489109290</deployedseed>
    <deployedseed>2029737254</deployedseed>
    <deployedseed>592971927</deployedseed>
    <deployedseed>203895701</deployedseed>
    <deployedseed>759734583</deployedseed>
    <deployedseed>212098549</deployedseed>
    <deployedseed>1885193105</deployedseed>
    <deployedseed>1242537142</deployedseed>
    <deployedseed>102858923</deployedseed>
    <deployedseed>76606666</deployedseed>
    <deployedseed>269742404</deployedseed>
    <deployedseed>1739079390</deployedseed>
    <deployedseed>1835600944</deployedseed>
    <deployedseed>45548534</deployedseed>
    <deployedseed>2072439273</deployedseed>
    <deployedseed>1149086188</deployedseed>
    <deployedseed>1894741600</deployedseed>
    <deployedseed>124027448</deployedseed>
    <deployedseed>1183695175</deployedseed>
    <deployedseed>64037497</deployedseed>
    <deployedseed>594317977</deployedseed>
    <deployedseed>1368042633</deployedseed>
    <deployedseed>1024149895</deployedseed>
    <deployedseed>323575135</deployedseed>
    <deployedseed>529402852</deployedseed>
    <deployedseed>979578253</deployedseed>
    <deployedseed>1069766617</deployedseed>
    <deployedseed>194659721</deployedseed>
    <deployedseed>1834680386</deployedseed>
    <deployedseed>1374550469</deployedseed>
    <deployedseed>1058255795</deployedseed>
    <deployedseed>996492180</deployedseed>
    <deployedseed>774448398</deployedseed>
    <deployedseed>1252942943</deployedseed>
    <deployedseed>2146933148</deployedseed>
    <deployedseed>1291996956</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta_2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
