<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Determining Continuity of a Piecewise Function</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p></p><p>Determine if the following function is continuous at \(x=4\) and \(x=6\):</p><p></p>\[ f(x) =&nbsp;\begin{cases} {@f1@} &amp; x\leq 6 \\ {@f2@} &amp; x&nbsp;&gt; 6&nbsp; \end{cases} \]<p>State "yes" or "no" for each case:<br></p><p></p><p>At \(x=4\): [[input:ans1]]&nbsp;</p><p></p><p>At \(x=6\): [[input:ans2]]&nbsp;</p><p></p><p>[[validation:ans1]]</p><p>[[validation:ans2]]<br></p><p>[[feedback:prt1]]<br></p><p>[[feedback:prt2]]<br></p><p></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p><b>Worked Solution</b></p>
<p>At \(x=4\), the function is \({@f1@}\), which is itself continuous. Therefore, \(f(x)\) is continuous at \(x=4\).</p><p>To the left of \(x = 6\), the function is {@f1@}, therefore, \( \displaystyle \lim_{x\rightarrow 6^-} f(x) = {@aa@} \cdot 6 = {@aa*6@}\).</p><p>To the right of \(x = 6\), the function is {@f2@}, therefore, \( \displaystyle \lim_{x\rightarrow 6^+} f(x) = {@bb@} \cdot 6 {@sig@} {@abs(cc)@} = {@bb*6+cc@}\).<br></p><p>Therefore,&nbsp;\(f(x)\) is not continuous at \(x=6\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[aa:rand(2)+2;
bb:rand_with_prohib(2,3,[aa]);
cc:rand([-3,-2,-1,1,2,3]);
sig:if cc<0 then "-" else "+";
f1:aa*x;
f2:bb*x+cc;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\( f(x) =&nbsp;\begin{cases} {@f1@} &amp; x\leq 6 \\ {@f2@} &amp; x&nbsp;&gt; 6&nbsp; \end{cases} \)<p>At \(x=4\): "yes"</p><p>At \(x=6\): "no"</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>string</type>
      <tans><![CDATA["yes"]]></tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>string</type>
      <tans><![CDATA["no"]]></tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>StringSloppy</answertest>
        <sans>ans1</sans>
        <tans><![CDATA["yes"]]></tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>StringSloppy</answertest>
        <sans>ans2</sans>
        <tans><![CDATA["no"]]></tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1915378923</deployedseed>
    <deployedseed>1559311121</deployedseed>
    <deployedseed>3781452</deployedseed>
    <deployedseed>420000154</deployedseed>
    <deployedseed>1977673831</deployedseed>
    <deployedseed>1761495663</deployedseed>
    <deployedseed>1458592833</deployedseed>
    <deployedseed>902839739</deployedseed>
    <deployedseed>172638604</deployedseed>
    <deployedseed>1626564011</deployedseed>
    <deployedseed>94523104</deployedseed>
    <deployedseed>623602132</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value><![CDATA["yes"]]></value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value><![CDATA["no"]]></value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
