<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Complex Number Operations</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p> Given the complex numbers z = {@z@} and w = {@w@}, evaluate and write in the form \(a+bi\), where \(a\) and \(b\) are real numbers <strong>(to write the complex number \(i\) type %i, and √a type sqrt(a))</strong></p>
<p><b>(i)</b> \(z+w\) = [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>
<p><b>(ii)</b> \(z-w\) = [[input:ans2]] [[validation:ans2]][[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>Given complex numbers:  \(z={@z@}\) and  \(w={@w@}\).</p>

<p><b>(i)</b><b> \( z + w \)</b></p>
 <p> \( z + w = ({@z@}) + ({@w@}) \)</p>

 </p> Combining the real and imaginary parts:</p>
<p>   \(\begin{align} z + w &= ({@a@}+{@c@}) + ({@b@}+{@d@})i \\
   z + w &= {@ta1@} \end{align}\)</p>
 <p>  So, \( z + w \) in the form \( a + bi \) is \( {@ta1@}\).</p>

<br><b>(ii)</b><b> \( z - w \)</b>
   <p>\( z - w = ({@z@}) - ({@w@}) \)</p>
 <p>  Combining the real and imaginary parts:</p>
 <p>  \( \begin{align} z - w  &=({@a@}-{@c@}) +({@b@}-{@d@})i \\
    z - w  &= {@ta2@} \end{align}\)</p>
  <p> So, \( z - w \) in the form \( a + bi \) is \( {@ta2@} \).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>a: rand(20); 
b: 5-rand(20); 
c: rand(15); 
d: 3-rand(10);
z: a+b*%i; 
w: c+d*%i;
addzw: z+w;
ta1: addzw;
subzw: z-w;
ta2: subzw;
prodzw: ratsimp(z*w);
ta3: prodzw;
magz: abs(z);
ta4: magz;
magzw: abs(z+w);
ta5: magzw;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@z@} and {@w@}. {@z+w=addzw@}, {@z-w=subzw@} and {@z*w=prodzw@}.</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<p>Correct answer, well done.</p>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<p>Your answer is partially correct.</p>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<p>Incorrect answer.</p>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>9</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>9</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty>0</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1465205372</deployedseed>
    <deployedseed>2027968813</deployedseed>
    <deployedseed>771647625</deployedseed>
    <deployedseed>835522529</deployedseed>
    <deployedseed>1679644908</deployedseed>
    <deployedseed>1410014578</deployedseed>
    <deployedseed>978680352</deployedseed>
    <deployedseed>1750677084</deployedseed>
    <deployedseed>1711711585</deployedseed>
    <deployedseed>1845386216</deployedseed>
    <deployedseed>1285079161</deployedseed>
    <deployedseed>1865719403</deployedseed>
    <deployedseed>1367970145</deployedseed>
    <deployedseed>1353833554</deployedseed>
    <deployedseed>1362074397</deployedseed>
    <deployedseed>395425621</deployedseed>
    <deployedseed>2130680333</deployedseed>
    <deployedseed>762869877</deployedseed>
    <deployedseed>1959706979</deployedseed>
    <deployedseed>598944666</deployedseed>
    <deployedseed>397435285</deployedseed>
    <deployedseed>644717157</deployedseed>
    <deployedseed>1545511710</deployedseed>
    <deployedseed>34888531</deployedseed>
    <deployedseed>208452657</deployedseed>
    <deployedseed>1729309609</deployedseed>
    <deployedseed>1262618251</deployedseed>
    <deployedseed>419216924</deployedseed>
    <deployedseed>494103277</deployedseed>
    <deployedseed>223075666</deployedseed>
    <deployedseed>1207316023</deployedseed>
    <deployedseed>69533595</deployedseed>
    <deployedseed>1471209892</deployedseed>
    <deployedseed>1209695612</deployedseed>
    <deployedseed>1535099554</deployedseed>
    <deployedseed>1709466054</deployedseed>
    <deployedseed>1132556958</deployedseed>
    <deployedseed>1523661678</deployedseed>
    <deployedseed>1270171847</deployedseed>
    <deployedseed>1939841818</deployedseed>
    <deployedseed>576939210</deployedseed>
    <deployedseed>795433428</deployedseed>
    <deployedseed>458931843</deployedseed>
    <deployedseed>1618608834</deployedseed>
    <deployedseed>558534571</deployedseed>
    <deployedseed>1105355762</deployedseed>
    <deployedseed>811233886</deployedseed>
    <deployedseed>8191866</deployedseed>
    <deployedseed>893305637</deployedseed>
    <deployedseed>1265754092</deployedseed>
    <deployedseed>1279627302</deployedseed>
    <deployedseed>1243484153</deployedseed>
    <deployedseed>1468703669</deployedseed>
    <deployedseed>2041875286</deployedseed>
    <deployedseed>1736711789</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
