<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Exponent Simplification and Writing in the Form \(a^n * b^m\)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Write \(\frac{{@term1@}^{{@cc@}} \cdot {@term2@}^{{@dd@}} \cdot {@term3@}}{{@term4@}^{-{@ee@}} \cdot {@term5@}^{{@ff@}}}\) in the form \({@var1@}^n \cdot {@var2@}^m\), where \(n,m\in \mathbb{Z}\)</p>

<p>[[input:ans1]] [[validation:ans1]]</p>
<p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution </h3>

<p>We write out all of the terms in the expression as powers of \(3\) and \(4\), and then simplify using the standard laws of incidies.</p>

<p>\(\begin{align*}
\frac{{@term1@}^{{@cc@}}\cdot {@term2@}^{{@dd@}}\cdot {@term3@}}{{@term4@}^{-{@ee@}}\cdot {@term5@}^{{@ff@}}} &= \frac{({@var1@}^{{@aa@}})^{{@cc@}}\cdot ({@var2@}^{{@bb@}})^{{@dd@}}\cdot {@var1@} \cdot {@var2@}}{{@var1@}^{-{@ee@}}\cdot {@var2@}^{{@ff@}}} \\

&= \frac{{@var1@}^{{@aa*cc@}}\cdot {@var2@}^{{@bb*dd@}}\cdot {@var1@}^1 \cdot {@var2@}^1}{{@var1@}^{-{@ee@}}\cdot {@var2@}^{{@ff@}}} \\

&= \frac{{@var1@}^{{@ev(aa*cc,simp)@}}\cdot {@var2@}^{{@ev(bb*dd,simp)@}}\cdot {@var1@}^1 \cdot {@var2@}^1}{{@var1@}^{-{@ee@}}\cdot {@var2@}^{{@ff@}}} \\

&= {@var1@}^{{@ev(aa*cc,simp)@}+1-(-{@ee@})} \cdot {@var2@}^{{@ev(bb*dd,simp)+1-ff@}} \\

&= {@var1@}^{{@pwr1@}} \cdot {@var2@}^{{@pwr2@}}

\end{align*}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>simp:false;
var1:rand([3,5]);
var2:rand([2,4])
aa:ev(rand(3)+2,simp);
bb:ev(rand(3)+2,simp);
cc:ev(rand(3)+2,simp);
dd:ev(rand(3)+2,simp);
ee:ev(rand(3)+2,simp);
ff:ev(rand(3)+8,simp);
term1:ev(var1^aa,simp);
term2:ev(var2^bb,simp);
term3:ev(var1*var2,simp);
term4:var1;
term5:var2;
pwr1:ev(aa*cc+1+ee,simp);
pwr2:ev(bb*dd+1-ff,simp);
ta1:var1^pwr1*var2^pwr2;

/*prt to check if the student simplifies the entire computation into a single digit rather than writting the answer in index form..*/
node2: simplify(ta1);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>random bases: {@var1@} and {@var2@}, question: \(\frac{({@var1@}^a)^c\cdot ({@var2@}^b)^d \cdot {@var1@}\cdot {@var2@}}{{@var1@}^{-e} \cdot {@var2@}^f}\), random [a,b,c,d,e,f] = [{@aa@},{@bb@},{@cc@},{@dd@},{@ee@},{@ff@}] </p>

<p>answer = \({@var1@}^x\cdot {@var2@}^y = {@ta1@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>0</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>EqualComAss</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>node2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<br>

<p>Good effort! but please remember to write the answer in <b>index form</b>. In this case, it should be \({@ta1@}\) and not \({@node2@}\).</p>
<p>Keep practicing and pay attention to the correct format.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1023823157</deployedseed>
    <deployedseed>79637087</deployedseed>
    <deployedseed>2095899117</deployedseed>
    <deployedseed>1441275504</deployedseed>
    <deployedseed>2147143528</deployedseed>
    <deployedseed>189748492</deployedseed>
    <deployedseed>1897931056</deployedseed>
    <deployedseed>1821840491</deployedseed>
    <deployedseed>1155862938</deployedseed>
    <deployedseed>1602650851</deployedseed>
    <deployedseed>2042676545</deployedseed>
    <deployedseed>666806727</deployedseed>
    <deployedseed>1669355556</deployedseed>
    <deployedseed>808312908</deployedseed>
    <deployedseed>263285500</deployedseed>
    <deployedseed>331479434</deployedseed>
    <deployedseed>1122542368</deployedseed>
    <deployedseed>1161788747</deployedseed>
    <deployedseed>1779996154</deployedseed>
    <deployedseed>124736529</deployedseed>
    <deployedseed>965424388</deployedseed>
    <deployedseed>1562028742</deployedseed>
    <deployedseed>1426659606</deployedseed>
    <deployedseed>2079185088</deployedseed>
    <deployedseed>1377827236</deployedseed>
    <deployedseed>32138703</deployedseed>
    <deployedseed>304532401</deployedseed>
    <deployedseed>210636437</deployedseed>
    <deployedseed>1877549170</deployedseed>
    <deployedseed>1704248525</deployedseed>
    <deployedseed>1637918353</deployedseed>
    <deployedseed>403273615</deployedseed>
    <deployedseed>1845549121</deployedseed>
    <deployedseed>103036752</deployedseed>
    <deployedseed>2142888389</deployedseed>
    <deployedseed>1437230398</deployedseed>
    <deployedseed>1348910855</deployedseed>
    <deployedseed>151250768</deployedseed>
    <deployedseed>1164584128</deployedseed>
    <deployedseed>1660515958</deployedseed>
    <deployedseed>1604146055</deployedseed>
    <deployedseed>747840785</deployedseed>
    <deployedseed>1094954173</deployedseed>
    <deployedseed>1524090497</deployedseed>
    <deployedseed>1885507965</deployedseed>
    <deployedseed>1524939586</deployedseed>
    <deployedseed>1746478381</deployedseed>
    <deployedseed>474271063</deployedseed>
    <deployedseed>1624951608</deployedseed>
    <deployedseed>122665774</deployedseed>
    <deployedseed>855578648</deployedseed>
    <deployedseed>786550909</deployedseed>
    <deployedseed>1668448444</deployedseed>
    <deployedseed>1187504853</deployedseed>
    <deployedseed>140477670</deployedseed>
    <deployedseed>26166572</deployedseed>
    <deployedseed>1151059886</deployedseed>
    <deployedseed>516306822</deployedseed>
    <deployedseed>1328046056</deployedseed>
    <deployedseed>1811872962</deployedseed>
    <deployedseed>325067994</deployedseed>
    <deployedseed>121899222</deployedseed>
    <deployedseed>721519110</deployedseed>
    <deployedseed>1249679487</deployedseed>
    <deployedseed>1905236633</deployedseed>
    <deployedseed>2098790566</deployedseed>
    <deployedseed>606872052</deployedseed>
    <deployedseed>1116163811</deployedseed>
    <deployedseed>909637566</deployedseed>
    <deployedseed>1847157625</deployedseed>
    <deployedseed>320803352</deployedseed>
    <deployedseed>37477017</deployedseed>
    <deployedseed>1721846934</deployedseed>
    <deployedseed>474990480</deployedseed>
    <deployedseed>1725572802</deployedseed>
    <deployedseed>1763416710</deployedseed>
    <deployedseed>1777348910</deployedseed>
    <deployedseed>1907168352</deployedseed>
    <deployedseed>318222883</deployedseed>
    <deployedseed>679048459</deployedseed>
    <deployedseed>2071281368</deployedseed>
    <deployedseed>1659835972</deployedseed>
    <deployedseed>350302124</deployedseed>
    <deployedseed>879283228</deployedseed>
    <deployedseed>668143608</deployedseed>
    <deployedseed>1906663504</deployedseed>
    <deployedseed>540858320</deployedseed>
    <deployedseed>305700900</deployedseed>
    <deployedseed>1008713312</deployedseed>
    <deployedseed>398500792</deployedseed>
    <deployedseed>1063081208</deployedseed>
    <deployedseed>338142652</deployedseed>
    <deployedseed>1658694644</deployedseed>
    <deployedseed>2026719239</deployedseed>
    <deployedseed>1272674177</deployedseed>
    <deployedseed>668157500</deployedseed>
    <deployedseed>2025455614</deployedseed>
    <deployedseed>390548884</deployedseed>
    <deployedseed>1867538469</deployedseed>
    <deployedseed>1996556672</deployedseed>
    <deployedseed>1204515212</deployedseed>
    <deployedseed>2028942927</deployedseed>
    <deployedseed>418219099</deployedseed>
    <deployedseed>21001644</deployedseed>
    <deployedseed>1665467111</deployedseed>
    <deployedseed>546842515</deployedseed>
    <deployedseed>496891471</deployedseed>
    <deployedseed>1105791498</deployedseed>
    <deployedseed>1169599848</deployedseed>
    <deployedseed>100545267</deployedseed>
    <deployedseed>1424726360</deployedseed>
    <deployedseed>1083374690</deployedseed>
    <deployedseed>381170123</deployedseed>
    <deployedseed>1712084765</deployedseed>
    <deployedseed>1400433739</deployedseed>
    <deployedseed>606133042</deployedseed>
    <deployedseed>1242165572</deployedseed>
    <deployedseed>1353340092</deployedseed>
    <deployedseed>2058201210</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>node2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
