<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Solving the system of linear equations using the method of Gaussian Elimination Test (Guided Question)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Solve the following system of equations using the method of row operations and matrix transformations:</p>

<p>\( \left \{ \begin{array}{ll}
\displaystyle {@A8[1,1]*x@} + {@A8[1,2]*y@}+ {@A8[1,3]*z@} &= {@A8[1,4]@}\\
\displaystyle  {@A8[2,1]*x@} + {@A8[2,2]*y@}+ {@A8[2,3]*z@} &= {@A8[2,4]@}\\
\displaystyle  {@A8[3,1]*x@} + {@A8[3,2]*y@}+ {@A8[3,3]*z@} &= {@A8[3,4]@}
\end{array}\right.
\)</p>
<p><em>Note: Fill in the blanks in the intermediate matrix after each step.</em></p>

<br>

<p>
<ol>
<li>Begin by writing down the augmented matrix for the given system of equations:</li>
<p> [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>

<br>

<li>Taking the augmented matrix you have written above, apply row operations to transform the matrix into its rref by following each of the following steps:</li>
<br>
<p>Step 1: \(R_1 \leftarrow {@1/mult_const[3]@}\times R_1 \)</p>
<p>[[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</p>

<p>Step 2: \(\displaystyle  R_2 \leftarrow R_2{@-add_const[6]@}R_1\)</p>
<p>[[input:ans4]] [[validation:ans4]] [[feedback:prt4]]<p>

<p>Step 3: \( \displaystyle  R_3 \leftarrow R_3{@-add_const[5]@}R_1\)</p>
<p>[[input:ans5]] [[validation:ans5]] [[feedback:prt5]]<p>

<p>Step 4: \(\displaystyle   R_2 \leftarrow {@1/mult_const[2]@}R_2 \)</p>
<p>[[input:ans6]] [[validation:ans6]] [[feedback:prt6]]<p>

<p>Step 5: \(\displaystyle  R_1 \leftarrow R_1{@-add_const[4]@}R_2  \)</p>
<p>[[input:ans7]] [[validation:ans7]] [[feedback:prt7]]<p>

<p>Step 6: \( \displaystyle  R_3 \leftarrow R_3{@-add_const[3]@}R_2 \)</p>
<p>[[input:ans8]] [[validation:ans8]] [[feedback:prt8]]<p>

<p>Step 7: \(\displaystyle  R_3 \leftarrow {@1/mult_const[1]@}R_3 \)</p>
<p>[[input:ans9]] [[validation:ans9]] [[feedback:prt9]]<p>

<p>Step 8: \(\displaystyle  R_1 \leftarrow R_1{@-add_const[2]@}R_3 \)</p>
<p>[[input:ans10]] [[validation:ans10]] [[feedback:prt10]]<p>

<p>Step 9: \(\displaystyle  R_2 \leftarrow R_2{@-add_const[1]@}R_3 \)</p>
<p>[[input:ans2]] [[validation:ans2]] [[feedback:prt2]]<p>

<br>

<li>Write the values of \(x, y, z\) consecutively in column form; </li>
<p>[[input:ans11]] [[validation:ans11]] [[feedback:prt11]]<p>
<p></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>We start from the following system of equations.</p>
<p>\( \left \{ \begin{array}{rc}
\displaystyle  {@A8[1,1]*x@} + {@A8[1,2]*y@}+ {@A8[1,3]*z@} &= {@A8[1,4]@}\\
\displaystyle  {@A8[2,1]*x@} + {@A8[2,2]*y@}+ {@A8[2,3]*z@} &= {@A8[2,4]@}\\
 \displaystyle  {@A8[3,1]*x@} + {@A8[3,2]*y@}+ {@A8[3,3]*z@} &= {@A8[3,4]@}
\end{array}\right.
\)</p>

<p>The Augmented matrix will be:  {@A8@}</p>

<br>

<p> We'll take the matrix from the previous step beginning with the augmented matrix above and perform the following operations.</p>
<p style="margin-left:30px">\(\begin{align*}
\textbf{Step 1} && \displaystyle R_1 &\leftarrow {@1/mult_const[3]@}\times R_1: && {@A9@} \\

\textbf{Step 2} &&  \displaystyle   R_2 &\leftarrow  R_2{@-add_const[6]@}R_1: && {@A10@} \\

\textbf{Step 3} && \displaystyle  R_3 &\leftarrow R_3{@-add_const[5]@}R_1 : && {@A11@} \\

\textbf{Step 4} && \displaystyle  R_2 &\leftarrow {@1/mult_const[2]@}R_2 : && {@A12@} \\

\textbf{Step 5} && \displaystyle   R_1 &\leftarrow R_1{@-add_const[4]@}R_2 : && {@A13@}\\

\textbf{Step 6} && \displaystyle  R_3 &\leftarrow R_3{@-add_const[3]@}R_2 : && {@A14@}\\

\textbf{Step 7} && \displaystyle  R_3 &\leftarrow {@1/mult_const[1]@}R_3 : && {@A15@} \\

\textbf{Step 8} && \displaystyle  R_1 &\leftarrow R_1{@-add_const[2]@}R_3 : && {@A16@}\\

\textbf{Step 9} & & \displaystyle  R_2 &\leftarrow R_2{@-add_const[1]@}R_3 : && {@A17@} 
\end{align*}\)</p>

<br>

<p>Therefore, the solutions of the system will be \(\displaystyle  x= {@ans_x@}, y={@ans_y@}, z={@ans_z@} \)</p>]]></text>
    </generalfeedback>
    <defaultgrade>2</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/* Starting from the answer- rref matrix; solution matrix */
ans_x:rand(10)-5;
ans_y:rand(10)-5;
ans_z:rand(10)-5;
rrefA: matrix(
    [1,0,0,ans_x],
    [0,1,0,ans_y],
    [0,0,1,ans_z]
);
ta2:rrefA;

/* ----- FUNCTIONS*/
/* multiplies a specific row in the given matrix by a constant value. */
row_times_const(matrix, row_no, const) := block(
    A:copymatrix(matrix),
    A[row_no]: const*A[row_no],
    return(A)
);

/* swaps two rows in the given matrix; row_no1 and row_no2 in matrix */
swap_rows(matrix, row_no1, row_no2) := block(
    A:copymatrix(matrix),
    A[row_no1]: matrix[row_no2],
    A[row_no2]: matrix[row_no1],
    return(A)
);

/* Add a multiple of one row to another row in the matrix.*/
sum_rows(matrix, row_no1, row_no2, const) := block(
    A:copymatrix(matrix),
    A[row_no1]: matrix[row_no1] + const*matrix[row_no2],
    return(A)
);
/* ----- END FUNCTIONS */

/* manipulation constants */
mult_const: random_permutation([rand(2)+2, rand(2)+2, rand(2)+2]);
add_const: append(
    random_permutation([1,1,rand(3)+1,rand(4)+1]), 
    [rand(3)+1, rand(3)+1]
);

/* Going from the solution matrix in RREF to the question matrix */
A1: sum_rows(rrefA, 2, 3, add_const[1]);
A2: sum_rows(A1, 1, 3, add_const[2]);
A3: row_times_const(A2, 3, mult_const[1]);

A3: sum_rows(A3, 3, 2, add_const[3]);
A4: sum_rows(A, 1, 2, add_const[4]);
A5: row_times_const(A4, 2, mult_const[2]);

A6: sum_rows(A5, 3, 1, add_const[5]);
A7: sum_rows(A6, 2, 1, add_const[6]);
A8: row_times_const(A7, 1, mult_const[3]);
ta1:A8;

/* From the question matrix, undoing each row operation to get to the RREF */
A9: row_times_const(A8, 1, 1/mult_const[3]); /*step 1*/
A10: sum_rows(A9, 2, 1, -add_const[6]); /*step 2*/
A11: sum_rows(A10, 3, 1, -add_const[5]); /*step 3*/

A12: row_times_const(A11, 2, 1/mult_const[2]); /*step 4*/
A13: sum_rows(A12, 1, 2, -add_const[4]); /*step 5*/
A14: sum_rows(A13, 3, 2, -add_const[3]); /*step 6*/

A15: row_times_const(A14, 3, 1/mult_const[1]); /*step 7*/
A16: sum_rows(A15, 1, 3, -add_const[2]); /*step 8*/
A17: sum_rows(A16, 2, 3, -add_const[1]); /*step 9*/

A18:  matrix(
    [ans_x],
    [ans_y],
    [ans_z]); /*Isolate the last column*/</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<br>Systems 
<p>\(
\begin{align*}
{@A8[1,1]*x@} + {@A8[1,2]*y@}+ {@A8[1,3]*z@} &= {@A8[1,4]@}\\
{@A8[2,1]*x@} + {@A8[2,2]*y@}+ {@A8[2,3]*z@} &= {@A8[2,4]@}\\
{@A8[3,1]*x@} + {@A8[3,2]*y@}+ {@A8[3,3]*z@} &= {@A8[3,4]@}
\end{align*}
\)</p>
<br>Matrix {@A@}
<br>RREF is {@ta2@}]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>matrix</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans10</name>
      <type>matrix</type>
      <tans>A16</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans11</name>
      <type>matrix</type>
      <tans>A18</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>matrix</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>matrix</type>
      <tans>A9</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>matrix</type>
      <tans>A10</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>matrix</type>
      <tans>A11</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6</name>
      <type>matrix</type>
      <tans>A12</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans7</name>
      <type>matrix</type>
      <tans>A13</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans8</name>
      <type>matrix</type>
      <tans>A14</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans9</name>
      <type>matrix</type>
      <tans>A15</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt10</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>A16_1: sum_rows(ans9, 1, 3, -add_const[2]); /*step 8*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans10</sans>
        <tans>A16</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt10-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt10-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans10</sans>
        <tans>A16_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt10-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt10-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt11</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans11</sans>
        <tans>A18</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt11-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt11-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>echelon(A)</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Finding the echelon form of a matrix is a precursor step to finding the Reduced Row Echelon Form (RREF).</p>
<p>In this case you should find the Reduced Row Echelon Form (RREF) of the matrix.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 2: Check for correct row operation on wrong ans 1*/
A9_1: row_times_const(ans1, 1, 1/mult_const[3]); /*step 1*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>A9</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>A9_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Ensure you are applying the operations to the right matrix. You seem to have computed the right row operation the previous input matrix. </p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 2: Check for correct row operation on wrong ans 3*/
A10_1: sum_rows(ans3, 2, 1, -add_const[6]); /*step 2*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>A10</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>A10_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Ensure you are applying the operations to the right matrix. However, what you have done here is correct.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Check for correct row operation on wrong ans5*/
A11_1: sum_rows(ans4, 3, 1, -add_const[5]); /*step 3*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>A11</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>A11_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Ensure you are applying the operations to the right matrix. However, what you have done here is correct.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt6</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Check for correct row operation on wrong ans6*/
A12_1: row_times_const(ans5, 2, 1/mult_const[2]); /*step 4*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans6</sans>
        <tans>A12</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt6-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans6</sans>
        <tans>A12_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Ensure you are applying the operations to the right matrix. However, what you have done here is correct.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt6-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt7</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>A13_1: sum_rows(ans6, 1, 2, -add_const[4]); /*step 5*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans7</sans>
        <tans>A13</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt7-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans7</sans>
        <tans>A13_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Ensure you are applying the operations to the right matrix. However, what you have done here is correct.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt7-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt8</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>A14_1: sum_rows(ans7, 3, 2, -add_const[3]); /*step 6*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans8</sans>
        <tans>A14</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt8-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt8-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans8</sans>
        <tans>A14_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt8-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt8-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt9</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>A15_1: row_times_const(ans8, 3, 1/mult_const[1]); /*step 7*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans9</sans>
        <tans>A15</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt9-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt9-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans9</sans>
        <tans>A15_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt9-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Ensure you are applying the operations to the right matrix. However, what you have done here is correct.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt9-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>450810240</deployedseed>
    <deployedseed>1433062839</deployedseed>
    <deployedseed>1456058640</deployedseed>
    <deployedseed>304663023</deployedseed>
    <deployedseed>715164182</deployedseed>
    <deployedseed>7655779</deployedseed>
    <deployedseed>1785766223</deployedseed>
    <deployedseed>1777496000</deployedseed>
    <deployedseed>252573374</deployedseed>
    <deployedseed>709722043</deployedseed>
    <deployedseed>1511534909</deployedseed>
    <deployedseed>1999961629</deployedseed>
    <deployedseed>321273679</deployedseed>
    <deployedseed>968448296</deployedseed>
    <deployedseed>677811195</deployedseed>
    <deployedseed>286719133</deployedseed>
    <deployedseed>1160654531</deployedseed>
    <deployedseed>1520589653</deployedseed>
    <deployedseed>475472611</deployedseed>
    <deployedseed>1574353359</deployedseed>
    <deployedseed>163962000</deployedseed>
    <deployedseed>273060524</deployedseed>
    <deployedseed>32801483</deployedseed>
    <deployedseed>1345774629</deployedseed>
    <deployedseed>1765729531</deployedseed>
    <deployedseed>375932214</deployedseed>
    <deployedseed>1341333361</deployedseed>
    <deployedseed>600600625</deployedseed>
    <deployedseed>245331930</deployedseed>
    <deployedseed>543460616</deployedseed>
    <deployedseed>285589254</deployedseed>
    <deployedseed>771657998</deployedseed>
    <deployedseed>696608401</deployedseed>
    <deployedseed>1724500074</deployedseed>
    <deployedseed>2089502629</deployedseed>
    <deployedseed>744336036</deployedseed>
    <deployedseed>820623409</deployedseed>
    <deployedseed>726709021</deployedseed>
    <deployedseed>1464950635</deployedseed>
    <deployedseed>2081674067</deployedseed>
    <deployedseed>370878676</deployedseed>
    <deployedseed>807487540</deployedseed>
    <deployedseed>1819966977</deployedseed>
    <deployedseed>1968630887</deployedseed>
    <deployedseed>1280173275</deployedseed>
    <deployedseed>2012374471</deployedseed>
    <deployedseed>1284252472</deployedseed>
    <deployedseed>19342960</deployedseed>
    <deployedseed>1205015128</deployedseed>
    <deployedseed>1829773014</deployedseed>
    <deployedseed>2037241525</deployedseed>
    <deployedseed>1102944234</deployedseed>
    <deployedseed>430766031</deployedseed>
    <deployedseed>1549459618</deployedseed>
    <deployedseed>1752986744</deployedseed>
    <deployedseed>173992583</deployedseed>
    <deployedseed>892650178</deployedseed>
    <deployedseed>138227923</deployedseed>
    <deployedseed>1994763406</deployedseed>
    <deployedseed>1775961359</deployedseed>
    <deployedseed>2032566193</deployedseed>
    <deployedseed>1539169647</deployedseed>
    <deployedseed>1476078765</deployedseed>
    <deployedseed>986281431</deployedseed>
    <deployedseed>1520237877</deployedseed>
    <deployedseed>779224217</deployedseed>
    <deployedseed>1380477256</deployedseed>
    <deployedseed>931911671</deployedseed>
    <deployedseed>1785186537</deployedseed>
    <deployedseed>86115361</deployedseed>
    <deployedseed>1024417869</deployedseed>
    <deployedseed>563871568</deployedseed>
    <deployedseed>1106218009</deployedseed>
    <deployedseed>1149891887</deployedseed>
    <deployedseed>71040048</deployedseed>
    <deployedseed>1579920579</deployedseed>
    <deployedseed>1522636757</deployedseed>
    <deployedseed>1767427347</deployedseed>
    <deployedseed>673377835</deployedseed>
    <deployedseed>759989984</deployedseed>
    <deployedseed>1576048929</deployedseed>
    <deployedseed>615246075</deployedseed>
    <deployedseed>1433456197</deployedseed>
    <deployedseed>1189779275</deployedseed>
    <deployedseed>585658185</deployedseed>
    <deployedseed>1886013441</deployedseed>
    <deployedseed>384283745</deployedseed>
    <deployedseed>761468959</deployedseed>
    <deployedseed>749806918</deployedseed>
    <deployedseed>1653334166</deployedseed>
    <deployedseed>727910890</deployedseed>
    <deployedseed>1272591669</deployedseed>
    <deployedseed>456613116</deployedseed>
    <deployedseed>1871929332</deployedseed>
    <deployedseed>146184351</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
