<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Solve the logarithmic equation, giving your answer in exact form</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Solve the following logarithmic equation, giving your answer in exact form in terms of \(e\).</p> 
<p>\(\ln({@x-aa@})-\ln(x) = {@bb@}\)</p>
<p><em>Note: \(e\) is entered as %e.</em></p>
<p>[[input:ans1]] [[validation:ans1]]</p>
<p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution:</h3>
<p>To solve the equation, we first combine the two logarithm functions using the standard laws of logarithms.
<p style="margin-left: 30px">\(\ln({@x-aa@}) - \ln(x) = {@bb@}\)</p>
<p style="margin-left: 30px">\(\ln \left({@(x-aa)/x@}\right) = {@bb@}\)</p>

<p><br>We may now apply the exponential function to both sides. In other words, we raise \(e\) to the power of both sides, which cancels out the logarithm funciton.</p>
<p style="margin-left: 30px">\({@(x-aa)/x@} = {@%e^bb@}\)</p>
<p style="margin-left: 30px">\({@x-aa@} = {@%e^bb@} x\)</p>

<p><br>Finally, we solve the resulting linear equation.</p>
<p style="margin-left: 30px">\({@x-%e^bb*x@} = {@aa@}\)</p>
<p style="margin-left: 30px">\(x ({@1-e^bb@}) = {@aa@}\)</p>
<p style="margin-left: 30px">\(\displaystyle x = \frac{{@aa@}}{{@1-%e^bb@}}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[pm1:rand([-1,1]);
aa:pm1*(rand(8)+1);
bb:rand_with_prohib(-8,8,[aa,0,-1,1]);
pm2:if bb<0 then -1 else 1;
ins1:x-aa;
ins2:x;
ta1:aa/(1-%e^(bb));]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Logarithmic equation \(\ln({@x-aa@})-\ln(x) = {@bb@}\)</p>
<p>Solution \(={@ta1@}\) </p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/* Node 1: Checks for the correct ta */
/*Node 2: Checks if the wrong student answer is 1*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1189850687</deployedseed>
    <deployedseed>78880471</deployedseed>
    <deployedseed>1187674644</deployedseed>
    <deployedseed>1806731727</deployedseed>
    <deployedseed>1986012027</deployedseed>
    <deployedseed>1944899933</deployedseed>
    <deployedseed>1263395384</deployedseed>
    <deployedseed>1904060465</deployedseed>
    <deployedseed>1034805624</deployedseed>
    <deployedseed>711344636</deployedseed>
    <deployedseed>1578240573</deployedseed>
    <deployedseed>353477210</deployedseed>
    <deployedseed>565821318</deployedseed>
    <deployedseed>1326997913</deployedseed>
    <deployedseed>1612837088</deployedseed>
    <deployedseed>514796353</deployedseed>
    <deployedseed>84899003</deployedseed>
    <deployedseed>1950322317</deployedseed>
    <deployedseed>1782357329</deployedseed>
    <deployedseed>2022712406</deployedseed>
    <deployedseed>1802149537</deployedseed>
    <deployedseed>324209285</deployedseed>
    <deployedseed>1293884713</deployedseed>
    <deployedseed>774693272</deployedseed>
    <deployedseed>607293813</deployedseed>
    <deployedseed>1320111979</deployedseed>
    <deployedseed>937749477</deployedseed>
    <deployedseed>1507431217</deployedseed>
    <deployedseed>1950753900</deployedseed>
    <deployedseed>2103704899</deployedseed>
    <deployedseed>1623766469</deployedseed>
    <deployedseed>1156587595</deployedseed>
    <deployedseed>1470495245</deployedseed>
    <deployedseed>1833232794</deployedseed>
    <deployedseed>355773093</deployedseed>
    <deployedseed>1150671408</deployedseed>
    <deployedseed>2022649975</deployedseed>
    <deployedseed>1850401343</deployedseed>
    <deployedseed>2132566830</deployedseed>
    <deployedseed>1464069239</deployedseed>
    <deployedseed>1789514167</deployedseed>
    <deployedseed>573063963</deployedseed>
    <deployedseed>1495814619</deployedseed>
    <deployedseed>289030035</deployedseed>
    <deployedseed>1636799799</deployedseed>
    <deployedseed>1691428153</deployedseed>
    <deployedseed>2085645700</deployedseed>
    <deployedseed>552861223</deployedseed>
    <deployedseed>1358276219</deployedseed>
    <deployedseed>2081579504</deployedseed>
    <deployedseed>197698440</deployedseed>
    <deployedseed>960164671</deployedseed>
    <deployedseed>12699359</deployedseed>
    <deployedseed>1839740254</deployedseed>
    <deployedseed>1601864952</deployedseed>
    <deployedseed>888548493</deployedseed>
    <deployedseed>1230208271</deployedseed>
    <deployedseed>575319244</deployedseed>
    <deployedseed>2084011690</deployedseed>
    <deployedseed>1088816698</deployedseed>
    <deployedseed>1662337072</deployedseed>
    <deployedseed>105907366</deployedseed>
    <deployedseed>1775952787</deployedseed>
    <deployedseed>917195623</deployedseed>
    <deployedseed>802943177</deployedseed>
    <deployedseed>596059798</deployedseed>
    <deployedseed>2086872327</deployedseed>
    <deployedseed>1584726670</deployedseed>
    <deployedseed>38511228</deployedseed>
    <deployedseed>82557517</deployedseed>
    <deployedseed>225335205</deployedseed>
    <deployedseed>470897888</deployedseed>
    <deployedseed>2075271242</deployedseed>
    <deployedseed>958890919</deployedseed>
    <deployedseed>209017781</deployedseed>
    <deployedseed>690792200</deployedseed>
    <deployedseed>1511311858</deployedseed>
    <deployedseed>961002989</deployedseed>
    <deployedseed>2005811096</deployedseed>
    <deployedseed>1188773726</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
