<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Modulus and Argument of Complex Numbers and Their Conjugates</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>1. Find the modulus \(r\) and argument \(\theta \) of \({@aa@}{@sigb@}{@bbi@}\).</p> 
<p>Give your answers in exact form (in terms of square roots and \(\pi\). Recall that \(\sqrt{a}\) is entered as sqrt(a) and \(\pi\) is entered as %pi.</p>
<p>\(r = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>
<p>\(\theta = \) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p>
<p>2. Find the modulus \(q\) and argument \(\delta \) of the complex conjugate of \({@aa@}{@sigb@}{@bbi@}\).</p> 
<p>Give your answers in exact form.</p>
<p>\(q = \) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</p>
<p>\(\delta = \) [[input:ans4]] [[validation:ans4]] [[feedback:prt4]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>1. The complex number in question is \({@aa@}{@sigb@}{@bbi@}\). If we consider it in the form \(a + b\cdot i\), then \(a = {@aa@}\) and \(b = {@bb/i@}\). So, the modulus of \({@aa@}{@sigb@}{@bbi@}\) is</p>
<p>\(r = \left| {@aa@}{@sigb@}{@bbi@}\right| = \sqrt{a^2 + b^2}\)</p>
<p>\( = \sqrt{\left( {@aa@}\right) ^2 + \left( {@bb/i@}\right) ^2}\)</p>
<p>\( = \sqrt{ {@aa^2@} + {@(bb/i)^2@}}\)</p>
<p>\( = \sqrt{{@aa^2+(bb/i)^2@}}\)</p>
<p>\( = {@ta1@}\)</p>
<br>
<p>The complex number in question is \({@aa@}{@sigb@}{@bbi@}\). If we consider it in the form \(a + b\cdot i\), then \(a = {@aa@}\) and \(b = {@bb/i@}\).</p> 
<p>We first note that, in the Argand Diagram, \({@aa@}{@sigb@}{@bbi@}\) lies in quadrant number {@quad@}, which means that the argument \(\theta \) is within the range \({@lb@}&lt;\theta &lt; {@ub@}\).</p>
So, \(\theta = \tan^{-1}\left( \frac{{@abs(bb/i)@}}{{@abs(aa)@}}\right) \)<p></p>
<p> \(= \tan^{-1}\left( {@abs(bb/i)/abs(aa)@}\right) \)</p>
<p> \(= {@atan(abs(bb/i)/aa)@}\).</p>
<p>2. So, considering the symmetry of the unit circle, for \({@aa@}{@sigb@}{@bbi@}\), \(\theta = {@ta2@}\).</p>
<br>
<p>The modulus of a complex number and its conjugate are the same as the distance from the origin in the argand diagram will not change. So, \(q = r = {@ta3@}\).</p>
<br>
<p>In the Argand Diagram, the complex conjugate of a number is its reflection on the \(x-\)axis. Therefore, its argument is the negative of the argument of the original number. In this case, \(\delta = -\theta = {@ta4@}\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[tt:rand([3,4,6]);
dumset:{};
for i:1 thru 2*tt do if gcd(i,tt)=1 then dumset:adjoin(i,dumset);
dumlist:listify(dumset);
pp:rand(dumlist);
arg:pp*%pi/tt;
mod:if tt=4 then rand([1,2,6,10,sqrt(2),2*sqrt(2),3*sqrt(2)/*,sqrt(6),sqrt(3)*/]) else rand([1,2,4,5,7,sqrt(3),2*sqrt(3),4*sqrt(3)/*,sqrt(2),sqrt(6)*/]);
cc:cos(arg);
ss:sin(arg);
aa:mod*cc;
bb:mod*ss*i;
sigb:if mod*ss<0 then "-" else "+";
bbi:if cabs(bb)=i then "i" else cabs(bb)*i;
ta1:mod;
ta2:if arg>%pi then arg-2*%pi else arg;
ta3:mod;
ta4:-ta2;

/*Other variables*/
lba:if aa<0 then "(" else "";
rba:if aa<0 then ")" else "";
lbb:if bb/i<0 then "(" else "";
rbb:if bb/i<0 then ")" else "";
quad:if aa>0 then if bb/i>0 then 1 else 4 else if bb/i>0 then 2 else 3;
if quad = 1 then (lb:0,ub:%pi/2) elseif quad = 2 then (lb:%pi/2,ub:%pi) elseif quad=3 then (lb:-%pi,ub:%pi/2) else (lb:%pi/2,ub:0);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Complex number: \({@aa@}{@sigb@}{@bbi@}\)</p>
<p>\(r = {@ta1@}\)
<p>\(\theta = {@ta2@} \)
<p>\(q = {@ta3@} \)
<p>\(\delta = {@ta4@} \)]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>arg</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer is correct but not in the conventional argument range of \(-\pi &lt; \theta \leq \pi\).</p>
<p></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>quad</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-3-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1204635627</deployedseed>
    <deployedseed>2125435186</deployedseed>
    <deployedseed>1321869247</deployedseed>
    <deployedseed>1407393607</deployedseed>
    <deployedseed>921228700</deployedseed>
    <deployedseed>1322123396</deployedseed>
    <deployedseed>724125201</deployedseed>
    <deployedseed>1621348486</deployedseed>
    <deployedseed>1272414060</deployedseed>
    <deployedseed>1632824363</deployedseed>
    <deployedseed>1018467456</deployedseed>
    <deployedseed>1943242548</deployedseed>
    <deployedseed>1171990202</deployedseed>
    <deployedseed>590106565</deployedseed>
    <deployedseed>1609760285</deployedseed>
    <deployedseed>332679522</deployedseed>
    <deployedseed>2145582676</deployedseed>
    <deployedseed>1393163970</deployedseed>
    <deployedseed>486349708</deployedseed>
    <deployedseed>999717798</deployedseed>
    <deployedseed>2146910560</deployedseed>
    <deployedseed>1629321053</deployedseed>
    <deployedseed>240961756</deployedseed>
    <deployedseed>72821631</deployedseed>
    <deployedseed>207671856</deployedseed>
    <deployedseed>1843375177</deployedseed>
    <deployedseed>527801812</deployedseed>
    <deployedseed>2013803900</deployedseed>
    <deployedseed>1774315859</deployedseed>
    <deployedseed>387576095</deployedseed>
    <deployedseed>1618115883</deployedseed>
    <deployedseed>72409879</deployedseed>
    <deployedseed>741672489</deployedseed>
    <deployedseed>1322410832</deployedseed>
    <deployedseed>724975140</deployedseed>
    <deployedseed>2105934495</deployedseed>
    <deployedseed>1788013521</deployedseed>
    <deployedseed>863421824</deployedseed>
    <deployedseed>133806463</deployedseed>
    <deployedseed>430928451</deployedseed>
    <deployedseed>71970050</deployedseed>
    <deployedseed>566408057</deployedseed>
    <deployedseed>143503053</deployedseed>
    <deployedseed>1742024786</deployedseed>
    <deployedseed>666962865</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-3-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
