<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Composition of Functions</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Given the following two functions,</p>
<p style="margin-left: 30px">\(f(x)={@fx@}\)
<br>\(g(x)={@gx@}\)</p>
<p>Determine the composition \(f(g(x))\).</p>

<p style="margin-left: 30px">\( f(g(x)) = \) [[input:ans1]] [[validation:ans1]]</p>
<p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>To find \(f(g(x))\) we need to apply the function \(f(x)\) to the function \(g(x)\). We compute the following,</p>
<p style="margin-left: 30px">\(\begin{align*}
f(g(x)) &= f({@gx@}) \\
&= {@aa*(x^2+cc)+bb@} \\
&= {@ta1@}
\end{align*}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>aa:rand(4)+2;
pm1:rand_with_prohib(-1,1,[0]);
bb:pm1*(rand(8)+1);
fx:aa*x+bb;

pm2:rand_with_prohib(-1,1,[0]);
cc:pm2*(rand(4)+1);
gx:x^2+cc;

ta1:expand(ev(fx,x=gx));
incorr_comp_order:expand(ev(gx,x=fx));</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\(f(x) = {@fx@}\)
<br>\(g(x) = {@gx@}\)
<br>\(f(g(x)) = {@ta1@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>incorr_comp_order</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have computed the function \(g(f(x))\).
<br>Take care with the order in which you apply the functions in the composition.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>3556340</deployedseed>
    <deployedseed>1815329883</deployedseed>
    <deployedseed>1136353422</deployedseed>
    <deployedseed>2084863717</deployedseed>
    <deployedseed>1853557674</deployedseed>
    <deployedseed>1847878784</deployedseed>
    <deployedseed>1877701699</deployedseed>
    <deployedseed>1035811904</deployedseed>
    <deployedseed>1547827146</deployedseed>
    <deployedseed>211571835</deployedseed>
    <deployedseed>2133775280</deployedseed>
    <deployedseed>197046438</deployedseed>
    <deployedseed>110269217</deployedseed>
    <deployedseed>1086159124</deployedseed>
    <deployedseed>1746070762</deployedseed>
    <deployedseed>248218954</deployedseed>
    <deployedseed>1031729113</deployedseed>
    <deployedseed>372238492</deployedseed>
    <deployedseed>2107960605</deployedseed>
    <deployedseed>373057455</deployedseed>
    <deployedseed>626991301</deployedseed>
    <deployedseed>1790487852</deployedseed>
    <deployedseed>1072898085</deployedseed>
    <deployedseed>551916043</deployedseed>
    <deployedseed>912776710</deployedseed>
    <deployedseed>1427356498</deployedseed>
    <deployedseed>1861290257</deployedseed>
    <deployedseed>1658631256</deployedseed>
    <deployedseed>122076496</deployedseed>
    <deployedseed>1244292697</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>incorr_comp_order</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
