<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Domain and Range of a Function</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find the domain and range of \(f(x)={@fx@}\).</p><p></p>
<p>Give your answer as interval of the real numbers.</p>
<p><em>Note: To enter intervals of you real line, you should write the relevant text in the table below.</em></p>

<table width="250" cellspacing="1" cellpadding="0" border="1" style="margin-left: 30px">
  <tr>
    <th style="text-align: center;">Intervals</th>
    <th style="text-align: center;">Input Syntax</th>
  </tr>
  <tr>
    <td style="text-align: center;">\( [a,b] \)</td>
    <td style="text-align: center;"><code>cc(a,b)</code></td>
  </tr>
  <tr>
    <td style="text-align: center;">\( (a,b) \)</td>
    <td style="text-align: center;"><code>oo(a,b)</code></td>
  </tr>
  <tr>
    <td style="text-align: center;">\( [a,b) \)</td>
    <td style="text-align: center;"><code>co(a,b)</code></td>
  </tr>
  <tr>
    <td style="text-align: center;">\( (a,b] \)</td>
    <td style="text-align: center;"><code>oc(a,b)</code></td>
  </tr>
 <tr>
    <td style="text-align: center;">\( (a, \infty) \)</td>
    <td style="text-align: center;"><code>oo(a,inf)</code></td>
  </tr>
 <tr>
    <td style="text-align: center;">\( (-\infty, b] \)</td>
    <td style="text-align: center;"><code>oc(-inf,b)</code></td>
  </tr>
 <tr>
    <td style="text-align: center;">\( (-\infty, \infty) \)</td>
    <td style="text-align: center;"><code>oo(-inf,inf)</code></td>
  </tr>
</table>
</p>

<br>

<p>(a) \(Dom(f)=\) [[input:ans1]] [[validation:ans1]]</p>
<p>[[feedback:prt1]]</p>
<p>(b) \(Range(f) =\) [[input:ans2]] [[validation:ans2]]</p>
<p>[[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>(a) To determine the domain of the function, we need to identify any restrictions on the values of \(x\) that would make the expression undefined.</p>
<p>Looking at the function expression \(f(x)={@fx@}\), we observe that there are no explicit limitations or restrictions on the variable \(x\). </p>
<p>Therefore, the domain of \(f(x)\) is all real numbers, which can be denoted as,</p>
<p style="margin-left: 30px">\( Dom(f) = \mathbb{R} = {@ta1@}\)</p>

<br>

<p>(b) Next, we determine the range of the function. 
<br>The range represents the set of all possible values that \(f(x)\) can attain as \(x\) varies.</p>

<p>Since we recognize that \(f(x)\) is a quadratic function, the range will be given by a <b>half-line</b>, \( [a, \infty) \) or \( (-\infty, a] \)
<br>Where the maximum or minimum value of \(f(x)\) is determined by the vertex of the quadratic function.</p>

<br>

<p>To determine the vertex of the quadratic function, we can "complete the square".
<br><em>Note: It is also possible to use a calculus method here to determine the turning point of the function, which is the same as the vertex.</em>

<p hidden>{@simp:false;@}</p>
<p style="margin-left: 30px">\(\begin{align*}
f(x) &= {@fx@} \\
&= {@line1@} \\
&= {@line2@}
\end{align*}\)</p>
<p hidden>{@simp:true;@}</p>

<br>

<p>Therefore, \(f(x)\) has a <b>minimum</b> value of {@yv@} since \(({@x - xv@})^2 \geq 0\)
<br>Hence, the vertex of \(f(x)\) occurs at \( ({@xv@}, {@yv@}) \).</p>

<p>Then, considering the shape of a quadratic function, the range of the function \(f(x)\) can be expressed as:</p>
<p style="margin-left: 30px">\(Range(f)= \{ y \in \mathbb{R}: y \geq {@yv@} \} ={@ta2@}\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>aa:rand(8)+1;
pm1:rand_with_prohib(-1,1,[0]);
bb:rand(8)+1;
pm2:rand_with_prohib(-1,1,[0]);
cc:rand(8)+1;
fx:expand((x-pm1*aa)*(x-pm2*bb)+cc);
xv:(pm1*aa+pm2*bb)/2;
yv:ev(fx,x=xv);
ta1:oo(-inf,inf);
ta2:co(yv,inf);

/*worked solution: unevaluated lines*/
simp:false;

line1:(x - ev((pm1*aa + pm2*bb)/2,simp))^2 - ev((pm1*aa + pm2*bb)^2/4,simp) + ev(pm1*pm2*aa*bb + cc,simp);

line2:(x - ev((pm1*aa + pm2*bb)/2,simp))^2 + ev(-(pm1*aa + pm2*bb)^2/4 + pm1*pm2*aa*bb + cc,simp);

simp:true;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\(f(x)={@fx@}\)</p>
<p>\( Dom(f) = {@ta1@}\) </p>
<p> \( Range(f)={@ta2@} \)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The domian of \(f(x)\) is the set of value of \(x\) for which we can evaluate \(f(x)\).
<br>In this case, there are no restrictions on the values for which \(f(x)\) can be evaluated and therefore,</p>
<p style="margin-left: 30px">\( Range(f) = \mathbb{R} = {@ta1@}\)</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>\(f(x)\) is a quadratic function and so has the shape of a parabola.</p>
<p>The minimum value of the function can be determined by finding the turning point or "vertex" of the parabola, either by completing the square, differentiating or by using a known formula for the vertex.</p>
<p>The range of the function is the set of potential \(y\)-coordinates of points on this graph. Therefore, once we have found the minimum value, \(a\) we can see that the range of \(f\) is given by \( [a, \infty) \).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1136150450</deployedseed>
    <deployedseed>1584124657</deployedseed>
    <deployedseed>1400740742</deployedseed>
    <deployedseed>1114688405</deployedseed>
    <deployedseed>1982282526</deployedseed>
    <deployedseed>1786519702</deployedseed>
    <deployedseed>1922180854</deployedseed>
    <deployedseed>1637284869</deployedseed>
    <deployedseed>1974802759</deployedseed>
    <deployedseed>1188800621</deployedseed>
    <deployedseed>722834619</deployedseed>
    <deployedseed>215196199</deployedseed>
    <deployedseed>2087044790</deployedseed>
    <deployedseed>472742202</deployedseed>
    <deployedseed>1411478884</deployedseed>
    <deployedseed>633816097</deployedseed>
    <deployedseed>444527968</deployedseed>
    <deployedseed>133838073</deployedseed>
    <deployedseed>1840184117</deployedseed>
    <deployedseed>757647400</deployedseed>
    <deployedseed>232098154</deployedseed>
    <deployedseed>701527995</deployedseed>
    <deployedseed>1569957539</deployedseed>
    <deployedseed>120683661</deployedseed>
    <deployedseed>1656918592</deployedseed>
    <deployedseed>911891248</deployedseed>
    <deployedseed>1725578808</deployedseed>
    <deployedseed>971794450</deployedseed>
    <deployedseed>1373943444</deployedseed>
    <deployedseed>2061078281</deployedseed>
    <deployedseed>1255935882</deployedseed>
    <deployedseed>1076180716</deployedseed>
    <deployedseed>1022669278</deployedseed>
    <deployedseed>1628618234</deployedseed>
    <deployedseed>1544800902</deployedseed>
    <deployedseed>943304391</deployedseed>
    <deployedseed>989110098</deployedseed>
    <deployedseed>137729688</deployedseed>
    <deployedseed>368150977</deployedseed>
    <deployedseed>290445326</deployedseed>
    <deployedseed>189086155</deployedseed>
    <deployedseed>190957148</deployedseed>
    <deployedseed>611109721</deployedseed>
    <deployedseed>742175475</deployedseed>
    <deployedseed>1231230868</deployedseed>
    <deployedseed>1674321676</deployedseed>
    <deployedseed>1260845548</deployedseed>
    <deployedseed>1288713620</deployedseed>
    <deployedseed>1211532054</deployedseed>
    <deployedseed>1190547497</deployedseed>
    <deployedseed>1848642253</deployedseed>
    <deployedseed>44331371</deployedseed>
    <deployedseed>591094082</deployedseed>
    <deployedseed>1669725951</deployedseed>
    <deployedseed>1780213809</deployedseed>
    <deployedseed>643818103</deployedseed>
    <deployedseed>584887336</deployedseed>
    <deployedseed>1747720466</deployedseed>
    <deployedseed>1531498524</deployedseed>
    <deployedseed>1759369883</deployedseed>
    <deployedseed>1989410958</deployedseed>
    <deployedseed>414564901</deployedseed>
    <deployedseed>1922503464</deployedseed>
    <deployedseed>1595706073</deployedseed>
    <deployedseed>1848317743</deployedseed>
    <deployedseed>644058483</deployedseed>
    <deployedseed>2023628073</deployedseed>
    <deployedseed>1931794977</deployedseed>
    <deployedseed>2020511882</deployedseed>
    <deployedseed>348125725</deployedseed>
    <deployedseed>1711866146</deployedseed>
    <deployedseed>1258651312</deployedseed>
    <deployedseed>1827154629</deployedseed>
    <deployedseed>898640747</deployedseed>
    <deployedseed>268415919</deployedseed>
    <deployedseed>2140649325</deployedseed>
    <deployedseed>870098884</deployedseed>
    <deployedseed>2054254935</deployedseed>
    <deployedseed>596657044</deployedseed>
    <deployedseed>1783382239</deployedseed>
    <deployedseed>1431857188</deployedseed>
    <deployedseed>191814343</deployedseed>
    <deployedseed>1807704925</deployedseed>
    <deployedseed>402205428</deployedseed>
    <deployedseed>1517307371</deployedseed>
    <deployedseed>1590776380</deployedseed>
    <deployedseed>1567952137</deployedseed>
    <deployedseed>2004716055</deployedseed>
    <deployedseed>2067620038</deployedseed>
    <deployedseed>1479303845</deployedseed>
    <deployedseed>470294420</deployedseed>
    <deployedseed>210113450</deployedseed>
    <deployedseed>1233701316</deployedseed>
    <deployedseed>741378413</deployedseed>
    <deployedseed>1741991803</deployedseed>
    <deployedseed>556095073</deployedseed>
    <deployedseed>568452425</deployedseed>
    <deployedseed>1939347256</deployedseed>
    <deployedseed>1662323243</deployedseed>
    <deployedseed>211978295</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
