<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Determining Real and Imaginary Parts of a Complex Function</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p></p><p></p><p>Find the real part \(u(x,y)\) and the imaginary part \(v(x,y)\) of the complex function \(f(z) = {@ f_z @}\).</p><p>\(u(x,y)\) = [[input:ans1]] [[validation:ans1]]&nbsp;</p><p>\(v(x,y)\) = [[input:ans2]] [[validation:ans2]]&nbsp;</p>
<p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked solution</h3>
<p>Find the real part \(u(x,y)\) and the imaginary part \(v(x,y)\) of the complex function \(f(z) =  \displaystyle {@ f_z @}\)</p>
<p>We have</p>
<p> \(\begin{align}
f(z) & = \displaystyle {@f_z@} \\
& = \displaystyle {@subst(x+y*%i,z,f_z)@} \quad (\text{by substituting } z=x+iy) \\
& = \displaystyle {@f_xy@}  \\
& = \displaystyle  {@ta1@}+ \left({@ta2@} \right){@%i@} \\
& = \displaystyle  {@ratsimp(ta1)@}+ \left({@ratsimp(ta2)@} \right){@%i@} \\
\end{align}\). </p>
<p> Since \(f(z)= u(x,y)+{@%i@} v(x,y) \), by identification, we have</p>
<p>\(u(x,y) =\displaystyle {@ratsimp(ta1)@}\) and \(v(x,y)=\displaystyle {@ratsimp(ta2)@}\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/* creating a complex function*/
aa: rand_with_prohib(-5,5,[0]);
root1: rand(7)-3; /* Random number between -3 and 3 */
root2: rand_with_prohib(-3, 3, [root1]);
root3: rand_with_prohib(-3, 3, [root1, root2]);

f_z: expand((aa*(z-root1))/(z-root2)); /* Expand makes sure f_z is expanded when displayed to the students */
zz: x + y*%i;
f_xy: expand((aa*(zz-root1))/(zz-root2)); /*writing f_z with =x+iy*/
ta1: realpart(f_xy);
ta2: imagpart(f_xy);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Real and imaginary parts of a complex function {@f_z@}.</p>
<p>Real part: \(\displaystyle {@ta1@}\)</p>
<p>Imaginary part:  \(\displaystyle {@ta2@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ratsimp(ta1)</tans>
      <boxsize>25</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ratsimp(ta2)</tans>
      <boxsize>25</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The real part of your function is incorrect. Check your evaluation again.<br></p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The imaginary part of your function is incorrect. Check your evaluation.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1143616786</deployedseed>
    <deployedseed>846364854</deployedseed>
    <deployedseed>855911396</deployedseed>
    <deployedseed>1954935471</deployedseed>
    <deployedseed>1989769512</deployedseed>
    <deployedseed>1441566886</deployedseed>
    <deployedseed>1937911060</deployedseed>
    <deployedseed>620611356</deployedseed>
    <deployedseed>1489402124</deployedseed>
    <deployedseed>1377982301</deployedseed>
    <deployedseed>1770815142</deployedseed>
    <deployedseed>1094522803</deployedseed>
    <deployedseed>1017386838</deployedseed>
    <deployedseed>896797335</deployedseed>
    <deployedseed>1489988513</deployedseed>
    <deployedseed>1276083825</deployedseed>
    <deployedseed>33487258</deployedseed>
    <deployedseed>853062040</deployedseed>
    <deployedseed>454116712</deployedseed>
    <deployedseed>1166637712</deployedseed>
    <deployedseed>268677503</deployedseed>
    <deployedseed>611376655</deployedseed>
    <deployedseed>650139652</deployedseed>
    <deployedseed>406580603</deployedseed>
    <deployedseed>137702058</deployedseed>
    <deployedseed>1189773012</deployedseed>
    <deployedseed>764053409</deployedseed>
    <deployedseed>2098680950</deployedseed>
    <deployedseed>2087622508</deployedseed>
    <deployedseed>678867730</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
