<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text><![CDATA[Finding Real and Imaginary Parts of a Complex Function"]]></text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p></p><p></p><p>Find the real part \(u(x,y)\) and the imaginary part \(v(x,y)\) of the complex function \(f(z) = {@ f_z @}\).</p><p>\(u(x,y)\) = [[input:ans1]] [[validation:ans1]]&nbsp;</p><p>\(v(x,y)\) = [[input:ans2]] [[validation:ans2]]&nbsp;</p><p></p>
<p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked solution</h3>
<br> To decompose a complex number into its real and imaginary parts, we need to  replace \(z\) with \( x + iy \)
<p></p>
<br> Therefore, \( f(z)\) will be;
<br> {@f_z@} =  {@f_xy@}
<br> \(f(z) = {@ta1@}+i({@ta2@})\).</p>
<p></p>
<br>From this expression, we can identify the real part \( u(x, y)\) and the imaginary part \(v(x, y)\) as follows:
<br> Real part: \(u(x,y)= {@ta1@}\) 
<br> Imaginary part: \(v(x,y)={@ta2@}\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/* creating a complex function*/
aa: rand_with_prohib(-5,5,[0]);
root1: rand(7)-3; /* Random number between -3 and 3 */
root2: rand_with_prohib(-3, 3, [root1]);
root3: rand_with_prohib(-3, 3, [root1, root2]);

f_z: expand(aa*(z-root1)*(z-root2)); /* Expand makes sure f_z is expanded when displayed to the students */
zz: x + y*%i;
f_xy: expand(aa*(zz-root1)*(zz-root2)); /*writing f_z with =x+iy*/
ta1: realpart(f_xy);
ta2: imagpart(f_xy);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Complex function: \({@ f_z @}\)</p>
<p>\(u(x,y) ={@ta1@}\) </p>
<p>\(v(x,y) ={@ta2@}\) </p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The real part of your function is incorrect. Check your evaluation again.<br></p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The imaginary part of your function is incorrect. Check your evaluation.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1101165113</deployedseed>
    <deployedseed>903583319</deployedseed>
    <deployedseed>316807687</deployedseed>
    <deployedseed>456580255</deployedseed>
    <deployedseed>1508240646</deployedseed>
    <deployedseed>576092040</deployedseed>
    <deployedseed>2118379942</deployedseed>
    <deployedseed>2051518779</deployedseed>
    <deployedseed>1903637478</deployedseed>
    <deployedseed>885949728</deployedseed>
    <deployedseed>1424978958</deployedseed>
    <deployedseed>1884265733</deployedseed>
    <deployedseed>957060326</deployedseed>
    <deployedseed>946302855</deployedseed>
    <deployedseed>1103573052</deployedseed>
    <deployedseed>1508016728</deployedseed>
    <deployedseed>210379783</deployedseed>
    <deployedseed>18409339</deployedseed>
    <deployedseed>451734564</deployedseed>
    <deployedseed>793431330</deployedseed>
    <deployedseed>1310524847</deployedseed>
    <deployedseed>654045431</deployedseed>
    <deployedseed>1391637005</deployedseed>
    <deployedseed>1638892524</deployedseed>
    <deployedseed>422855800</deployedseed>
    <deployedseed>1486017970</deployedseed>
    <deployedseed>172713713</deployedseed>
    <deployedseed>1370947871</deployedseed>
    <deployedseed>1154673262</deployedseed>
    <deployedseed>147461199</deployedseed>
    <deployedseed>1053522611</deployedseed>
    <deployedseed>205510888</deployedseed>
    <deployedseed>446249449</deployedseed>
    <deployedseed>1147618571</deployedseed>
    <deployedseed>765307366</deployedseed>
    <deployedseed>1604031112</deployedseed>
    <deployedseed>506551311</deployedseed>
    <deployedseed>28929925</deployedseed>
    <deployedseed>1211471379</deployedseed>
    <deployedseed>359463667</deployedseed>
    <deployedseed>386058440</deployedseed>
    <deployedseed>1097637677</deployedseed>
    <deployedseed>900067638</deployedseed>
    <deployedseed>353904666</deployedseed>
    <deployedseed>2102324039</deployedseed>
    <deployedseed>740682147</deployedseed>
    <deployedseed>1869687934</deployedseed>
    <deployedseed>824361549</deployedseed>
    <deployedseed>1094890717</deployedseed>
    <deployedseed>215357612</deployedseed>
    <deployedseed>1404425486</deployedseed>
    <deployedseed>641136978</deployedseed>
    <deployedseed>1783621212</deployedseed>
    <deployedseed>1862050280</deployedseed>
    <deployedseed>347147202</deployedseed>
    <deployedseed>1895205149</deployedseed>
    <deployedseed>2092129468</deployedseed>
    <deployedseed>1566852107</deployedseed>
    <deployedseed>451974908</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
