<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Set of Real numbers</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Simplify as much as possible the set of real numbers \(\displaystyle \frac{{@aa@}+{@bb@}-{@cc@}}{{@dd@}+{@ee@}-{@bb@}}\)</p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>To solve the set of real number  \( \displaystyle \frac{{@aa@}+{@bb@}-{@cc@}}{{@dd@}+{@ee@}-{@bb@}}\)</p>
<p>We solve the numerator first by operating on \({@aa@}+{@bb@}-{@cc@}={@sol_1@}\)</p>
<p>Solving the denominator \({@dd@}+{@ee@}-{@bb@}={@sol_2@}\)</p>
<p>Diving the result of the numerator and denominator, we have</p>
<p>\( \begin{align}
 \displaystyle \frac{{@aa@}+{@bb@}-{@cc@}}{{@dd@}+{@ee@}-{@bb@}} & = \displaystyle \frac{{@sol_1@}}{{@sol_2@}}\\
&= \displaystyle {@ta@}
\end{align} \)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>aa:rand(4)+2;
bb:rand(10)+6/8;
cc:rand(4)+aa;
dd:rand(30)+10/12;
ee:rand(14)+8/6;
/*solution*/
sol_1:aa+bb-cc;
sol_2:dd+ee-bb;
ta:sol_1/sol_2;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Evaluate \(\frac{{@aa@}+{@bb@}-{@cc@}}{{@dd@}+{@ee@}-{@bb@}}\)</p>
<p>{@ta@}</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>224814912</deployedseed>
    <deployedseed>1622231901</deployedseed>
    <deployedseed>48945777</deployedseed>
    <deployedseed>1808930778</deployedseed>
    <deployedseed>692898261</deployedseed>
    <deployedseed>1320030898</deployedseed>
    <deployedseed>1851781083</deployedseed>
    <deployedseed>1697685622</deployedseed>
    <deployedseed>286434030</deployedseed>
    <deployedseed>1269108785</deployedseed>
    <deployedseed>300276303</deployedseed>
    <deployedseed>1843893148</deployedseed>
    <deployedseed>1352592930</deployedseed>
    <deployedseed>271202070</deployedseed>
    <deployedseed>217375176</deployedseed>
    <deployedseed>1773180785</deployedseed>
    <deployedseed>1012038762</deployedseed>
    <deployedseed>1856813924</deployedseed>
    <deployedseed>1114175044</deployedseed>
    <deployedseed>1431004045</deployedseed>
    <deployedseed>1534002884</deployedseed>
    <deployedseed>837694040</deployedseed>
    <deployedseed>637350561</deployedseed>
    <deployedseed>22541145</deployedseed>
    <deployedseed>1026936991</deployedseed>
    <deployedseed>73730026</deployedseed>
    <deployedseed>1753121790</deployedseed>
    <deployedseed>809604017</deployedseed>
    <deployedseed>276805627</deployedseed>
    <deployedseed>913799639</deployedseed>
    <deployedseed>1065770688</deployedseed>
    <deployedseed>346789844</deployedseed>
    <deployedseed>1578501807</deployedseed>
    <deployedseed>1454829202</deployedseed>
    <deployedseed>699449602</deployedseed>
    <deployedseed>434431958</deployedseed>
    <deployedseed>60182992</deployedseed>
    <deployedseed>1182737983</deployedseed>
    <deployedseed>1245800101</deployedseed>
    <deployedseed>1163565233</deployedseed>
    <deployedseed>4157665</deployedseed>
    <deployedseed>1346147441</deployedseed>
    <deployedseed>2120469473</deployedseed>
    <deployedseed>147176740</deployedseed>
    <deployedseed>1476664313</deployedseed>
    <deployedseed>1079416100</deployedseed>
    <deployedseed>467595893</deployedseed>
    <deployedseed>246479151</deployedseed>
    <deployedseed>1769559750</deployedseed>
    <deployedseed>77002980</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
