<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the modulus of a complex number</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Find the modulus of the complex number \(\displaystyle z = \frac{{@num@}}{{@den@}}\)</p>
<p><em>Note: Give your answers in exact form i.e using fractions or writing in square root form</em></p>
<p>\(\displaystyle   |z| = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>To find the modulus of a complex number, you can use the formula:</p>
 
<p>\(|z| = \sqrt{{\text{Re}(z)}^2 + {\text{Im}(z)}^2}\),</p>

<p>where \(\text{Re}(z)\) is the real part of the complex number and \(\text{Im}(z)\) is the imaginary part.</p>

<p>Given the complex number \(\displaystyle z = \frac{{@num@}}{{@den@}}\), let's simplify it first:</p>

<p>\(\displaystyle z = \frac{{@num@}}{{@den@}}\)</p>

<p>To get rid of the denominator, you can multiply both the numerator and the denominator by \({@den_conj@}\):</p>

<p>\(\displaystyle z = \left( \frac{{@num@}}{{@den@}} \right) \left( \frac{{@den_conj@}}{{@den_conj@}} \right) \)</p>
<p>\(\displaystyle z = \frac{{@expand(num1)@}}{{@den1@}}\)</p>
<p>\(\displaystyle z = {@expand(z_conjugate)@}\)</p>

<p>Now we have the real and imaginary parts of the complex number:</p>

<p>\( \displaystyle \text{Re}(z) = {@re_z@}\)</p>
<p>\(\displaystyle \text{Im}(z) = {@im_z@}\)</p>

<p>Now we can plug these values into the modulus formula:</p>

<p>\(\displaystyle \left | z \right |= \sqrt{( {@re_z@})^2 + \left({@im_z@}\right)^2}\)</p>
<p>\(\displaystyle \left | z \right | = \sqrt{ {@re_z^2@} + {@im_z^2@}}\)</p>
<p>\(\displaystyle |z| = \sqrt{ {@re_z^2 + im_z^2@}}\)</p>
<p>\(\displaystyle |z| = {@mod@}\)</p>

<p>So, the modulus of the complex number \(\displaystyle z = \ \frac{{@num@}}{{@den@}}\) is \( \displaystyle {@mod@}\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/* Define the complex number */
a:rand(50)+1;
b:rand_with_prohib(-50,50,[0,a]);
num:(a + b*%i);
den:(-b*%i);
den_conj:conjugate(den);
z: num/den;


/* Simplify the complex number */
z_simplified: ratsimp(z);

/* Multiply by conjugate to get rid of the denominator */
num1:num*den_conj;
den1:den*den_conj;
z_conjugate:num1/den1;

/* Simplify the conjugate expression */
z_conjugate_simplified: (num1/den1);

/* Separate real and imaginary parts */
re_z: realpart(z_conjugate_simplified);
im_z: imagpart(z_conjugate_simplified);

/* Calculate the modulus */
mod: sqrt(re_z^2 + im_z^2);

/* Print the result */
ta:mod;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Find the modulus of the complex number \(z ={@z@}\).</p>
<p>\(\displaystyle  |z| = {@ta@}\) </p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
