<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Modulus of complex numbers</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p> Given the complex numbers \(z={@z@}\) and \(w={@w@}\)</p> 
<br> Evaluate:
<br>\(|z|=\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
<br>
<br>\(|z+w|=\) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked solution.</h3>
<p>\(\begin{align}
\displaystyle |z| &= \sqrt{z \overline{z}}\\
& = \sqrt{({@z@})(\overline{@z@})}\\
& =  \sqrt{({@z@}) ({@conjugate(z)@})}\\
& ={@ta4@}
\end{align}\)</p>
<p>We have \(z+w=({@z@})+({@w@})={@z+w@}\). So</p>
<p>\(\begin{align}
\displaystyle |z+w|& = \sqrt{({@z+w@})(\overline{@z+w@})}\\
& =  \sqrt{({@z+w@}) ({@conjugate(z+w)@})}\\
& ={@ta5@}
\end{align}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>a: rand(20); 
b: 5-rand(20); 
c: 1+rand(15); 
d: -(1+rand(10));
z: a+b*%i; 
w: c+d*%i;
addzw: z+w;
ta1: addzw;

magz: abs(z);
ta4: magz;
magzw: abs(z+w);
ta5: magzw;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\(|z|={@ta4@}\)</p>
<p>\(|z+w|={@ta5@}\) </p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<p>Correct answer, well done.</p>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<p>Your answer is partially correct.</p>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<p>Incorrect answer.</p>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta5</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>3</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta5</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1961539180</deployedseed>
    <deployedseed>728618471</deployedseed>
    <deployedseed>2020413942</deployedseed>
    <deployedseed>515560716</deployedseed>
    <deployedseed>736520072</deployedseed>
    <deployedseed>2136414882</deployedseed>
    <deployedseed>444408289</deployedseed>
    <deployedseed>584325453</deployedseed>
    <deployedseed>211335723</deployedseed>
    <deployedseed>99514093</deployedseed>
    <deployedseed>1507489169</deployedseed>
    <deployedseed>1825362433</deployedseed>
    <deployedseed>2013966556</deployedseed>
    <deployedseed>718661810</deployedseed>
    <deployedseed>1754785940</deployedseed>
    <deployedseed>264308950</deployedseed>
    <deployedseed>1015848009</deployedseed>
    <deployedseed>101919395</deployedseed>
    <deployedseed>1321726377</deployedseed>
    <deployedseed>1703303772</deployedseed>
    <deployedseed>954512539</deployedseed>
    <deployedseed>651464439</deployedseed>
    <deployedseed>1676649496</deployedseed>
    <deployedseed>1727232132</deployedseed>
    <deployedseed>663934491</deployedseed>
    <deployedseed>1433597478</deployedseed>
    <deployedseed>1941315637</deployedseed>
    <deployedseed>457289249</deployedseed>
    <deployedseed>965365397</deployedseed>
    <deployedseed>1838579638</deployedseed>
    <deployedseed>1518109527</deployedseed>
    <deployedseed>237105265</deployedseed>
    <deployedseed>1526632193</deployedseed>
    <deployedseed>1537049646</deployedseed>
    <deployedseed>514591161</deployedseed>
    <deployedseed>863929482</deployedseed>
    <deployedseed>1873909087</deployedseed>
    <deployedseed>424083330</deployedseed>
    <deployedseed>1665961615</deployedseed>
    <deployedseed>550802966</deployedseed>
    <deployedseed>1870935931</deployedseed>
    <deployedseed>1030508546</deployedseed>
    <deployedseed>931733007</deployedseed>
    <deployedseed>156624563</deployedseed>
    <deployedseed>2145203030</deployedseed>
    <deployedseed>1220259633</deployedseed>
    <deployedseed>1441022166</deployedseed>
    <deployedseed>1149317329</deployedseed>
    <deployedseed>1024131798</deployedseed>
    <deployedseed>1138339755</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta4</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta5</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-4-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
