<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the determinants of two matrices with no relationship</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Suppose that \({@mata@} = {@aa@}\). Find the determinant of the matrix below.</p>

<br>

<p style="margin-left: 30px">\({@matb@} = \) [[input:ans1]] [[validation:ans1]]</p>
<p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution:</h3>

<p>To get the determinant of the given matrix, we will perform row operations on {@mata@} to give the matrix \({@matb@}\)</p>
<p>Recalling the known ways in which the row operations change the determinant of a matrix, this allows us to calculate the unknown determinant in terms of the given one.</p>

<br>

<p><b>Step 1:</b> Interchange row 2 and row 3</p>
<p style="margin-left:30px">\( 
\begin{vmatrix} a & b & c \\ d & e & f \\ g & h & i \end{vmatrix} 
\xrightarrow{R2 \leftrightarrow R3} 
\begin{vmatrix} a & b & c \\ g & h & i \\ d & e & f \end{vmatrix} 
\)</p>
<p>Since the action of interchanging two rows changes the determinant by a factor of \(-1\), the determinant after step 1 is \({@-aa@}\).</p>

<p><b>Step 2:</b> Scale the second row by {@bb@}</p>
<p style="margin-left:30px">\( 
\begin{vmatrix} a & b & c \\ g &  h & i \\ d & e & f \end{vmatrix} 
\xrightarrow{R2 \rightarrow {@bb@}R2} 
\begin{vmatrix} a & b & c \\  {@bb*g@} & {@bb*h@} &  {@bb*i@}  \\ d & e & f \end{vmatrix} 
\)</p>
<p>Since the action of scaling a row changes the determinant by the same factor as the scaling, the determinant after step 2 is \({@bb@} \times {@-aa@} = {@bb*-aa@}\)</p>

<p><b>Step 3:</b> Scale the second row by {@cc@}</p>
<p style="margin-left:30px">\( 
\begin{vmatrix} a & b & c \\ {@bb*g@} & {@bb*h@} &  {@bb*i@}  \\ d & e & f \end{vmatrix} 
\xrightarrow{R3 \rightarrow {@cc@} R3} 
\begin{vmatrix} a & b & c \\ {@bb*g@} & {@bb*h@} &  {@bb*i@}  \\ {@cc *d@} & {@cc*e@} & {@cc*f@} \end{vmatrix} 
\)</p>
<p>Since the action of scaling a row changes the determinant by the same factor as the scaling, the determinant after step 3 is \( {@cc@} \times {@bb*-aa@} = {@cc*bb*(-aa)@} \)
</p>

<p><b>Step 4:</b> Subtract the first row from the third row</p>
<p style="margin-left:30px">\( 
\begin{vmatrix} a & b & c \\ {@bb*g@} & {@bb*h@} &  {@bb*i@}  \\ {@cc *d@} & {@cc*e@} & {@cc*f@} \end{vmatrix} 
\xrightarrow{R3 \rightarrow R3 - R1} 
\begin{vmatrix} a & b & c \\ {@bb*g@} & {@bb*h@} &  {@bb*i@}  \\ {@cc *d@}-a & {@cc*e@}-b & {@cc*f@}-c  \end{vmatrix} \)</p>
<p>Since the action of replacing a row with a sum or difference of itself and another row does not change the determinant, the determinant after step 4 is still \( {@cc*bb*(-aa)@} \)
</p>

<br>

<p>So, the determinant of the final matrix is \( {@cc*(bb*-aa)@} \)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>ordergreat(i,h,g,f,e,d,c,b,a);
aa:rand_with_prohib(-20,20,[0]);
bb:rand(3)+2;
cc:rand(3)+2;
mata:matrix([a,b,c],[d,e,f],[g,h,i]);
matb:matrix([a,b,c],[bb*g,bb*h,bb*i],[cc*d-a,cc*e-b,cc*f-c]);
ta:-aa*bb*cc;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>If \({@mata@} = {@aa@}\), then \( {@matb@}={@ta@} \).</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>|</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[lbr1:if aa<0 then "(" else "";
rbr1:if aa<0 then ")" else "";
/* Node 1:Checks for the correct answer */]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2027316945</deployedseed>
    <deployedseed>148410126</deployedseed>
    <deployedseed>1017308477</deployedseed>
    <deployedseed>1328985140</deployedseed>
    <deployedseed>1428398297</deployedseed>
    <deployedseed>1744931421</deployedseed>
    <deployedseed>1325373505</deployedseed>
    <deployedseed>1354537035</deployedseed>
    <deployedseed>349087377</deployedseed>
    <deployedseed>2036964546</deployedseed>
    <deployedseed>1317057670</deployedseed>
    <deployedseed>39879499</deployedseed>
    <deployedseed>406775023</deployedseed>
    <deployedseed>1612120082</deployedseed>
    <deployedseed>82385871</deployedseed>
    <deployedseed>849192105</deployedseed>
    <deployedseed>1456964259</deployedseed>
    <deployedseed>800785771</deployedseed>
    <deployedseed>1396668593</deployedseed>
    <deployedseed>856827857</deployedseed>
    <deployedseed>1145250140</deployedseed>
    <deployedseed>1181047018</deployedseed>
    <deployedseed>352579398</deployedseed>
    <deployedseed>1461195401</deployedseed>
    <deployedseed>24431540</deployedseed>
    <deployedseed>108486645</deployedseed>
    <deployedseed>1731923812</deployedseed>
    <deployedseed>997644557</deployedseed>
    <deployedseed>892900217</deployedseed>
    <deployedseed>1110623562</deployedseed>
    <deployedseed>639035346</deployedseed>
    <deployedseed>1488235512</deployedseed>
    <deployedseed>882854171</deployedseed>
    <deployedseed>1782153968</deployedseed>
    <deployedseed>1207134657</deployedseed>
    <deployedseed>1390471</deployedseed>
    <deployedseed>216522144</deployedseed>
    <deployedseed>784143753</deployedseed>
    <deployedseed>297218961</deployedseed>
    <deployedseed>1421430599</deployedseed>
    <deployedseed>1017797283</deployedseed>
    <deployedseed>1789064679</deployedseed>
    <deployedseed>1003583432</deployedseed>
    <deployedseed>1704882353</deployedseed>
    <deployedseed>782363421</deployedseed>
    <deployedseed>1110930712</deployedseed>
    <deployedseed>1156410642</deployedseed>
    <deployedseed>2123276207</deployedseed>
    <deployedseed>1303749143</deployedseed>
    <deployedseed>260259903</deployedseed>
    <deployedseed>1593200975</deployedseed>
    <deployedseed>1929370198</deployedseed>
    <deployedseed>18538495</deployedseed>
    <deployedseed>989458123</deployedseed>
    <deployedseed>1693159360</deployedseed>
    <deployedseed>10399670</deployedseed>
    <deployedseed>485591125</deployedseed>
    <deployedseed>60298661</deployedseed>
    <deployedseed>78559356</deployedseed>
    <deployedseed>1948661959</deployedseed>
    <deployedseed>1725180513</deployedseed>
    <deployedseed>1709320666</deployedseed>
    <deployedseed>2082693986</deployedseed>
    <deployedseed>238221843</deployedseed>
    <deployedseed>1071954148</deployedseed>
    <deployedseed>1763391242</deployedseed>
    <deployedseed>662211194</deployedseed>
    <deployedseed>614969333</deployedseed>
    <deployedseed>1781021683</deployedseed>
    <deployedseed>1538483160</deployedseed>
    <deployedseed>1860622687</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
