<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the Distance Between Two Points in R3</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p> Let  \(\mathbf{u} = {@u@}\) and \(\mathbf{v} = {@v@}\) in points defined in \(\mathbb{R}^3\). </p>

<p>Find the distance between \(\mathbf{u}\) and \(\mathbf{v}\).</p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>
<p>(Note: Give your answer in exact form)</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3><b>Worked Solution</b></h3>
<p>The distance between two vectors \(\mathbf{u}\) and&nbsp;\(\mathbf{v}\) is<br>\[ d(\mathbf{u},&nbsp;\mathbf{v}) = ||&nbsp;\mathbf{u} -\mathbf{v}||. \]</p><p>In this question we have \(\mathbf{u} = {@u@}\) and let \(\mathbf{v} = {@v@}\) so
\[ &nbsp;{@u@}-{@v@} = {@ws2@}.\]
So 
\[ ||{@ws2@}|| = {@ws3@} = {@ev(ws3,simp)@}.\] 
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[simp:true;
avar:rand_with_step(2,6,1); 
bvar:-rand_with_step(1,5,2);
cvar:rand_with_step(1,7,1);
u:matrix([avar,1,bvar]);
v:matrix([1,0, cvar]);
v3:u-v;
simp:false;
ta:sqrt(ev(v3.v3,simp));
ws2:ev(u-v,simp);
ws3:sqrt(apply("+",maplist(lambda([ex], ex^2),first(args(ws2)))));]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\(\mathbf{u} = {@u@}\) and \(\mathbf{v} = {@v@}\)</p>
<p>Distance between the points given by \(\mathbf{u}\) and \(\mathbf{v}\): {@ta@}</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>0</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>389744176</deployedseed>
    <deployedseed>1375754776</deployedseed>
    <deployedseed>2018716204</deployedseed>
    <deployedseed>1908070679</deployedseed>
    <deployedseed>1587910110</deployedseed>
    <deployedseed>1774046049</deployedseed>
    <deployedseed>870220166</deployedseed>
    <deployedseed>275676341</deployedseed>
    <deployedseed>784936241</deployedseed>
    <deployedseed>1399291022</deployedseed>
    <deployedseed>1525890805</deployedseed>
    <deployedseed>1042778948</deployedseed>
    <deployedseed>1703821793</deployedseed>
    <deployedseed>1957434598</deployedseed>
    <deployedseed>1927737733</deployedseed>
    <deployedseed>722850565</deployedseed>
    <deployedseed>1802042864</deployedseed>
    <deployedseed>329139336</deployedseed>
    <deployedseed>1837157265</deployedseed>
    <deployedseed>2092522300</deployedseed>
    <deployedseed>1341886533</deployedseed>
    <deployedseed>186680495</deployedseed>
    <deployedseed>1351334665</deployedseed>
    <deployedseed>2050811716</deployedseed>
    <deployedseed>1806560628</deployedseed>
    <deployedseed>959638542</deployedseed>
    <deployedseed>642645981</deployedseed>
    <deployedseed>2089376923</deployedseed>
    <deployedseed>1105923819</deployedseed>
    <deployedseed>2129000008</deployedseed>
    <deployedseed>648642843</deployedseed>
    <deployedseed>312189588</deployedseed>
    <deployedseed>328914819</deployedseed>
    <deployedseed>1551756687</deployedseed>
    <deployedseed>770503033</deployedseed>
    <deployedseed>320092785</deployedseed>
    <deployedseed>778242778</deployedseed>
    <deployedseed>1912213996</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev((v3.v3),simp)^0.5</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(significantfigures(ans1,4),simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
