<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Complex Number: Evaluating Fourth Roots</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find the four fourth roots of \(z={@z_real@} + {@z_imag@}i\). Gill all roots using principal argument!.</p><p>Recall that \(\sqrt[n]{r}\) is entered as \(r^{1/n}\).</p>
<p><em>Note: Write \(z\) in exponential form first using principal argument, and the four fourth roots of \(z\) in trigonometric form</em></p>
<br>\(r =\)  [[input:ans1]] and \(\theta\) = [[input:ans2]] [[validation:ans1]]&nbsp;[[validation:ans2]]
<br>\(z =\) [[input:ans3]][[validation:ans3]]
<p>[[feedback:prt3]]</p>
<br>The four fourth roots of \(z\) in trigonometric form are:
<br>\(z_{0}\) = [[input:ans4]] [[validation:ans4]][[feedback:prt4]]
<br>\(z_{1}\) = [[input:ans5]] [[validation:ans5]][[feedback:prt5]]
<br>\(z_{2}\) = [[input:ans6]] [[validation:ans6]][[feedback:prt6]]
<br>\(z_{3}\) = [[input:ans7]] [[validation:ans7]][[feedback:prt7]]
<p>[[feedback:prt1]] </p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked solution</h3>
<p>Find the three square roots of  \(z={@z_real@} {@if z_imag<0 then "-" else "+"@} {@ev(abs(z_imag),simp)@}i\). Give all roots using principal argument!. </p>
<p><em>Note:  Write \(z\) in exponential form first using its principal argument, and the three square roots of \(z\) in trigonometric form. <br>Recall that \(\sqrt[n]{r}\) is entered as \(r^{1/n}\).</em></p>
<p>Recall that if \(z=x+iy\), \(r = |z| = \sqrt{x^{2}+y^{2}}\) and depending on the quadrant in which your complex number lies, \(\theta\) is obtained from \(\tan^{-1}(y/x)\).</p>
<p><span>[[ if test="quadrant=0" ]]</span></p>
<p>In this case, your complex number lies in the first quadrant where \(\tan^{-1}(y/x) = {@arg@}\)  . </p>
<p><span>[[ elif test="quadrant=1" ]]</span></p>
<p>In this case, your complex number lies in Quadrant 2 where  \(\tan^{-1}(y/x) = -{@angle@}\)  and \(\theta = \pi+(-{@angle@}) = {@arg@}\). </p>
<p><span>[[ elif test="quadrant=2" ]]</span></p>
<p>In this case, your complex number lies in Quadrant 3 where  \(\tan^{-1}(y/x) = {@angle@}\)  and \(\theta = \pi+{@angle@} = {@arg@}\). </p>
<p><span>[[ else ]]</span></p>
<p>In this case, your complex number lies in Quadrant 4 where  \(\tan^{-1}(y/x) = -{@angle@}\)  and \(\theta = 2\pi+(-{@angle@}) = {@arg@}\). </p>
<p><span>[[/ if]]</span></p>
<p>So \(r = {@radius@}\) and \(\theta = {@ta2@}\) and hence in  exponential form, \(z=\displaystyle {@radius@} e^{i({@arg@})}\).</p>
<p> </p>
<p>We apply the following formula to get the roots, \( z_{k}= \displaystyle |z|^{@1/n@}e^{({@Theta/n@}+{@2*k*%pi/n@})i}\) where \(n=4\), and \(k=0,1,2,3\). So,</p>
 </p>
<p>\( \begin{align}
z_{k}= \displaystyle {@radius@}^{@1/4@}e^{({@ev(ta2/4,simp)@}+{@2*k*%pi/4@})i}
\end{align}
\)</p>
 </p>
<p>\( \begin{align}
\text{For} \; & k=0, & \displaystyle z_{0}= \displaystyle {@ev(radius^(1/4),simp)@} e^{{@argta4@}i}={@ev(radius^(1/4),simp)@} \left( \cos({@argta4@}) + i \sin({@argta4@}) \right) \\
\text{For} \; & k=1, & \displaystyle z_{1}= \displaystyle {@ev(radius^(1/4),simp)@} e^{{@argta5@}i} = {@ev(radius^(1/4),simp)@} \left( \cos({@argta5@}) + i \sin({@argta5@}) \right) \\
\text{For} \; & k=2, & \displaystyle z_{2}= \displaystyle {@ev(radius^(1/4),simp)@} e^{{@argta6@}i} = {@ev(radius^(1/4),simp)@} \left( \cos({@argta6@}) + i \sin({@argta6@}) \right) \\
\text{For} \; & k=3, & \displaystyle z_{3}= \displaystyle {@ev(radius^(1/4),simp)@} e^{{@argta7@}i} = {@ev(radius^(1/4),simp)@} \left( \cos({@argta7@}) + i \sin({@argta7@}) \right)
\end{align}
\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* For each of the 4 quadrants, there are 3 different argument values to choose from so that the complex number has nice values in both cartesian and polar form:
https://en.wikipedia.org/wiki/List_of_trigonometric_identities#/media/File:Unit_circle_angles_color.svg */
quadrant: rand(4);
angle: rand([%pi/2, %pi/3, %pi/4, %pi/6]);
arg: if -%pi<(angle+%pi/2*quadrant) and (angle+%pi/2*quadrant)<=%pi then (angle+%pi/2*quadrant) else (angle+%pi/2*quadrant)-2*%pi;
radius: rand_with_step(1, 9, 1);
zz: radius*%e^(%i*arg);

/*zz_cartesian does not always display the real part first then the imaginary part*/
z_real: realpart(zz);
z_imag: imagpart(zz);
zz_cartesian: z_real + z_imag * %i;

/* polarform(zz) does not give the complex number in the form zz (above) in fact it gives the complex number in cartesian form!*/
ta1: radius;

ta2: arg;

simp:false;
ta3: radius*%e^(ta2*%i);

ta89:radius*(cos(ta2) + %i*sin(ta2));

/*Fourth roots*/
simp: false;
argta4: if -%pi<(arg/4) and (arg/4)<=%pi then ev((arg/4),simp) else ev(((arg/4)-2*%pi),simp); 
ta4: radius^(1/4)*(cos(argta4) + %i *sin(argta4));
argta5: if -%pi<((arg+2*%pi)/4) and ((arg+2*%pi)/4)<=%pi then ev(((arg+2*%pi)/4),simp) else ev((((arg+2*%pi)/4)-2*%pi),simp); 
ta5: radius^(1/4)*(cos(argta5) + %i* sin(argta5));
argta6: if -%pi<((arg+4*%pi)/4) and ((arg+4*%pi)/4)<=%pi then ev(((arg+4*%pi)/4),simp) else ev((((arg+4*%pi)/4)-2*%pi),simp);
ta6: radius^(1/4)*(cos(argta6) + %i* sin(argta6));
argta7: if -%pi<((arg+6*%pi)/4) and ((arg+6*%pi)/4)<=%pi then ev(((arg+6*%pi)/4),simp) else ev((((arg+6*%pi)/4)-2*%pi),simp);
ta7: radius^(1/4)*(cos(argta7) + %i* sin(argta7));
teacher_roots: {ta4, ta5, ta6, ta7}; /*teacher fourth roots*/]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Given: \(z={@z_real@} + {@z_imag@}i\). </p>
<p>Exponential form: \(z ={@ta3@}\)</p>
<p>The four fourth roots of \(z\) in trigonometric form are:<br></p>
<p>\(z_{0}={@ta4@}\) </p>
<p>\(z_{1}={@ta5@}\) </p>
<p>\(z_{2}={@ta6@}\)</p>
<p>\(z_{3}={@ta7@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<p>Correct answer, well done.</p>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<p>Your answer is partially correct.</p>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<p>Incorrect answer.</p>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>18</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>ta5</tans>
      <boxsize>18</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6</name>
      <type>algebraic</type>
      <tans>ta6</tans>
      <boxsize>18</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans7</name>
      <type>algebraic</type>
      <tans>ta7</tans>
      <boxsize>18</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*student polar form for z*/
s_radius: ans1;
s_arg: ans2;
s_zz: s_radius*%e^(%i*s_arg);

student_roots: {ans4, ans5, ans6, ans7}; /* Make a SET of student answers (set means that duplicates are removed) */

correct_answers:intersection(student_roots,teacher_roots); 
number_correct:cardinality(correct_answers);

/*student calculation for the square root*/
s_ans4: s_radius^(1/4)*(cos((s_arg)/4) + %i *sin((s_arg)/4));
s_ans5: s_radius^(1/4)*(cos((s_arg+2*%pi)/4) + %i* sin((s_arg+2*%pi)/4));
s_ans6: s_radius^(1/4)*(cos((s_arg+4*%pi)/4) + %i* sin((s_arg+4*%pi)/4));
s_ans7: s_radius^(1/4)*(cos((s_arg+6*%pi)/4) + %i* sin((s_arg+6*%pi)/4));

/*a student can reverse the order in which they display these roots. Should that matter? If so, how do we address that in the response trees?*/

/*I am afraid that the feedback for the quadrants is not displaying. Should the two response trees be tied together for this to happen? A curious guess.*/

/*Node 1: Checks for the correct value of r */
/*Node 2: Checks for the correct measured angle from the positive real axis.*/
/*Node 3: Checks for the correct polar form using incorrect values for r and theta*/
/*Node 4: Checks for the correct polar form of the number*/
/*Node 5: checks if all roots are correct */
/*Node 6: checks if three roots are correct */
/*Node 7: checks if two roots are correct*/
/*Node 8: checks if one root is correct */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.3</truescore>
        <truepenalty>0</truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Check that you have calculated the value of \(r\) correctly.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>3</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Check that you have correctly measured the angle from the positive real axis.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.1</truescore>
        <truepenalty/>
        <truenextnode>4</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your values for r and \(\theta\) are incorrect, but you have correctly written the polar form using these values.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.2</truescore>
        <truepenalty/>
        <truenextnode>4</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Well done.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your values for r and \(\theta\) are correct, but you have not correctly written the polar form for the number. Check your syntax!</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description/>
        <answertest>Sets</answertest>
        <sans>student_roots</sans>
        <tans>teacher_roots</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You got all roots correctly.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>5</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>5</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>number_correct</sans>
        <tans>3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.75</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-6-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You got three roots correctly. Good job. You get 3/4 points.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>6</falsenextnode>
        <falseanswernote>prt1-6-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>6</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>number_correct</sans>
        <tans>2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-7-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You got two roots correctly. You get 1/2 points.</p>
<p>&nbsp;</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>7</falsenextnode>
        <falseanswernote>prt1-7-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>7</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>number_correct</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-8-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You got one root correct. You get 1/4 points.</p>
<p>&nbsp;</p>
<p>&nbsp;</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-8-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>ta5</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt6</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans6</sans>
        <tans>ta6</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt6-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt7</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans7</sans>
        <tans>ta7</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt7-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1131682157</deployedseed>
    <deployedseed>711802986</deployedseed>
    <deployedseed>319650492</deployedseed>
    <deployedseed>1170978194</deployedseed>
    <deployedseed>254548658</deployedseed>
    <deployedseed>567211467</deployedseed>
    <deployedseed>253214367</deployedseed>
    <deployedseed>1393931887</deployedseed>
    <deployedseed>48613211</deployedseed>
    <deployedseed>1625562072</deployedseed>
    <deployedseed>347596057</deployedseed>
    <deployedseed>1038380994</deployedseed>
    <deployedseed>79234418</deployedseed>
    <deployedseed>1122349339</deployedseed>
    <deployedseed>2088435604</deployedseed>
    <deployedseed>1878225035</deployedseed>
    <deployedseed>1522964326</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ta5</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>ta6</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>ta7</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-5-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt5-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt6</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt6-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt7</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt7-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
