<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Evaluating Cube Roots of Complex Numbers</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find the three cube roots of&nbsp; \(z={@z_real@} + {@z_imag@}i\). Give all roots using principal argument!</p>
<p><em>Note:  Write \(z\) in exponential form first using its principal argument, and the three square roots of \(z\) in trigonometric form. <br>Recall that \(\sqrt[n]{r}\) is entered as \(r^{1/n}\).</em></p>
<p>r = [[input:ans1]]&nbsp; and \(\theta\) = [[input:ans2]] [[validation:ans1]]&nbsp;[[validation:ans2]]</p>
<p>z = [[input:ans3]][[validation:ans3]]</p>
<p>[[feedback:prt1]]</p>
<p><span>[[ comment ]] TODO: Use EqualComAssRules for ans3 after server upgrade. [[/ comment ]]</span></p>
<p>(b) The three cube roots of \(z\) in trigonometric are;</p>
<p>\(z_{0}\) = [[input:ans4]] [[validation:ans4]]</p>
<p>\(z_{1}\) = [[input:ans5]] [[validation:ans5]]</p>
<p>\(z_{2}\) = [[input:ans6]] [[validation:ans6]]</p>
<p>[[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked solution</h3>
<p>Find the three square roots of  \(z={@z_real@} {@if z_imag<0 then "-" else "+"@} {@ev(abs(z_imag),simp)@}i\). Give all roots using principal argument!. </p>
<p><em>Note:  Write \(z\) in exponential form first using its principal argument, and the three square roots of \(z\) in trigonometric form. <br>Recall that \(\sqrt[n]{r}\) is entered as \(r^{1/n}\).</em></p>
<p>Recall that if \(z=x+iy\), \(r = |z| = \sqrt{x^{2}+y^{2}}\) and depending on the quadrant in which your complex number lies, \(\theta\) is obtained from \(\tan^{-1}(y/x)\).</p>
<p><span>[[ if test="quadrant=0" ]]</span></p>
<p>In this case, your complex number lies in the first quadrant where \(\tan^{-1}(y/x) = {@arg@}\)  . </p>
<p><span>[[ elif test="quadrant=1" ]]</span></p>
<p>In this case, your complex number lies in Quadrant 2 where  \(\tan^{-1}(y/x) = -{@angle@}\)  and \(\theta = \pi+(-{@angle@}) = {@arg@}\). </p>
<p><span>[[ elif test="quadrant=2" ]]</span></p>
<p>In this case, your complex number lies in Quadrant 3 where  \(\tan^{-1}(y/x) = {@angle@}\)  and \(\theta = \pi+{@angle@} = {@arg@}\). </p>
<p><span>[[ else ]]</span></p>
<p>In this case, your complex number lies in Quadrant 4 where  \(\tan^{-1}(y/x) = -{@angle@}\)  and \(\theta = 2\pi+(-{@angle@}) = {@arg@}\). </p>
<p><span>[[/ if]]</span></p>
<p>So \(r = {@radius@}\) and \(\theta = {@ta2@}\) and hence in  exponential form, \(z=\displaystyle {@radius@} e^{i({@arg@})}\).</p>
<p> </p>
<p>We apply the following formula to get the roots, \( z_{k}= \displaystyle |z|^{@1/n@}e^{({@Theta/n@}+{@2*k*%pi/n@})i}\) where \(n=3\), and \(k=0,1,2\). So,</p>
 </p>
<p>\( \begin{align}
z_{k}= \displaystyle {@radius@}^{@1/3@}e^{({@ev(ta2/3,simp)@}+{@2*k*%pi/3@})i}
\end{align}
\)</p>
 </p>
<p>\( \begin{align}
\text{For} \; & k=0, & \displaystyle z_{0}= \displaystyle {@ev(radius^(1/3),simp)@} e^{{@arg4@}i}={@ev(radius^(1/3),simp)@} \left( \cos({@arg4@}) + i \sin({@arg4@}) \right) \\
\text{For} \; & k=1, & \displaystyle z_{1}= \displaystyle {@ev(radius^(1/3),simp)@} e^{{@arg5@}i} = {@ev(radius^(1/3),simp)@} \left( \cos({@arg5@}) + i \sin({@arg5@}) \right) \\
\text{For} \; & k=2, & \displaystyle z_{2}= \displaystyle {@ev(radius^(1/3),simp)@} e^{{@arg6@}i} = {@ev(radius^(1/3),simp)@} \left( \cos({@arg6@}) + i \sin({@arg6@}) \right)
\end{align}
\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* For each of the 4 quadrants, there are 3 different argument values to choose from so that the complex number has nice values in both cartesian and polar form: https://en.wikipedia.org/wiki/List_of_trigonometric_identities#/media/File:Unit_circle_angles_color.svg */
quadrant: rand(4);

angle: rand([%pi/2, %pi/3, %pi/4, %pi/6]);
arg: if -%pi<(angle+%pi/2*quadrant) and (angle+%pi/2*quadrant)<=%pi then (angle+%pi/2*quadrant) else (angle+%pi/2*quadrant)-2*%pi;
radius: rand_with_step(1, 9, 1);
zz: radius*%e^(%i*arg);

/*zz_cartesian does not always display the real part first then the imaginary part*/
z_real: realpart(zz);
z_imag: imagpart(zz);
zz_cartesian: z_real + z_imag * %i;

/* polarform(zz) does not give the complex number in the form zz (above) in fact it gives the complex number in cartesian form!*/
ta1: radius;
ta2: arg;

simp:false;
ta3: radius*%e^(ta2*%i);

/*cube roots*/
wk: %e^(((ta2+2*k*%pi)/3)*%i);
arg4: if -%pi<(arg/3) and (arg/3)<=%pi then ev((arg/3),simp) else ev((arg/3)-2*pi,simp);
ta4: radius^(1/3)*(cos(arg4)+%i *sin(arg4));

arg5: if -%pi<((arg+2*%pi)/3) and ((arg+2*%pi)/3)<=%pi then ev(((arg+2*%pi)/3),simp) else ev(((arg+2*%pi)/3)-2*pi,simp);
ta5: radius^(1/3)*(cos(arg5)+%i* sin(arg5));

arg6: if -%pi<((arg+4*%pi)/3) and ((arg+4*%pi)/3)<=%pi then ev(((arg+4*%pi)/3),simp) else ev(((arg+4*%pi)/3)-2*pi,simp);
ta6: radius^(1/3)*(cos(arg6)+%i* sin(arg6));
teacher_roots: {ta4, ta5, ta6}; /* Set of cube roots for roots z*/]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\(z={@z_real@} + {@z_imag@}i\). </p>
<p> \(r = {@ta1@} \) and \( \theta = {@ta2@} \) </p>
<p> \( z = {@radius*%e^(arg*%i)@} \) </p>
<p>The three cube roots of \(z\) in trigonometric are;</p>
<p>\(z_{0} = {@ta4@} \)</p>
<p>\(z_{1} = {@ta4@} \)</p>
<p>\(z_{2} = {@ta4@} \)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<p>Correct answer, well done.</p>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<p>Your answer is partially correct.</p>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<p>Incorrect answer.</p>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>radius*%e^(arg*%i)</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>18</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>ta5</tans>
      <boxsize>18</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6</name>
      <type>algebraic</type>
      <tans>ta6</tans>
      <boxsize>18</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct value of r */
/*Node 2: Checks for the correct measured angle from the positive real axis*/
/*Node 3: Checks for the correct polar form using the incorrect values for r and theta */
/*Node 4: Checks for the correct polar form*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.3</truescore>
        <truepenalty>0</truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Check that you have calculated the value of \(r\) correctly.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>3</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Check that you have correctly measured the angle from the positive real axis.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>radius*%e^(arg*%i)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your values for r and \(\theta\) are incorrect, but you have correctly written the polar form using these values.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.2</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Well done.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your values for r and \(\theta\) are correct, but you have not correctly written the polar form for the number. Check your syntax!</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>simp: true;

/*student polar form for z*/
s_radius: ans1;
s_arg: ans2;
s_zz: s_radius*%e^(%i*s_arg);

student_roots: {ans4, ans5, ans6}; /* Make a SET of student answers (set means that duplicates are removed) */

correct_answers: intersection(student_roots,teacher_roots); 
number_correct: cardinality(correct_answers);

/* student calculation for the square root */
s_ans4: s_radius^(1/3)*(cos((s_arg)/3) + %i *sin((s_arg)/3));
s_ans5: s_radius^(1/3)*(cos((s_arg+2*%pi)/3) + %i* sin((s_arg+2*%pi)/3));
s_ans6: s_radius^(1/3)*(cos((s_arg+4*%pi)/3) + %i* sin((s_arg+4*%pi)/3));

/* Proof of concept of using set comparison and carrying student errors forward. */
/* This is VERY crude and needs polishing. */
student_roots_based_on_their_polar_form: {s_ans4, s_ans5, s_ans6};

/*Node 1: Checks if all the three roots are correct */
/*Node 2: Checks if two roots are correct*/
/*Node 3: Checks if one root is correct*/
/*Node 4: Checks for the correct cube root of the polar form thou with incorrect polar form */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>student_roots</sans>
        <tans>teacher_roots</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><span>You found all three roots! You get full marks.</span></p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>number_correct</sans>
        <tans>2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.67</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><span>You got two out of three correct. Way to go!! You get 2/3 marks</span></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>number_correct</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.33</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><span>You got one out of three correct. Nice attempt. You get 1/3 marks</span></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt2-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description/>
        <answertest>Sets</answertest>
        <sans>student_roots</sans>
        <tans>student_roots_based_on_their_polar_form</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><span>It looks like your polar form was incorrect, but based on it, you managed to compute the cube roots of your polar form correctly.</span></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1407213946</deployedseed>
    <deployedseed>704834624</deployedseed>
    <deployedseed>640074683</deployedseed>
    <deployedseed>1159624988</deployedseed>
    <deployedseed>1166434098</deployedseed>
    <deployedseed>822431356</deployedseed>
    <deployedseed>1755439995</deployedseed>
    <deployedseed>1522004085</deployedseed>
    <deployedseed>1572060925</deployedseed>
    <deployedseed>1326911407</deployedseed>
    <deployedseed>1387166939</deployedseed>
    <deployedseed>722979611</deployedseed>
    <deployedseed>60958757</deployedseed>
    <deployedseed>1611721521</deployedseed>
    <deployedseed>592448594</deployedseed>
    <deployedseed>913342531</deployedseed>
    <deployedseed>843726792</deployedseed>
    <deployedseed>790412937</deployedseed>
    <deployedseed>1450177629</deployedseed>
    <deployedseed>1920168092</deployedseed>
    <deployedseed>110573939</deployedseed>
    <deployedseed>3532375</deployedseed>
    <deployedseed>76629978</deployedseed>
    <deployedseed>53281423</deployedseed>
    <deployedseed>1396624657</deployedseed>
    <deployedseed>1988170897</deployedseed>
    <deployedseed>1134940153</deployedseed>
    <deployedseed>1943351950</deployedseed>
    <deployedseed>1817091293</deployedseed>
    <deployedseed>502568235</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>radius*%e^(arg*%i)</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ta5</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>ta6</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-4-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
