<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Quadratic Word Problems</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Solve the following word problems leading to quadratic equations.</p>

<br>

<p>(a) The difference between a number and its positive square root is {@difference@}.</p>
<p>Find the number.</p>
<p style="margin-left: 30px">[[input:ans1]] [[validation:ans1]]</p>
<p>[[feedback:prt1]]</p>

<br>

<p>(b) The sides of an equilateral triangle are each shortened by \({@c@} \text{cm}\), \({@a@} \text{cm}\) and \({@b@} \text{cm}\) respectively. The result is a right-angled triangle.</p>
<p>Find the length of the sides of the equilateral triangle.</p>
<p style="margin-left: 30px">[[input:ans2]] [[validation:ans2]] \(\text{cm}\)</p>
<p>[[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>(a) We will call the required number \(n\).</p>
<p>If the difference between the number and it’s positive square root is {@difference@}, we have,</p>

<p style="margin-left: 30px">\(\begin{align*} 
n-\sqrt{n} = {@difference@}
& \Rightarrow n-{@difference@} = \sqrt{n} \\
& \Rightarrow (n-{@difference@})^{2} = (\sqrt{n})^{2} \\
& \Rightarrow {@expand((n-difference)^2)@} = n \\
& \Rightarrow {@expand((n-difference)^2) - n@} = 0
\end{align*}\)</p>

<p>So, we have derived a quadratic equation for \(n\) which we may solve by factorizing,</p>

<p style="margin-left: 30px">\(\begin{align*}
{@expand((n-difference)^2)-n@} &= {@eqn@} \\
&= 0
\end{align*}\)</p>
<p style="margin-left: 30px">\(\Rightarrow n = {@n1@}, n = {@n2@}\)</p>

<br>

<p>By squaring the equation, we may have introduced additional solutions.
<br>So, we need to verify that the solutions we have derived actually solve the original problem.</p>

<p>Firstly, we consider \(n = {@ta1@}\),</p>
<p style="margin-left: 30px">\({@ta1@} - \sqrt{{@ta1@}} = {@difference@}\)</p>
<p>Therefore, this solution <b>does</b> solve the original problem.</p>

<p>Then, we consider \(n = {@wrong_sol@}\),</p>
<p style="margin-left: 30px">\({@wrong_sol@} - \sqrt{{@wrong_sol@}} = {@wrong_sol-sqrt(wrong_sol)@} \neq {@difference@}\)</p>
<p>Therefore, this solution <b>does not</b> solve the original problem.</p>

<br>

<p>Hence, the only solution to the original problem is {@ta@}.</p>

<br>

<p>(b) We will call the length of the sides of the equilateral triangle \(x\).</p>
<p>Then, the sides of right-angled triangle have lengths \({@hypotenuse@}, {@height@}\) and \({@base@}\).</p>

<p>The side with length {@hypotenuse@} is the longest side in the right-angled triangle, so we can apply Pythagoras' Theorem so give,</p>

<p style="margin-left: 30px">\(\begin{align*}
a^2 + b^2  = c^2 &\Rightarrow {@height^2@} + {@base^2@} = {@hypotenuse^2@} \\
&\Rightarrow ({@expand(height^2)@}) + ({@expand(base^2)@}) = ({@expand(hypotenuse^2)@}) \\
&\Rightarrow {@expr4@} = 0 
\end{align*}\)</p>

<p>So, we have derived a quadratic equation for \(x\) which we may solve by factorizing,</p>

<p style="margin-left: 30px">\(\begin{align*}
{@expr4@} &= {@expr@} \\
&= 0
\end{align*}\)</p>
<p style="margin-left: 30px">\(\Rightarrow x = {@x1@}, x = {@x2@}\)</p>

<br>

<p>Considering now the context of the problem, since they represent the lengths of the sides of a triangle the values of \({@hypotenuse@}, {@height@}\) and \({@base@}\) must be <b>positive</b>.</p>
<p>Therefore, \(x \neq {@x2@}\)</p>

<p>So, the length of each side of the equilateral triangle is \({@ta2@} \text{cm}\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>2</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*Question 1 variables*/
aa:rand(5)+1;
bb:aa+1;
n1:aa^2;   /*first value of n*/
n2:bb^2; /*second value of n*/
eqn:(n-n1)*(n-n2); 
eqn1:expand(eqn);
cc:coeff(eqn1, n, 0);
difference: sqrt(cc);
ta1: if n1-sqrt(n1)=difference then n1 else n2;
wrong_sol: if n1-sqrt(n1)=difference then n2 else n1;

/*Question 2 variables*/
x2:rand(10)+13;
x1:x2+4;
expr:(x-x2)*(x-x1);
expr1:expand(eqn1);
c:x1-5; hypotenuse:x-c;
a:x1-4; height:x-a;
b:x1-3; base:x-b;
expr2:height^2+base^2=hypotenuse^2;
expr3:expand(height^2)+expand(base^2)=expand(hypotenuse^2);
expr4:expand(height^2+base^2-hypotenuse^2);
expr5:solve(expr4=0);
ta2:x1;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Difference between the two numbers: {@difference@} </p>
<p>Equation: {@eqn1@}</p>
<p>Number: {@ta1@}</p>

<br>

<p>Side A of the triangle: {@c@} cm</p>
<p>Side B of the triangle:{@a@} cm</p>
<p>Side C of the triangle:{@b@} cm</p>
<p>length of the triangle={@ta2@}</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*node 1 checks for correct answer*/
/*node 2: inputting the solution that does not work */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>wrong_sol</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>The answer that you have given is a correct root of the quadratic equation which represents this problem.
<br>However, this solution <b>does not work</b> in the context of the original problem, since we require,</p>
<p style="margin-left: 30px">\(n-\sqrt{n}={@difference@}\)</p>
<p>Whereas for your answer,</p>
<p  style="margin-left: 30px">\({@wrong_sol@} - \sqrt{{@wrong_sol@}} = {@wrong_sol - sqrt(wrong_sol)@} \neq {@difference@}\)</p>
<p>The additional solution has been introduced by squaring the equation in the process of deriving the quadratic.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*node 1 checks for correct answer*/
/*node 2 mistake: inputting the wrong value of x*/
mistake_2:x2;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>mistake_2</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>The answer that you have given is a correct root of the quadratic equation which represents this problem.
<br>However, this solution <b>does not work</b> in the context of the original problem, since we require that \({@hypotenuse@}, {@height@}\) and \({@base@}\) are the lengths of the sides of a triangle and therefore <b>positive</b>.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1206777755</deployedseed>
    <deployedseed>916589448</deployedseed>
    <deployedseed>1478231218</deployedseed>
    <deployedseed>2122217328</deployedseed>
    <deployedseed>763535125</deployedseed>
    <deployedseed>516751806</deployedseed>
    <deployedseed>1817909481</deployedseed>
    <deployedseed>40675066</deployedseed>
    <deployedseed>1336065435</deployedseed>
    <deployedseed>1110760029</deployedseed>
    <deployedseed>1451791203</deployedseed>
    <deployedseed>545397653</deployedseed>
    <deployedseed>1953506160</deployedseed>
    <deployedseed>1803155021</deployedseed>
    <deployedseed>557802148</deployedseed>
    <deployedseed>262638499</deployedseed>
    <deployedseed>945316491</deployedseed>
    <deployedseed>194792166</deployedseed>
    <deployedseed>244205018</deployedseed>
    <deployedseed>1548097842</deployedseed>
    <deployedseed>1400278927</deployedseed>
    <deployedseed>2050928679</deployedseed>
    <deployedseed>614838921</deployedseed>
    <deployedseed>1762214784</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
