<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text><![CDATA[Solve the inequality of the form p < b-ax]]></text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p> Solve \({@ee@} \leq {@bb-aa*x@} < {@dd@}\), as a subset of the real numbers represented by inequalities.
<br>Also write this inequalities as interval of the real numbers.</p>
</p>

<br>

<p>(a) Give your answer as inequalities for \(x\).</p>
<p><em>Note: To enter inequalities, you should write the relevant symbols in the table below.</em></p>

<table cellspacing="1" cellpadding="0" border="1" style="margin-left: 30px">
  <tr>
    <th width="200px" style="text-align: center;">Inequalities symbol</th>
    <th width="150px" style="text-align: center;">Input Syntax</th>
  </tr>
  <tr>
    <td style="text-align: center;">\( \geq \)</td>
    <td style="text-align: center;"><code>>=</code></td>
  </tr>
  <tr>
    <td style="text-align: center;">\( \leq \)</td>
    <td style="text-align: center;"><code><=</code></td>
  </tr>
  <tr>
    <td style="text-align: center;">\( <\)</td>
    <td style="text-align: center;"><code><</code></td>
  </tr>
  <tr>
    <td style="text-align: center;">\( >\)</td>
    <td style="text-align: center;"><code>></code></td>
  </tr>
 <tr>
    <td style="text-align: center;">\( a < x < b \)</td>
    <td style="text-align: center;"><code>a < x and x < b</code></td>
  </tr>
 <tr>
    <td style="text-align: center;">\( x > a\) or \(x < b \)</td>
    <td style="text-align: center;"><code>a < x or x < b</code></td>
  </tr>
</table>

<br>

<p style="margin-left: 30px">[[input:ans1]] [[validation:ans1]]</p>
<p>[[feedback:prt1]]</p>

<br>

<p>(b) Give your answer as interval in the simplified form.</p>
<p><em>Note: To enter intervals of you real line, you should write the relevant text in the table below.</em></p>

<table width="250" cellspacing="1" cellpadding="0" border="1" style="margin-left: 30px">
  <tr>
    <th style="text-align: center;">Intervals</th>
    <th style="text-align: center;">Input Syntax</th>
  </tr>
  <tr>
    <td style="text-align: center;">\( [a,b] \)</td>
    <td style="text-align: center;"><code>cc(a,b)</code></td>
  </tr>
  <tr>
    <td style="text-align: center;">\( (a,b) \)</td>
    <td style="text-align: center;"><code>oo(a,b)</code></td>
  </tr>
  <tr>
    <td style="text-align: center;">\( [a,b) \)</td>
    <td style="text-align: center;"><code>co(a,b)</code></td>
  </tr>
  <tr>
    <td style="text-align: center;">\( (a,b] \)</td>
    <td style="text-align: center;"><code>oc(a,b)</code></td>
  </tr>
 <tr>
    <td style="text-align: center;">\( (a, \infty) \)</td>
    <td style="text-align: center;"><code>oo(a,inf)</code></td>
  </tr>
 <tr>
    <td style="text-align: center;">\( (-\infty, b] \)</td>
    <td style="text-align: center;"><code>oc(-inf,b)</code></td>
  </tr>
 <tr>
    <td style="text-align: center;">\( (-\infty, \infty) \)</td>
    <td style="text-align: center;"><code>oo(-inf,inf)</code></td>
  </tr>
</table>

<br>

<p style="margin-left:30px">[[input:ans2]] [[validation:ans2]]</p>
<p>[[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>(a) We want to solve \({@ee@} \leq {@bb-aa*x@} < {@dd@}\)
<br>Separating both sides of the inequality, we have that,</p>

<p style="margin-left: 30px">\((1) : {@ee@} \leq {@bb-aa*x@}\)</p>
<p style="margin-left: 30px">\((2) : {@bb-aa*x@} < {@dd@}\)</p>

<p>First, we solve the direction \({@ee@} \leq {@bb-aa*x@}\)</p>

<p style="margin-left: 30px">\(\begin{align*}
{@ee@} \leq {@bb-aa*x@} &\implies {@aa*x@} \leq {@bb-ee@} \\
& \implies x \leq \textstyle {@(bb-ee)/aa@} \\
& \implies x \in \textstyle {@oc(-inf, ((bb-ee)/aa))@}
\end{align*}\)</p>

<p>Next, we solve the direction \({@bb-aa*x@} < {@dd@}\)</p>
<p style="margin-left: 30px">\(\begin{align*}
{@bb-aa*x@} < {@dd@} & \implies {@-aa*x@} < {@dd-bb@} \\
& \implies {@aa*x@} > {@-dd+bb@} \\
& \implies x > \textstyle {@(-dd+bb)/aa@} \\
& \implies x \in \textstyle {@oo(((-dd+bb)/aa), inf)@}
\end{align*}\)
</p>

<p>Hence, the solution of the inequality is given by,</p>
<p style="margin-left: 30px">{@ta@}</p>

<br>

<p>(b) Representing these inequalities as an interval of the real line, we have,</p>
<p style="margin-left: 30px">\(\begin{align*}
\{x \in \mathbb{R} : \textstyle {@ta@} \} &= \textstyle {@int_3@} \\
&= \textstyle {@int_4@}
\end{align*}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[aa:rand(9)+2;
bb:rand_with_prohib(-6,6,[0]);
ee:rand([-10,-2]);
dd:rand([1,10]);
quest:ee<=bb-aa*x and bb-aa*x<dd;
ta:x>(bb-dd)/aa and x<=(bb-ee)/aa;
/*Express answer in terms of real intervals*/
int_1: oc(-inf,((bb-ee)/aa));
int_2: oo(((-dd+bb)/aa),inf);
int_3:%intersection(int_1,int_2);
int_4:oc(((-dd+bb)/aa),((bb-ee)/aa));]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Question: \({@ee@} \leq {@bb-aa*x@} &lt; {@dd@}\) </p>
<p>Answer: \({@ta@}, {@int_4@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>int_4</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Worked Solution:</p><p>\({@cc@} \leq {@bb-aa*x@} &lt; {@dd@}\)<br></p><p>We solve both inequalities separately:</p><p>\({@cc@} \leq {@bb-aa*x@}\)<br></p><p>\({@cc-bb@} \leq {@-aa*x@}\)</p><p>\({@(bb-cc)/aa@} \geq x\)</p><p>and,</p><p>\({@bb-aa*x@} &lt; {@dd@}\)<br></p><p>\({@-aa*x@} &lt; {@dd-bb@}\)</p><p>\(x &gt; {@(bb-dd)/aa@}\).</p><p>Therefore, the answer to this inequality is the interval \({@(bb-dd)/aa@} &lt; x \leq {@(bb-cc)/aa@}\), which is entered as \(x &gt;{@(bb-dd)/aa@}\) and \(x\leq {@(bb-cc)/aa@}\).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>int_4</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>int_3</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<b>Your answer is correct.
<br>However, you can simplify the intersection you have given into a single interval.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1285286328</deployedseed>
    <deployedseed>658618425</deployedseed>
    <deployedseed>475394919</deployedseed>
    <deployedseed>1755014324</deployedseed>
    <deployedseed>250347540</deployedseed>
    <deployedseed>682437141</deployedseed>
    <deployedseed>809445815</deployedseed>
    <deployedseed>724524359</deployedseed>
    <deployedseed>534577887</deployedseed>
    <deployedseed>1532629935</deployedseed>
    <deployedseed>1148886307</deployedseed>
    <deployedseed>379957600</deployedseed>
    <deployedseed>1768630900</deployedseed>
    <deployedseed>1933469280</deployedseed>
    <deployedseed>1800857423</deployedseed>
    <deployedseed>1815961250</deployedseed>
    <deployedseed>444433628</deployedseed>
    <deployedseed>1231266427</deployedseed>
    <deployedseed>1859131752</deployedseed>
    <deployedseed>461114974</deployedseed>
    <deployedseed>543935406</deployedseed>
    <deployedseed>650615788</deployedseed>
    <deployedseed>1042944950</deployedseed>
    <deployedseed>388426223</deployedseed>
    <deployedseed>1847746818</deployedseed>
    <deployedseed>901717928</deployedseed>
    <deployedseed>821570399</deployedseed>
    <deployedseed>511685307</deployedseed>
    <deployedseed>1813589889</deployedseed>
    <deployedseed>1904700846</deployedseed>
    <deployedseed>778071903</deployedseed>
    <deployedseed>1578312635</deployedseed>
    <deployedseed>479014977</deployedseed>
    <deployedseed>693223155</deployedseed>
    <deployedseed>452176948</deployedseed>
    <deployedseed>246376950</deployedseed>
    <deployedseed>61105159</deployedseed>
    <deployedseed>1162320955</deployedseed>
    <deployedseed>398420216</deployedseed>
    <deployedseed>143158500</deployedseed>
    <deployedseed>349328998</deployedseed>
    <deployedseed>1016048505</deployedseed>
    <deployedseed>889561990</deployedseed>
    <deployedseed>1439053533</deployedseed>
    <deployedseed>1470186678</deployedseed>
    <deployedseed>214496181</deployedseed>
    <deployedseed>1723736159</deployedseed>
    <deployedseed>1751133190</deployedseed>
    <deployedseed>2007681139</deployedseed>
    <deployedseed>1120904914</deployedseed>
    <deployedseed>613201881</deployedseed>
    <deployedseed>382477192</deployedseed>
    <deployedseed>685399787</deployedseed>
    <deployedseed>994679998</deployedseed>
    <deployedseed>7761923</deployedseed>
    <deployedseed>632676099</deployedseed>
    <deployedseed>316381454</deployedseed>
    <deployedseed>1339061063</deployedseed>
    <deployedseed>321993396</deployedseed>
    <deployedseed>841701799</deployedseed>
    <deployedseed>1682461049</deployedseed>
    <deployedseed>468033015</deployedseed>
    <deployedseed>1839687231</deployedseed>
    <deployedseed>192739894</deployedseed>
    <deployedseed>1120915452</deployedseed>
    <deployedseed>968958305</deployedseed>
    <deployedseed>1655185521</deployedseed>
    <deployedseed>2141176686</deployedseed>
    <deployedseed>1586159357</deployedseed>
    <deployedseed>946162166</deployedseed>
    <deployedseed>1520216822</deployedseed>
    <deployedseed>1653923652</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
