<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Find the cardinality of the largest possible set</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>(a) Consider the following sets: </p>
<p>\(A={@setA@}\) and \(B={@setB@} \).</p>
<p>Find the cardinality of the largest possible set \(C\) such that \(C \subseteq A\) and \(C \subseteq B\).</p><p>[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>

<br>

<p>(b) For this part, the universal set is {@univ2@}.</p>
<p>Find examples of sets \(D\) and \(E\) such that \(|D| = {@kk@}\), \(|E| = {@ll@}\) and \(|D \cup E| = {@mm@}\)</p>
<p><em>Note: A set is entered in the blanks using curly brackets notation and double quotes.  As example, we have: <code>{"apple","orange","mango"}</em></code>.</p>
<p>\(D = \) [[input:ans2]] [[validation:ans2]]</p>
<p> \(E = \) [[input:ans3]]  [[validation:ans3]] </p>
<p> [[feedback:prt2]] </p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked solution</h3>
<p>(a) Consider the following sets: </p>
<p>\(A={@setA@}\) and \(B={@setB@} \).</p>
<p>To find the cardinality of the largest possible set \(C\) such that \(C \subseteq A\) and \(C \subseteq B\), the first step is to determine the set corresponding to \( A \cap B\).</p>
<p>We have: \( A \cap B = {@intAB@}\) and Card\((A \cap B)={@ta1@} \). Therefore, the cardinality of the largest possible set \(C\) satisfying the given condition above is \({@ta1@}\).</p>

<br>

<p>(b) For this part, the universal set is {@univ2@}.</p>
<p>Let find examples of sets \(D\) and \(E\) such that \(|D| = {@kk@}\), \(|E| = {@ll@}\) and \(|D \cup E| = {@mm@}\)</p>
<p>There are many possible ways to write the sets \(D\) and \(E\) in this case. Thus, we can choose:</p>
<p>\(D={@set_D@}\) and \(E={@set_E@}\). Then  \(D \cup E = {@set_DUE@}\).]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[aa:rand(12)+1;
bb:rand_with_prohib(1,12,[aa]);
cc:rand_with_prohib(1,12,[aa,bb]);
dd:rand_with_prohib(1,12,[aa,bb,cc]);
ee:rand_with_prohib(1,12,[aa,bb,cc,dd]);
ff:rand_with_prohib(1,12,[aa,bb,cc,dd,ee]);
gg:rand_with_prohib(1,12,[aa,bb]);
hh:rand_with_prohib(1,12,[aa,bb,gg]);
ii:rand_with_prohib(1,12,[aa,bb,gg,hh]);
jj:rand_with_prohib(1,12,[aa,bb,gg,hh,ii]);
setA:set(aa,cc,dd,ee,ff);
setB:set(bb,gg,hh,ii,jj);
intAB:intersect(setA,setB);
ta1:cardinality(intAB);
list1:["apple","banana","coconut","grape","lemon","lime","mango","melon","orange","pear"];
univ2:setify(list1);
/*Worked solution*/
numList1: [rand(2)+1,rand(4)+1,rand(3)+4,3+rand(2)+5,rand(5)+5,rand(8)+3,rand(6)+4,rand(6)+2,rand(7)+1,rand(9)+2,rand(10)+1];
evens: [];
odds:[];
for i in numList1 do (
    if evenp(i) then
        evens:append(evens,[i])
    else
        odds:append(odds,[i])
);
evens;
odds;

D:makelist(list1[e],e,evens);
E:makelist(list1[j],j,odds);
set_D: setify(D);
set_E: setify(E);
set_DUE: union(set_D,set_E);
kk: cardinality(set_D);
ll: cardinality(set_E);
mm: cardinality(set_DUE);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Largest cardinality of \(C\) is \({@ta1@}\).</p>
<p> \( \(D={@set_D@}\) </p>
<p>\(E={@set_E@}\). </p>
<p> \(D \cup E = {@set_DUE@}\).</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>set_D</tans>
      <boxsize>25</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>set_E</tans>
      <boxsize>25</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>You should get \({@ta1@}\). Check the worked solution.</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node1*/
is_card_D: is(cardinality(ans2)=kk);
is_card_E: is(cardinality(ans3)=ll);
find_ans2_U_ans3: union(ans2,ans3);
is_card_union: is(cardinality(find_ans2_U_ans3)=mm);
cor_ans: is(is_card_D=true and is_card_E=true and is_card_union=true);

/*Node 2*/
cor_ans1: is(cardinality(ans2)=kk) and is(cardinality(ans3) # ll) and is(intersect(ans2, univ2)=ans2);

/*Node 3*/
cor_ans2: is(cardinality(ans2) # kk) and is(cardinality(ans3)=ll) and is(intersect(ans3, univ2)=ans3);

/*Node 4*/
cor_ans3: is(intersect(ans2, univ2)=ans2) and is(intersect(ans3, univ2)=ans3) and is(cardinality(ans2)=kk) and is(cardinality(ans3)=ll) and is(cardinality(union(ans2,ans3)) # mm);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>cor_ans</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>cor_ans1</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1/2</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>The entry \(E={@ans3@}\) has cardinality \({@cardinality(ans3)@} \neq {@ll@}\)  which should not be the case. However \(|D|={@cardinality(ans2)@}\), and its elements are elements of the universal set too, which is correct. </p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>cor_ans2</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1/2</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>The entry \(D={@ans2@}\) has cardinality {@cardinality(ans2)@} \neq {@kk@} \), which should not be the case. However \(|E|={@cardinality(ans3)@}\), and the elements of \(E\) are elements of the universal set too, which is correct. </p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt2-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>cor_ans3</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>3/4</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-4-T</trueanswernote>
        <truefeedback format="html">
          <text>The given sets \(D\) and \(E\) are correct, but their union is given by \({@union(ans2,ans3)@}\), which has cardinality \( {@cardinality(union(ans2,ans3))@} \neq {@mm@} \)</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1907970449</deployedseed>
    <deployedseed>1854194736</deployedseed>
    <deployedseed>1643680735</deployedseed>
    <deployedseed>1451282738</deployedseed>
    <deployedseed>1449022998</deployedseed>
    <deployedseed>524137079</deployedseed>
    <deployedseed>106402998</deployedseed>
    <deployedseed>950700224</deployedseed>
    <deployedseed>1133872362</deployedseed>
    <deployedseed>449653837</deployedseed>
    <deployedseed>725445227</deployedseed>
    <deployedseed>162729186</deployedseed>
    <deployedseed>1524350619</deployedseed>
    <deployedseed>1803776643</deployedseed>
    <deployedseed>1735593358</deployedseed>
    <deployedseed>895191355</deployedseed>
    <deployedseed>56730266</deployedseed>
    <deployedseed>1134142747</deployedseed>
    <deployedseed>521307395</deployedseed>
    <deployedseed>673817500</deployedseed>
    <deployedseed>1952454234</deployedseed>
    <deployedseed>1097910112</deployedseed>
    <deployedseed>2100079145</deployedseed>
    <deployedseed>944408126</deployedseed>
    <deployedseed>973267828</deployedseed>
    <deployedseed>2101745162</deployedseed>
    <deployedseed>1097176574</deployedseed>
    <deployedseed>2020200489</deployedseed>
    <deployedseed>1489880754</deployedseed>
    <deployedseed>621540958</deployedseed>
    <deployedseed>1388844779</deployedseed>
    <deployedseed>758809384</deployedseed>
    <deployedseed>595088113</deployedseed>
    <deployedseed>213603275</deployedseed>
    <deployedseed>49000482</deployedseed>
    <deployedseed>1967474469</deployedseed>
    <deployedseed>983003913</deployedseed>
    <deployedseed>1797452668</deployedseed>
    <deployedseed>1281600618</deployedseed>
    <deployedseed>2042990009</deployedseed>
    <deployedseed>1532822317</deployedseed>
    <deployedseed>551562539</deployedseed>
    <deployedseed>1726367360</deployedseed>
    <deployedseed>810444967</deployedseed>
    <deployedseed>667018907</deployedseed>
    <deployedseed>1538332265</deployedseed>
    <deployedseed>80854437</deployedseed>
    <deployedseed>303450872</deployedseed>
    <deployedseed>243809733</deployedseed>
    <deployedseed>2097196159</deployedseed>
    <deployedseed>289349975</deployedseed>
    <deployedseed>1434188767</deployedseed>
    <deployedseed>1464051462</deployedseed>
    <deployedseed>652681880</deployedseed>
    <deployedseed>1679640571</deployedseed>
    <deployedseed>1974989046</deployedseed>
    <deployedseed>991283768</deployedseed>
    <deployedseed>1308580772</deployedseed>
    <deployedseed>591020836</deployedseed>
    <deployedseed>2059289862</deployedseed>
    <deployedseed>1124927751</deployedseed>
    <deployedseed>834327533</deployedseed>
    <deployedseed>2085032970</deployedseed>
    <deployedseed>152937834</deployedseed>
    <deployedseed>1824808859</deployedseed>
    <deployedseed>456455746</deployedseed>
    <deployedseed>629411701</deployedseed>
    <deployedseed>1563761906</deployedseed>
    <deployedseed>1480997485</deployedseed>
    <deployedseed>1779987646</deployedseed>
    <deployedseed>1203812470</deployedseed>
    <deployedseed>1081860675</deployedseed>
    <deployedseed>899245224</deployedseed>
    <deployedseed>1797771115</deployedseed>
    <deployedseed>1014231025</deployedseed>
    <deployedseed>811273719</deployedseed>
    <deployedseed>1652925496</deployedseed>
    <deployedseed>975901992</deployedseed>
    <deployedseed>1226631807</deployedseed>
    <deployedseed>1833546344</deployedseed>
    <deployedseed>731704964</deployedseed>
    <deployedseed>931620524</deployedseed>
    <deployedseed>1731612496</deployedseed>
    <deployedseed>920355461</deployedseed>
    <deployedseed>1139456560</deployedseed>
    <deployedseed>159771906</deployedseed>
    <deployedseed>557353800</deployedseed>
    <deployedseed>1444606545</deployedseed>
    <deployedseed>269064060</deployedseed>
    <deployedseed>120064275</deployedseed>
    <deployedseed>864181849</deployedseed>
    <deployedseed>2072876825</deployedseed>
    <deployedseed>998524692</deployedseed>
    <deployedseed>1783370935</deployedseed>
    <deployedseed>644953448</deployedseed>
    <deployedseed>574636514</deployedseed>
    <deployedseed>1862031878</deployedseed>
    <deployedseed>1277608111</deployedseed>
    <deployedseed>562596665</deployedseed>
    <deployedseed>1075373862</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>set_D</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>set_E</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>set_D</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>set_E</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
