<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Set Definition and Statement Evaluation</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>(a) Set \(A\) is defined as \(A = \{ x: x \in \mathbb{Z} \text{ and } {@lower_bound_a@} \leq x \leq {@upper_bound_a@}  \} \). 
<br>Write down set \(A\) by specifying its elements.</p>
<p>\(A = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>

<br>

<p>(b) Set \(B\) is defined as \(B = \{ x: x \in \mathbb{N} \text{ and } {@lower_bound_b@} \leq x \leq {@upper_bound_b@} \} \). 
<br>Write down set \(B\) by specifying its elements.</p>
<p>\(B = \) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p>

<br>

<p>(c) Select which of the following statements is true:</p>
<p>[[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked solution</h3>

<p>(a) The elements of \(A\) are those which are elements of \(\mathbb{Z} = \{ ... , -2, -1, 0, 1, 2, ... \}\) and which are in the interval \({@lower_bound_a@} \leq x \leq {@upper_bound_a@}\).
<br>Therefore, \(A = {@ta1@}\)</p>

<br>

<p>(b) The elements of \(B\) are those which are elements of \(\mathbb{N} = \{1, 2, 3, ... \}\) and which are in the interval \({@lower_bound_b@} \leq x \leq {@upper_bound_b@}\).
<br>Therefore, \(B = {@ta2@}\)</p>

<br>

<p>(c) We should check to see whether the elements of \(A\) and \(B\) are contained in the other set respectively.
<br>{@if tac=2 then 
sconcat("We can see that every element of the set \\(B\\) is also an element of the set \\(A\\). 
<br>However, not every element of the set \\(A\\) is an element of \\(B\\), for example ", listify(setdifference(ta1,ta2))[1], ".")
else 
sconcat("We can see that not every element of the set \\(B\\) is an element of the set \\(A\\), for example ",  listify(setdifference(ta2,ta1))[1], ".
<br>Also, not every element of the set \\(A\\) is an element of \\(B\\), for example ", listify(setdifference(ta1,ta2))[1], ".")@}
</p>]]></text>
    </generalfeedback>
    <defaultgrade>3</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Setting parameters for part (a) */

upper_bound_a: rand(5)+5;
lower_bound_a: -rand(4)-2;

talist1:makelist(i,i,lower_bound_a,upper_bound_a);
ta1:setify(talist1);

/* Setting parameters for part (b) */

upper_bound_b: rand(5)+5;
lower_bound_b: rand_with_prohib(-6,-1,[lower_bound_a]);

talist2:makelist(i,i,1,upper_bound_b);
ta2:setify(talist2);

/* Determine the correct statement of (c) */

choices: [
  [1, false, "\\(A\\) is a proper subset of \\(B\\)"], 
  [2, is(upper_bound_a >= upper_bound_b), "\\(B\\) is a proper subset of \\(A\\)"], 
  [3, false, "\\(A\\) is equal to \\(B\\)"], 
  [4, is(upper_bound_a < upper_bound_b), "\\(A\\) is not a subset of \\(B\\) and \\(B\\) is not a subset of \\(A\\)"]
];

ta3:random_permutation(choices);
tac: if upper_bound_a >= upper_bound_b then 2 else 4;

/* Answer-independent mistakes */

/* (a) */

/* N2: Correct answer given with with strings not numbers */
talistword1:makelist(string(i),i,lower_bound_a,upper_bound_a);
taword1:setify(talistword1);

/* N3: Treats bounds as strict */
part_a_n3_1: setify(makelist(i,i,lower_bound_a + 1,upper_bound_a));
part_a_n3_2: setify(makelist(i,i,lower_bound_a,upper_bound_a - 1));
part_a_n3_3: setify(makelist(i,i,lower_bound_a + 1,upper_bound_a - 1));

/* (b) */
/* N2: Correct answer given with with strings not numbers */
talistword2:makelist(string(i),i,1,upper_bound_b);
taword2:setify(talistword2);

/* N3: Includes non-positive elements (but otherwise correct) */
part_b_n3_1: setify(makelist(i,i,0,upper_bound_b));

/* N4: Treats upper bound as strict */
part_b_n4_1: setify(makelist(i,i,1,upper_bound_b-1));

/* N4-F: Otherwise incorrect */

/* (c) */

/* N2: Error carried forward */
part_c_ecf_ans1: {0};
part_c_ecf_ans2: {0};
part_c_ecf_ans3: 3;

/* N3: Gives A is a proper subset of B */
part_c_n3_1: 1;

/* N4: Gives A is equal to B */
part_c_n4_1: 3;

/* N5: Gives incorrect of two potentially correct answers */
part_c_n5_1: if tac=2 then 4 else 2;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\(A = {@ta1@}\)</p>
<p>\(B = {@ta2@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>25</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>{element 1, element 2, ... }</syntaxhint>
      <syntaxattribute>1</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>25</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>{element 1 , element 2 , ... }</syntaxhint>
      <syntaxattribute>1</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>radio</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/* Node 3: the students treats the inequalities as strict */
node3_test: if (ans1=part_a_n3_1 or ans1=part_a_n3_2 or ans1=part_a_n3_3) then 1 else 0;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>Sets</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>Sets</answertest>
        <sans>ans1</sans>
        <tans>taword1</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer is correct but as the elements of the set are numbers you do not need to contain them is quotation marks, "". 
<br>"" is used to enter words in an answer.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>node3_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Be careful that the interval in the definition of \(A\) has <b>weak</b> inequalities, so \(A\) should include the endpoints of the interval {@lower_bound_a@} and {@upper_bound_a@}.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>You have not correctly determined the set \(A\).
<br>The elements of \(A\) are those which are elements of \(\mathbb{Z} = \{ ... , -2, -1, 0, 1, 2, ... \}\) and which are in the interval \({@lower_bound_a@} \leq x \leq {@upper_bound_a@}\).
<br>Therefore, \(A = {@ta1@}\)</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/* Node 3: Check whether the answer includes non-positive elements but is otherwise correct */
non_pos_int: setify(makelist(i,i,lower_bound_b,0))
node3_test: if cardinality(intersection(ans2,non_pos_int)) # 0 and ta2=setdifference(ans2,non_pos_int) then 1 else 0;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>Sets</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>Sets</answertest>
        <sans>ans2</sans>
        <tans>taword2</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer is correct but as the elements of the set are numbers you do not need the "" symbol. "" is used to enter words in an answer.<br></p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>node3_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have correctly evaluated the interval which defines \(B\), however you have not considered that \(B\) must only contain elements in the set \(\mathbb{N} = \{1, 2, 3, ...\}\)</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt2-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>part_b_n4_1</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Be careful that the interval in the definition of \(B\) has <b>weak</b> inequalities, so \(B\) should include the endpoint of the interval {@upper_bound_b@}.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-4-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>You have not correctly determined the set \(B\).
<br>The elements of \(B\) are those which are elements of \(\mathbb{N} = \{1, 2, 3, ... \}\) and which are in the interval \({@lower_bound_b@} \leq x \leq {@upper_bound_b@}\).
<br>Therefore, \(B = {@ta2@}\)</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/* Node 2: Error carried forwards */
ecf: if (subsetp(ans1,ans2) and subsetp(ans2,ans1)) then 3 
  else if subsetp(ans1,ans2) then 1
  else if subsetp(ans2,ans1) then 2
  else 4;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>tac</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ecf</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>From the sets you have given in parts (a) and (b), you have correctly identified that 
{@if ecf=1 then "\\(A\\) is a proper subset of \\(B\\)" 
  else if ecf=2 then "\\(B\\) is a proper subset of \\(A\\)"
  else if ecf=3 then "\\(A\\) is equal to \\(B\\)"
  else if ecf=4 then "\\(A\\) is not a subset of \\(B\\) and \\(B\\) is not a subset of \\(A\\)"@}.
<br>However, this is not the case for the correctly stated sets \(A\) and \(B\).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>part_c_n3_1</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>It is not true that \(A\) is a proper subset of \(B\).
<br>For example, the element {@listify(setdifference(ta1,ta2))[1]@} belongs to \(A\) but not to \(B\).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt3-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>part_c_n4_1</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>It is not true that \(A\) is equal to \(B\).
<br>For example, the element {@listify(setdifference(ta1,ta2))[1]@} belongs to \(A\) but not to \(B\).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt3-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>part_c_n5_1</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>It is not true that {@if ans3=2 then "\\(B\\) is a proper subset of \\(A\\)" else "\\(A\\) is not a subset of \\(B\\) and \\(B\\) is not a subset of \\(A\\)"@}.
<br>{@if ans3=2 then sconcat("For example, the element ", listify(setdifference(ta2,ta1))[1],  " belongs to \\(B\\) but not to \\(A\\)") else "Every element of \\(B\\) is contained in \\(A\\) so in fact, \\(B\\) is a subset of\\(A\\)"@}.]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-5-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2127909311</deployedseed>
    <deployedseed>449025244</deployedseed>
    <deployedseed>603987329</deployedseed>
    <deployedseed>1838425399</deployedseed>
    <deployedseed>446984294</deployedseed>
    <deployedseed>1217932310</deployedseed>
    <deployedseed>1316426087</deployedseed>
    <deployedseed>2005074002</deployedseed>
    <deployedseed>613262710</deployedseed>
    <deployedseed>325362992</deployedseed>
    <deployedseed>1204028178</deployedseed>
    <deployedseed>1868525495</deployedseed>
    <deployedseed>1042535862</deployedseed>
    <deployedseed>1987063754</deployedseed>
    <deployedseed>2079459088</deployedseed>
    <deployedseed>1930880556</deployedseed>
    <deployedseed>540112801</deployedseed>
    <deployedseed>758136104</deployedseed>
    <deployedseed>2084893277</deployedseed>
    <deployedseed>316309478</deployedseed>
    <deployedseed>579861685</deployedseed>
    <deployedseed>2111529325</deployedseed>
    <deployedseed>1949077809</deployedseed>
    <deployedseed>314010044</deployedseed>
    <deployedseed>1589711653</deployedseed>
    <deployedseed>1881505389</deployedseed>
    <deployedseed>1674005578</deployedseed>
    <deployedseed>684959999</deployedseed>
    <deployedseed>1690901667</deployedseed>
    <deployedseed>1969842543</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>first(mcq_correct(ta3))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>taword1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value/>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value/>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>part_a_n3_1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value/>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value/>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-3-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>{0}</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value/>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value/>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-3-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>5</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value/>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>taword2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value/>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>6</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value/>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>part_b_n3_1</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value/>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt2-3-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>7</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value/>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>part_b_n4_1</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value/>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt2-4-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>8</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>part_c_ecf_ans1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>part_c_ecf_ans2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>part_c_ecf_ans3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-3-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt2-4-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt3-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>9</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>part_c_n3_1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt3-3-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>10</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>part_c_n4_1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt3-4-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>11</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>part_c_n5_1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt3-5-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
